# Assignment 01

### Exercise 1: Functions

Implement the function ``euclidean_distance_2d``, which computes the Euclidean distance between two points ``(x_1, y_1)`` and ``(x_2, y_2)`` in 2-D Euclidean space.
The function takes the coordinates ``x_1, y_1, x_2, y_2`` as four parameters with numeric values (``float``).

Then, compute the distance between the already defined 2-D points ``point_1`` and ``point_2`` using the previously implemented function. 

*Hint*: The Euclidean distance ``d`` can be computed using this formula: $d = \sqrt{|x_1-x_2|^2 + |y_1-y_2|^2}$

<!-- grade_id: cell-68371f066db3bc37, points: 0, grade: false, solution: true -->
```python
import math

def euclidean_distance_2d(x_1, y_1, x_2, y_2):
    ### BEGIN SOLUTION
    distance = math.sqrt((x_2 - x_1)**2 + (y_2 - y_1)**2)
    ### END SOLUTION
    return distance

point_1 = [0.23, -1.98]
point_2 = [-10.2548876, 4]

dist = None
# Compute Euclidean distance between point_1 and point_2 and save as dist.
### BEGIN SOLUTION
dist = euclidean_distance_2d(point_1[0], point_1[1], point_2[0], point_2[1])
### END SOLUTION
```

<!-- grade_id: cell-1083788a70e7a1fb, points: 1, grade: true, solution: false -->
```python
assert dist
assert euclidean_distance_2d(0, 0, 0, 0) == 0
assert math.isclose(euclidean_distance_2d(0.123456, -0.987654, -0.123456, 0.987654), 1.9906800924829686, rel_tol=1e-5)
assert math.isclose(dist, 12.070346638959206, rel_tol=1e-5)
print("TESTS PASSED!")
```

### Exercise 2: Dictionaries

#### 2 a)

There already exists a list of keys `keys_list`, and a list of values `values_list`.

Fill the empty dictionary `my_dict` such that the i-th entry of the list `keys_list` is a key of the dictionary with the value assigned to it corresponding to the i-th entry of the list `values_list`.

There may not be any other keys or values contained in `my_dict`.

Use the given lists to fill the dictionary and do not retype the contained values!

<!-- grade_id: cell-e3ea4658fa9266cd, points: 0, grade: false, solution: true -->
```python
my_dict = {}

keys_list = [1, 2, 3, 4, 5]
values_list = ["monday", "tuesday", "wednesday", "thursday", "friday"]

### BEGIN SOLUTION
for key, value in zip(keys_list, values_list):
    my_dict[key] = value
### END SOLUTION
```

<!-- grade_id: cell-26cb72fbcbf64dd0, points: 1, grade: true, solution: false -->
```python
assert my_dict[1] == "monday"
assert my_dict[3] == "wednesday"
assert len(my_dict) == 5
assert 0 not in my_dict
assert 6 not in my_dict
print("TESTS PASSED!")
```

#### 2 b)

Given is a dictionary `measurements_dict` with values corresponding to measured temperature values in °C.

Your task is to convert the temperature values within the dictionary to temperatures in °Fahrenheit.
The corresponding keys must be preserved.
(You are supposed to change the measured values in-place within the existing dictionary.)

**Attention**: The keys and values you are working with must always be taken from the existing dictionary and may not be re-entered manually!

*Hint*: Given a temperature value $c$ in ° Celsius, it can be converted to a temperature value $f$ in ° Fahrenheit using the following formula:
<p style="text-align: center;">
$f = \frac{9}{5} \cdot c + 32$
</p>

<!-- grade_id: cell-0ed48280a015f3e3, points: 0, grade: false, solution: true -->
```python
measurements_dict = {
    "1": 23.5,
    "2": 23.9,
    "3": 24.7,
    "4": 25.1,
    "5": 25.0,
    "6": 24.9,
    "7": 22.3,
    "8": 19.6,
    "9": 12.5,
    "10": 5.0
}

### BEGIN SOLUTION
for key in measurements_dict:
    measurements_dict[key] = measurements_dict[key] * (9/5) + 32
### END SOLUTION
```

<!-- grade_id: cell-b16c3535e364ff1c, points: 1, grade: true, solution: false -->
```python
assert len(measurements_dict) == 10
assert sum(measurements_dict.values()) == 691.7
print("TESTS PASSED!")
```

#### 2 c)

Implement the function `add_entry_for_key`, which takes the parameters `my_dict`, `new_entry` and a `key`.
The parameter `my_dict` can be assumed to be a dictionary containing only lists as values.
The function should add the value `new_entry` to the list stored under the key `key`.
If `key` is not yet within the keys of the dictionary, a new entry with key `key` and the list only containing `new_entry` as value should be added to `my_dict`.

Then, use the function `add_entry_for_key` to add to the dictionary according to the table below:
* Add Tom's grade to the dictionary `exam_results`.
* Add the list of the student's seat numbers to the dictionary `exam_results` under the key `seat_number`.

| Name | Grade | Seat Number |
| ---- | ---- | --------- |
| Lisa | 2,0 | 11 |
| Lena | 2,7 | 21 |
| Luisa | 1,3 | 31 |
| Tom | 1,7 | 41 |

*Hint*: Consider the order of the lists!

<!-- grade_id: cell-5ba830463fba67c0, points: 0, grade: false, solution: true -->
```python
def add_entry_for_key(my_dict, key, new_entry):
    ### BEGIN SOLUTION
    if key not in my_dict:
        my_dict[key] = []
    my_dict[key].append(new_entry)
    ### END SOLUTION
    
exam_results = {
    "students": ["Lisa", "Lena", "Luisa"],
    "grades": [2.0, 2.7, 1.3]
}

add_entry_for_key(exam_results, "students", "Tom")

### BEGIN SOLUTION
add_entry_for_key(exam_results, "grades", 1.7)
add_entry_for_key(exam_results, "seat_number", 11)
add_entry_for_key(exam_results, "seat_number", 21)
add_entry_for_key(exam_results, "seat_number", 31)
add_entry_for_key(exam_results, "seat_number", 41)
### END SOLUTION
```

<!-- grade_id: cell-4f94ae01b36b616e, points: 1, grade: true, solution: false -->
```python
assert "students" in exam_results
assert "grades" in exam_results
assert "seat_number" in exam_results
assert exam_results["seat_number"][3] == 41
assert len(exam_results) == 3
for key, value in exam_results.items():
    assert len(value) == 4
print("TESTS PASSED!")
```

### Exercise 3: Conditions and Loops

You are given the following three dictionaries:

* The dictionary `inventory` contains, under the respective part number as the key, the quantity of parts available in inventory.
* The dictionary `requirements` contains a list of parts needed to manufacture a product (part number is the key; the value indicates the quantity of that part type required).
* The dictionary `price_list` contains, under a part number as the key, the price for one unit of that part type.

Implement a code block that calculates the cost of purchasing the missing parts required to produce a product according to `requirements`.
Determine which parts, and in what quantity, need to be procured so that the total (inventory plus new purchases) is sufficient for production.
Then calculate the price of the parts that need to be newly procured and store this in the variable `missing_parts_costs`.

**Important**: Be sure to retrieve all values from the existing dictionaries. You must not hardcode any values (str or numeric) in your solution code.

<!-- grade_id: cell-cb23657f6d285409, points: 0, grade: false, solution: true -->
```python
inventory = {
    "12343": 2,
    "12344": 7,
    "12347": 4,
    "12348": 1,
    "12349": 12,
    "12350": 4,
    "12351": 8,
    "12352": 2,
    "12355": 2,
    "12357": 4,
    "12358": 6
}

requirements = {
    "12344": 4,
    "12346": 4,
    "12347": 4,
    "12348": 4,
    "12349": 4,
    "12352": 6,
    "12353": 6,
    "12354": 6,
    "12358": 6
}

price_list = {
    "12343": 2.99,
    "12344": 2.99,
    "12345": 3.87,
    "12346": 3.87,
    "12347": 1.72,
    "12348": 1.79,
    "12349": 4.14,
    "12350": 4.36,
    "12351": 5.12,
    "12352": 2.19,
    "12353": 2.88,
    "12354": 6.48,
    "12355": 2.78,
    "12356": 5.07,
    "12357": 4.88,
    "12358": 3.49
}

missing_parts_costs = 0

### BEGIN SOLUTION
for part_no, required_amount in requirements.items():
    missing_amount = required_amount
    if part_no in inventory:
        missing_amount = max(required_amount - inventory[part_no], 0)
    part_cost = price_list[part_no]
    missing_parts_costs = missing_parts_costs + part_cost * missing_amount
### END SOLUTION
    
print("Summed costs for missing parts: {}".format(missing_parts_costs))
```

<!-- grade_id: cell-eb42148c379a5022, points: 1, grade: true, solution: false -->
```python
import math
assert math.isclose(missing_parts_costs, 85.77)
print("TESTS PASSED!")
```