<img align="right" width=160 src="https://upload.wikimedia.org/wikipedia/commons/thumb/3/3a/Logo_KIT.svg/1920px-Logo_KIT.svg.png"/>

# IDS Assignments

**Allgemeine Hinweise**:

* Im Rahmen der Tutorien haben Sie 60 min Zeit für die Bearbeitung der Aufgaben. Bitte wenden Sie sich umgehend an Ihren Tutor, sobald die Bearbeitung abgeschlossen ist.
* Bitte überprüfen Sie, dass die enthaltenen Tests fehlerfrei ausgeführt werden können.
* Bei Unklarheiten, Fragen oder Problemen, wenden Sie sich bitte **unmittelbar** an Ihren Tutor (nicht erst zum Ende des Tutoriums) und holen Sie sich Unterstützung ein.
* Das Arbeiten in 2-er Teams ist erlaubt und erwünscht! :)


**General Notes**:

* During the tutorials, you have 60 minutes to work on the tasks. Please contact your tutor immediately once you have completed the work.
* Make sure that the included tests can run without errors.
* If anything is unclear, or if you have questions or problems, please contact your tutor immediately (not just at the end of the tutorial) and ask for support.
* Working in pairs (2-person teams) is allowed and encouraged! :)

## Assignment 02

### Aufgabe 1: Arbeiten mit Dateien und Daten

#### 1 A)

Die Datei `temperatures_daily.csv` im Unterordner `data` enthält die in Rheinstetten aufgezeichneten täglichen Tageshöchsttemperaturen des Zeitraums 02.08.2023 - 01.02.2025.

Dabei hat jede Zeile das Format `<date>;<temperature>`.

Lesen Sie die Datei ein und speichern Sie alle enthaltenen Daten als Datums-Temperatur-Tupel (`(date, temperature)`) in der Liste `single_temps`.
Der zweite Eintrag des Tupels `temperature` soll dabei den Datentyp `float` haben.

Das Öffnen der Datei, sowie die Iteration über die einzelnen Zeilen ist bereits implementiert.
Um die Aufgabe zu lösen, müssen Sie also lediglich Datum und Temperatur aus den Zeilen-Strings extrahieren, ggf. den Datentyp umwandeln und in das Ziel-Array `single_temps` schreiben.

*Hinweis*: Es ist empfehlenswert die Datei zunächst z.B. im Texteditor zu öffnen und manuell anzuschauen. Beachten Sie insbesondere, dass die Datei eine Kopfzeile enthält.

*Hinweis*: Die Funktion `datetime.date.fromisoformat(<str>)` kann verwendet werden, um den Datums-String in ein Objekt der Klasse `datetime.date` umzuwandeln. Alternativ kann auch mit den Zeichenketten weitergearbeitet werden.

**Achtung**: Die Datei ist nicht fehlerfrei. Beispielsweise können einzelne Werte fehlen oder fehlerhaft sein. Eine Formatierung der Werte zu `float` ist aber für alle Einträge möglich.

<!-- grade_id: cell-695438bf79cc75c0, points: 0, grade: false, solution: true -->
```python
import datetime

single_temps = []

input_file = "./data/temperatures_daily.csv"

# Open temperature file
with open(input_file, "r") as file_opened:
    # Iterate over lines (Str type) in the file
    for i, line in enumerate(file_opened):
        # YOUR CODE HERE
        
        # END OF YOUR CODE
    # Close file
    file_opened.close()
```

<!-- grade_id: cell-45aa9d5504a453b4, points: 1, grade: true, solution: false -->
```python
assert len(single_temps) == 547
for date, temp in single_temps:
    assert type(temp) == float
assert single_temps[0][1] == 18.2
assert single_temps[-1][1] == 1.4
```

#### 1 B)

Vervollständige die folgende Implementierung, welche für alle enthaltenen Monate (08/2023 - 02/2025) die durchschnittliche Tageshöchsttemperatur berechnen soll.

Ziel ist die Befüllung des Dictionary `average_high`, sodass der Schlüssel (Key) dem jeweiligen Monat (Format `mm/YYYY`, Datentyp `str`) und der Wert der berechneten Durchschnittstemperatur (`float`) der Werte des Monats entspricht.

Um die durchschnittlichen Höchstwerte monatsweise zu berechnen, kann es hilfreich sein, zunächst alle Temperaturwerte nach Monaten zu gruppieren.
Daher befüllen wir zunächst ein Dictionary `temps_by_month`, welches unter dem Monatsschlüssel (`mm/YYYY`) eine Liste aller in diesem Monat gemessenen Werte enthalten soll.

**Achtung**: Die Datei ist nicht fehlerfrei. Beispielsweise können einzelne Werte fehlen oder fehlerhaft sein. Gehen Sie damit wiefolgt um:
* Aufgrund fehlender Werte ist bei der Durchschnittsberechnung die Anzahl tatsächlich vorhandener Werte zu berücksichtigen.
* Fehlerhafte Werte sind schwieriger zu identifizieren und können zunächst bei der Berechnung ignoriert werden.

<!-- grade_id: cell-479353e3a94b8137, points: 0, grade: false, solution: true -->
```python
temps_by_month = {}
average_high = {}

for date, temp in single_temps:
    # Build month-Str as instructed (mm/YYYY)
    # YOUR CODE HERE
    
    # END OF YOUR CODE
    if month not in temps_by_month:
        temps_by_month[month] = []
    temps_by_month[month].append(temp)

# Iterate over temps_by_month, compute the monthly average high for each month and write it into the Dict average_high
# YOUR CODE HERE

# END OF YOUR CODE
```

<!-- grade_id: cell-2315c38b77b04863, points: 1, grade: true, solution: false -->
```python
import math

assert len(average_high) == 19
for key, value in average_high.items():
    assert type(key) is str
    assert type(value) is float
assert "10/2023" in average_high
assert math.isclose(average_high["11/2023"], 7.08, rel_tol=1e-5)
assert math.isclose(average_high["05/2024"], 18.2903, rel_tol=1e-5)
assert math.isclose(average_high["02/2025"], 1.4, rel_tol=1e-5)
```

#### 1 C) * BONUS *

*\* Diese Aufgabe ist optional. Wenn sie erfolgreich gelöst wird kann sie eine beliebige andere Teilaufgabe dieses Aufgabenblatts ersetzen.*

Schreibe nun die zuvor berechneten Monatsdurchschnittstemperaturen in eine Datei mit Namen ``temperatures_monthly_averages.csv`` in den Unterordner ``data``.

Der Inhalt der Datei soll dem csv-Format mit Trennzeichen `;` entsprechen und die beiden Spalten "Month" und "Average Temperature" beinhalten.

* Es muss also zunächst eine Kopfzeile mit den beiden Spaltennamen (Trennzeichen beachten!) geschrieben werden
* Danach folgt eine Zeile je Monat im korrekten Format
* Orientiere Dich sich im Zweifel am Format der in Aufgabenteil A) eingelesenen Datei


<!-- grade_id: cell-f7f05da090a0da2f, points: 0, grade: false, solution: true -->
```python
output_file = "./data/temperatures_monthly_averages.csv"

with open(output_file, "w") as file_opened:
    # YOUR CODE HERE
    
    # END OF YOUR CODE
    file_opened.close()
```

<!-- grade_id: cell-9e442d447aa30579, points: 1, grade: true, solution: false -->
```python
import os
assert os.path.isfile(output_file)
with open(output_file, "r") as file_reader:
    content = file_reader.read()
    file_reader.close()
lines = content.split("\n")
assert "Month" in lines[0]
assert "Average Temperature" in lines[0]

### Check value consistency
input_temp_sum = sum([entry for entry in average_high.values()])

temp_sum = 0
for line in lines[1:]:
    if len(line) > 0:
        temp_sum += float(line.rsplit(";")[1])
assert math.isclose(temp_sum, input_temp_sum, rel_tol=1e-3)
```

### Aufgabe 2: Objekte

#### 2 A)

Betrachte folgende, noch sehr einfache Implementierung einer Klasse `Car`.

Vervollständige die Magic Method `__eq__` so, dass sie genau dann `True` zurück gibt, wenn die beiden verglichenen Instanzen der Klasse nicht unterscheidbar sind, d.h. die Werte aller Attribute identisch sind.

Anschließend sollen drei Instanzen der Klasse `Car` wiefolgt erzeugt werden:
* `car_1` und `car_2` sind nicht "identisch" (Vergleich mit `__eq__` ergibt `False`) 
* `car_1` und `car_3` sind "identisch" (Vergleich mit `__eq__` ergibt `True`) 

<!-- grade_id: cell-7d02de331f43302f, points: 0, grade: false, solution: true -->
```python
class Car:
    
    def __init__(self, color):
        self.color = color
        self.mileage = 0
        
    def __eq__(self, other):
        # YOUR CODE HERE
        
        # END OF YOUR CODE
    
    def drive(self, miles):
        self.mileage += miles
    
car_1 = None
car_2 = None
car_3 = None

# YOUR CODE HERE

# END OF YOUR CODE
```

<!-- grade_id: cell-4724e0fd3c4b4ff4, points: 1, grade: true, solution: false -->
```python
assert type(car_1) == Car
assert type(car_2) == Car
assert type(car_3) == Car
assert car_1 == car_3
assert car_1 != car_2
```

#### 2 B)

In dieser Aufgabe soll der Klasse `Car` Logik zum Treibstoffverbrauch hinzugefügt werden.

In jeder Instanz ist der Treibstoffverbrauch des Modells pro 100 km gefahrener Strecke (`fuel_consumption_liter_per_100km`) hinterlegt.
Außerdem erhält jede Instanz neben einem Gesamtkilometerzähler (`mileage`) ein Attribut, welches den aktuellen Treibstofffüllstand (`fuel_liters`) enthält.

Es sollen nun die Methoden `drive()`, `fuel_up()` und `get_remaining_distance()` implementiert werden.

Bei Aufruf von `drive` wird eine entsprechende Fahrtstrecke zurückgelegt, d.h. der Kilometerzähler muss erhöht werden und der aktuelle Treibstofffüllstand entsprechend des Verbrauchs aktualisiert werden.

Ein Aufruf von `fuel_up` bedeutet ein Nachfüllen von Treibstoff, hier soll der aktuelle Füllstand entsprechend angepasst werden.

Die Methode `get_remaining_distance` gibt die Strecke in Kilometern zurück, die mit dem aktuellen Füllstand bewältigt werden können.

*Hinweis*: Hier wird ein konstanter Verbrauch des Fahrzeugs, welcher lediglich von der zurückzulegenden Strecke abhängt, angenommen.

<!-- grade_id: cell-6f85adabf0914287, points: 0, grade: false, solution: true -->
```python
class Car:
    
    def __init__(self, color, fuel_consumption):
        self.color = color
        self.fuel_consumption_liter_per_100km = fuel_consumption
        self.fuel_liters = 0
        self.mileage = 0
    
    def drive(self, distance_km):
        # YOUR CODE HERE
        
        # END OF YOUR CODE
        
    def fuel_up(self, liters_added):
        # YOUR CODE HERE
        
        # END OF YOUR CODE
        
    def get_remaining_distance(self):
        # YOUR CODE HERE
        
        # END OF YOUR CODE
```

<!-- grade_id: cell-728105f777ae311f, points: 1, grade: true, solution: false -->
```python
my_car = Car("red", 5.2)
my_car.fuel_up(40)
my_car.drive(120)
assert my_car.mileage == 120
assert my_car.fuel_liters == 33.76
assert math.isclose(my_car.get_remaining_distance(), 649.2307, abs_tol=0.01)
```