<img align="right" width=160 src="https://upload.wikimedia.org/wikipedia/commons/thumb/3/3a/Logo_KIT.svg/1920px-Logo_KIT.svg.png"/>

# IDS Assignments

**Allgemeine Hinweise**:

* Im Rahmen der Tutorien haben Sie 60 min Zeit für die Bearbeitung der Aufgaben. Bitte wenden Sie sich umgehend an Ihren Tutor, sobald die Bearbeitung abgeschlossen ist.
* Bitte überprüfen Sie, dass die enthaltenen Tests fehlerfrei ausgeführt werden können.
* Bei Unklarheiten, Fragen oder Problemen, wenden Sie sich bitte **unmittelbar** an Ihren Tutor (nicht erst zum Ende des Tutoriums) und holen Sie sich Unterstützung ein.
* Das Arbeiten in 2-er Teams ist erlaubt und erwünscht! :)


**General Notes**:

* During the tutorials, you have 60 minutes to work on the tasks. Please contact your tutor immediately once you have completed the work.
* Make sure that the included tests can run without errors.
* If anything is unclear, or if you have questions or problems, please contact your tutor immediately (not just at the end of the tutorial) and ask for support.
* Working in pairs (2-person teams) is allowed and encouraged! :)

## Assignment 03: Object-Oriented Programming II


General Hint: Use the `print()`-function statements to display variables and debug your code. 

### Übung 1: Klassen und Vererbung

#### 1 A)

In the following, we look at the class `Smart_Device`, which abstracts an intelligent smart home device.

The class has a boolean attribute  `on`, which indicates whether the device is currently switched on or off. 
There is also a `max_energy_consumption_kwh` attribute, which specifies the maximum power consumption of the device in kilowatts.

What you need to do:
-  Implement the method `turn_on_off` in the class, which flips the value boolean attribute `on` from `False` to `True` and the other way around.
-  Create an instance of `Smart_Device`, whose name is `my_device`, and `max_energy_consumption_kW` is 30. 
-  Then use the method `turn_on_off` to turn it on. Finally, print out the value of the `on` attribute. 


<!-- grade_id: cell-fdaca7585491de91, points: 0, grade: false, solution: true -->
```python
class Smart_Device:
    
    def __init__(self, max_energy_consumption_kW):
        self.on = False
        self.max_energy_consumption_kW = max_energy_consumption_kW
        
    # YOUR CODE HERE
    
    # END OF YOUR CODE

my_device = None

# YOUR CODE HERE

# END OF YOUR CODE
```

<!-- grade_id: cell-8e474901b3e4855c, points: 0, grade: true, solution: false -->
```python
assert my_device
assert hasattr(Smart_Device, "turn_on_off")
assert callable(Smart_Device.turn_on_off)
assert my_device.on
print("TEST PASSED!")
```

#### 1 B)

Now in this part, we will implement a class called `Smart_Light` that inherits from `Smart_Device`

What you need to do:
-  Implement the class `Smart_Light`, which inherits from the class `Smart_Device`
-  Implement a constructor that takes the same arguments as the superclass, in addition to two new attributes, namely `brightness` and `color`
-  Set the default `brightness` to 70 and default `color` to `white`

Hint: Use the `super()` function to initialize the variables of the superclass. 


<!-- grade_id: cell-35d9408f080d9e80, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

# END OF YOUR CODE
```

<!-- grade_id: cell-5831f0378532c42f, points: 0, grade: true, solution: false -->
```python
import inspect 

assert Smart_Light
assert issubclass(Smart_Light, Smart_Device)

signature = inspect.signature(Smart_Light.__dict__["__init__"])
params = signature.parameters.values()

param_names = [param.name for param in params]
assert "self" in param_names
assert "max_energy_consumption_kW" in param_names
assert "brightness" in param_names
assert "color" in param_names

assert signature.parameters["brightness"].default == 70
assert signature.parameters["color"].default == "white"

print("TEST PASSED!")
```

#### 1 C)

Now we continue the implementation of the class `Smart_Light` from part B. Copy your solution into the next cell. 
We want to implement instance methods for setting the brightness and color.

What you need to do:
- Copy your code from part B into the cell below.
- Add the following new methods to the same class `Smart_Light`:
  - Method `set_color`: This method sets the light color by reassigning the `color` attribute. 
  - Method `set_brightness`: This method sets the brightness by reassigning the `brightness` attribute. Make sure that the brightness is between 0 and 100.
- Finally, create an instance of my `Smart_Light` named `my_light`, with `max_energy_consumption_kW` set to 0.03, `brightness` set to 75, and color set to "yellow"

<!-- grade_id: cell-3b8e1af625b8b02f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

# END OF YOUR CODE
```

<!-- grade_id: cell-e046b90f84bba01d, points: 0, grade: true, solution: false -->
```python
assert my_light
assert isinstance(my_light, Smart_Device)

assert hasattr(my_light, "max_energy_consumption_kW")
assert hasattr(my_light, "brightness")
assert hasattr(my_light, "color")

assert hasattr(my_light, "set_color")
assert hasattr(my_light, "set_brightness")

assert my_light.max_energy_consumption_kW == 0.03
assert my_light.color == "yellow"
assert my_light.brightness == 75

import inspect 

signature = inspect.signature(Smart_Light.__dict__["set_color"])
params = signature.parameters.values()
param_names = [param.name for param in params]
assert "self" in param_names
assert "color" in param_names

signature = inspect.signature(Smart_Light.__dict__["set_brightness"])
params = signature.parameters.values()
param_names = [param.name for param in params]
assert "self" in param_names
assert "brightness" in param_names

print("TEST PASSED!")
```

### Exercise 2: Magic Methods

#### 2 A)
Given is the class `Robot`, which abstracts a robot. The class has the attributes `name`, `battery_capacity`, and `task`

Implement the methods `__str__` and `__repr__` as follows:
- The magic method ``__str__`` shoudl return a string as `my name is <name>, my task is <task> and I can do work for <battery_capacity> hours`
- The magic method ``__repr__`` should return a string that has `name`, `task`, and `battery_capacity`. (e.g., `"spot,cleaning,50"`)

Now create an instance named `my_robot` whose name is `Jack` and his task is `teaching`, and has a battery capacity of `12` hours. Finally, convert this class into a string and save it in a new variable called `str_jack`. 

<!-- grade_id: cell-01a92def60cc0550, points: 0, grade: false, solution: true -->
```python
class Robot:
    
    def __init__(self, name, battery_capacity=50, task="cleaning"):
        self.name = name
        self.battery_capacity = battery_capacity
        self.task = task
        
    def __str__(self):
        # YOUR CODE HERE
        
        # END OF YOUR CODE

    def __repr__(self):
        # YOUR CODE HERE
        
        # END OF YOUR CODE

my_robot = None
str_jack = None
# YOUR CODE HERE

# END OF YOUR CODE

```

<!-- grade_id: cell-e10d9c426270062a, points: 0, grade: true, solution: false -->
```python
assert my_robot
assert str_jack

assert isinstance(my_robot, Robot)
assert str_jack == "my name is Jack, my task is teaching and I can do work for 12 hours"

assert (my_robot.__repr__() == "Jack,teaching,12") or (my_robot.__repr__() == "Jack, teaching, 12")
print("TEST PASSED!")
```

#### 2 B)
Given is the class `Grade`, which abstracts the grade of a student. The possible grades are `A+`, `B+`, `C+`, `D`, with `A+` as the best grade and `D` as the worst grade. 

We want to implement the methods `__lt__`, `__gt__`  and `__eq__` in order to compare grades to each other. For Example the if we are comparing `A+` and `B+`:
- `A+` > `B+` must return `True`
- `A+ ` < `B+` must return `False`
- `A+` == `B+` must return `False`

Hints:
- You may initialize additional attributes in the constructor. However, the arguments must remain the same.
- Maybe dictionaries can help you?

<!-- grade_id: cell-299917d4b5e5f00c, points: 0, grade: false, solution: true -->
```python
class Grade():
    possible_grades = ["A+", "B+", "C+", "D"]
    def __init__(self, grade):
        if grade in self.possible_grades:
            self.grade = grade
        else:
            raise ValueError(f"grade must be on in {self.possible_grades}")
        # YOUR CODE HERE
        
        # END OF YOUR CODE

    # YOUR CODE HERE
    
    # END OF YOUR CODE
```

<!-- grade_id: cell-80c70901041cfed5, points: 1, grade: true, solution: false -->
```python
assert Grade

assert Grade("A+") == Grade("A+")
assert Grade("B+") == Grade("B+")
assert Grade("C+") == Grade("C+")
assert Grade("D") == Grade("D")

assert Grade("A+") != Grade("B+")
assert Grade("A+") != Grade("C+")
assert Grade("A+") != Grade("D")
assert Grade("B+") != Grade("C+")
assert Grade("B+") != Grade("D")
assert Grade("C+") != Grade("D")

assert Grade("A+") > Grade("B+")
assert Grade("A+") > Grade("C+")
assert Grade("A+") > Grade("D")
assert Grade("B+") > Grade("C+")
assert Grade("B+") > Grade("D")
assert Grade("C+") > Grade("D")

assert not Grade("A+") < Grade("B+") 
assert not Grade("A+") < Grade("C+")
assert not Grade("A+") < Grade("D")
assert not Grade("B+") < Grade("C+")
assert not Grade("B+") < Grade("D")
assert not Grade("C+") < Grade("D")

print("TEST PASSED!")
```