<img align="right" width=160 src="https://upload.wikimedia.org/wikipedia/commons/thumb/3/3a/Logo_KIT.svg/1920px-Logo_KIT.svg.png"/>

# IDS Assignments

**Allgemeine Hinweise**:

* Im Rahmen der Tutorien haben Sie 60 min Zeit für die Bearbeitung der Aufgaben. Bitte wenden Sie sich umgehend an Ihren Tutor, sobald die Bearbeitung abgeschlossen ist.
* Bitte überprüfen Sie, dass die enthaltenen Tests fehlerfrei ausgeführt werden können.
* Bei Unklarheiten, Fragen oder Problemen, wenden Sie sich bitte **unmittelbar** an Ihren Tutor (nicht erst zum Ende des Tutoriums) und holen Sie sich Unterstützung ein.
* Das Arbeiten in 2-er Teams ist erlaubt und erwünscht! :)


**General Notes**:

* During the tutorials, you have 60 minutes to work on the tasks. Please contact your tutor immediately once you have completed the work.
* Make sure that the included tests can run without errors.
* If anything is unclear, or if you have questions or problems, please contact your tutor immediately (not just at the end of the tutorial) and ask for support.
* Working in pairs (2-person teams) is allowed and encouraged! :)

## Assignment 04: Algorithmen

### Aufgabe 1: Umwandlung zwischen Zahlensystemen

#### 1 A) Binär nach Dezimal

Implementiere einen Algorithmus zur Umwandlung einer binären Zeichenkette in das Dezimalsystem. Implementiere die Funktion `binary_to_decimal` weiter.

Die Formel lautet:

$(b_{n-1}\ldots b_1 b_0)_2 = \left(\sum_{i=0}^{n-1} b_i 2^i\right)_{10}$

**Hinweis:** Die Ausgabe des Algorithmus ist eine ganze Zahl


<!-- grade_id: cell-42a04ad1b4d179e1, points: 0, grade: false, solution: true -->
```python
def binary_to_decimal(binary_string):
    decimal_value = 0
    # YOUR CODE HERE
    
    # END OF YOUR CODE
    return decimal_value
```

<!-- grade_id: cell-f0c46b0dfe143394, points: 1, grade: true, solution: false -->
```python
assert binary_to_decimal

assert binary_to_decimal("101") == 5
assert binary_to_decimal("10110") == 22
print("TEST PASSED!")
```

### 1 B) Das Horner-Schema

Schreiben Sie eine Funktion namens `horners_method`. Diese Funktion nimmt eine reguläre Dezimalzahl und konvertiert sie in ein anderes Zahlensystem (wie binär, oktal oder hexadezimal) unter Verwendung einer Technik, die Horner's Methode genannt wird. Die Funktion sollte Ihnen dann die umgewandelte Zahl in diesem neuen System zurückgeben.

Die Funktion benötigt zwei Informationen (Argumente):

1.  `number`: Dies ist die Dezimalzahl, die Sie umrechnen möchten.
2.  `base`: Dies ist die Basis des neuen Zahlensystems, in das Sie konvertieren möchten (z.B. 2 für binär, 8 für oktal, 16 für hexadezimal). Wenn Sie keine Basis (`base`) angeben, geht die Funktion automatisch davon aus, dass Sie in das Binärsystem (Basis 2) konvertieren wollen.

**Besonderer Hinweis für Hexadezimal:** Wenn Sie nach Hexadezimal (Basis 16) konvertieren, erhalten Sie manchmal Reste, die 10 oder mehr betragen. Für diese müssen Sie ein spezielles Dictionary namens `dec_hex_dict` verwenden, um diese Zahlen in die korrekten Hexadezimalbuchstaben umzuwandeln (z.B. ‚A‘ für 10, ‚B‘ für 11, usw.).


<!-- grade_id: cell-08c1cf4030c62ef6, points: 0, grade: false, solution: true -->
```python
dec_hex_dict = {
    '0': '0',
     '1': '1',
     '2': '2',
     '3': '3',
     '4': '4',
     '5': '5',
     '6': '6',
     '7': '7',
     '8': '8',
     '9': '9',
     '10': 'A',
     '11': 'B',
     '12': 'C',
     '13': 'D',
     '14': 'E',
     '15': 'F'
}

# YOUR CODE HERE

# END OF YOUR CODE
```

<!-- grade_id: cell-7b12b07614728f1c, points: 0, grade: true, solution: false -->
```python
assert horners_method
assert horners_method(22) == "10110"
assert horners_method(31) == "11111"
assert horners_method(22, 8) == "26"
assert horners_method(31, 8) == "37"
assert horners_method(10, 10) == "10"
assert horners_method(511, 16) == "1FF"
assert horners_method(4090, 16) == "FFA"
print("TEST PASSED!")
```

#### Aufgabe 2: Swap Sort

Der Sortieralgorithmus *Swap Sort* kann ebenfalls zum Sortieren von Listen mit vergleichbaren Elementen benutzt werden.
Dabei darf die Liste jedoch keine Duplikate enthalten.

Folgende Implementierung des Algorithmus ist gegeben:

```python
def swap_sort(list_to_sort):
    i = 0
    while i < len(list_to_sort):
        m = 0
        for other_entry in list_to_sort:
            if other_entry < list_to_sort[i]:
                m += 1
        if i == m:
            i += 1
        else:
            h = list_to_sort[m]
            list_to_sort[m] = list_to_sort[i]
            list_to_sort[i] = h
    return list_to_sort
```

Erweitere die bestehende Implementierung so, dass neben der sortierten Liste auch die Anzahl der durchgeführten Iterationen zurückgegeben wird.
Kopiere dazu zunäcsht obige Implementierung in die folgende Codezelle.

Es sollen keine Codezeilen des Algorithmus verändert werden, sondern lediglich neue Zeilen hinzugefügt werden.

*Hinweis*: Mit **einer Iteration** des Algorithmus ist dabei ein Durchlauf bestehend aus dem Zählen der kleineren Elemente der Liste und anschließendem Tauschen zweier Elemente bzw. Erhöhung des Indizes ``i`` gemeint.

Befülle die Liste ``unsorted`` anschließend so mit den Zahlen ``5,4,1,3,2``, dass die maximale Anzahl von Iterationen erreicht wird. Wie viele Iterationen sind das?

<!-- grade_id: cell-7223e0f31a61643f, points: 0, grade: false, solution: true -->
```python
def swap_sort(list_to_sort):
    # YOUR CODE HERE
    
    # END OF YOUR CODE
    return list_to_sort, iteration_counter

# YOUR CODE HERE

# END OF YOUR CODE
sorted_list, iteration_number = swap_sort(unsorted)
print(iteration_number)
```

<!-- grade_id: cell-62f3175b4df8e270, points: 1, grade: true, solution: false -->
```python
import random
unsorted = list(range(10))
assert swap_sort(unsorted)[1] == 10
random.shuffle(unsorted)
assert swap_sort(unsorted)[1] >= 10
for i in range(len(unsorted) - 1):
    assert unsorted[i] <= unsorted[i+1]

print("TEST PASSED!")
```

### Aufgabe 3: Sieb des Eratosthenes

In dieser Aufgabe soll ein Algorithmus zur Bestimmung von Primzahlen implementiert werden, das sogenannte *Sieb des Eratosthenes*.

Der Algorithmus funktioniert wie folgt:
* Zur Bestimmung aller Primzahlen bis zu einer ganzen Zahl ``n > 1``, starte mit einer Liste aller ganzen Zahlen von ``2`` bis ``n``.
* Beginnend mit ``i=2`` streiche alle Vielfachen von ``i`` aus der Liste
* Wiederhole den vorangegangenen Schritt mit dem nächst größeren Eintrag der Liste
* Ist ``i`` das größte Element der Liste, so ist der Algorithmus abgeschlossen. Die Liste enthält nun alle Primzahlen zwischen ``2`` und ``n``.

Implementiere den Algorithmus.

<!-- grade_id: cell-7b127a8ea604e223, points: 0, grade: false, solution: true -->
```python
def sieve_of_eratosthenes(n):
    prime_numbers = list(range(2, n+1))
    # YOUR CODE HERE
    
    # END OF YOUR CODE
    return prime_numbers
```

<!-- grade_id: cell-73f72fff1cd7a1cf, points: 2, grade: true, solution: false -->
```python
assert len(sieve_of_eratosthenes(10)) == 4
assert 2 in sieve_of_eratosthenes(10)
assert 3 in sieve_of_eratosthenes(10)
assert 5 in sieve_of_eratosthenes(10)
assert 7 in sieve_of_eratosthenes(10)
assert len(sieve_of_eratosthenes(29)) == 10

print("TEST PASSED!")
```