## Assignment 05: Data Science

### Exercise 1: Creating DataFrames

*Hint*: For solving this exercise, the contents of the on-site session 05 from 28/05/2025 can be helpful (see ILIAS).

Consider the following tabular data:

| Country | Population | Area |
| --- | ---: | ---: |
| Germany	| 84.552.242 | 357.022 |
| France	| 66.548.530 | 551.500 |
| China | 1.419.321.278 | 9.596.960 |
| India	| 1.450.935.791	| 3.287.263 |

The table contains the population size and area (in $km^2$) for a few countries.

**a)**

Create an instance named `country_df` of the pandas class DataFrame which contains the data of the table above.
To do so, first create a dictionary (`country_dict`) and use this dictionary for the initialization of the DataFrame.

Lastly, use the `display()` function to display the DataFrame.

<!-- grade_id: cell-01940a3cfd34615e, points: 0, grade: false, solution: true -->
```python
import pandas as pd

### BEGIN SOLUTION
country_dict = {
    "Country": ["Germany", "France", "China", "India"],
    "Population": [84552242, 66548530, 1419321278, 1450935791],
    "Area": [357022, 551500, 9596960, 3287263]
}

country_df = pd.DataFrame(country_dict)
display(country_df)
### END SOLUTION
```

<!-- grade_id: cell-933c08e46fa7fec4, points: 0, grade: true, solution: false -->
```python
assert type(country_dict) is dict
assert type(country_df) is pd.DataFrame

assert country_df.shape == (4,3)
assert set(country_df.columns) == {"Country", "Population", "Area"}
assert country_df.iloc[0]["Population"] == 84552242
assert country_df.iloc[1]["Population"] == 66548530
assert country_df.iloc[2]["Population"] == 1419321278
assert country_df.iloc[3]["Population"] == 1450935791
assert country_df.iloc[0]["Area"] == 357022
assert country_df.iloc[1]["Area"] == 551500
assert country_df.iloc[2]["Area"] == 9596960
assert country_df.iloc[3]["Area"] == 3287263
assert not country_df.isna().any(axis=None)
```

**b)**

Manually create a csv file named `country_data.csv` which contains the data of the above table in proper csv format (using the separator `,`).

Then, you should be able to read and display the data using the test cell below.

*Hint:* You do not need to write any Python code for this exercise. Manually create the csv file in the same folder where this Jupyter Notebook is located. Thus, there is no solution code cell below, but only the test cell trying to read the file to be created.

<!-- grade_id: cell-41fec27d71f0d2bf, points: 0, grade: true, solution: false -->
```python
import pandas as pd

# Read csv file and display contents
country_csv_df = pd.read_csv("country_data.csv", sep=",")
display(country_csv_df)

# Tests concerning contained data
assert type(country_csv_df) is pd.DataFrame

assert country_df.shape == (4,3)
assert set(country_df.columns) == {"Country", "Population", "Area"}
assert country_df.iloc[0]["Population"] == 84552242
assert country_df.iloc[1]["Population"] == 66548530
assert country_df.iloc[2]["Population"] == 1419321278
assert country_df.iloc[3]["Population"] == 1450935791
assert country_df.iloc[0]["Area"] == 357022
assert country_df.iloc[1]["Area"] == 551500
assert country_df.iloc[2]["Area"] == 9596960
assert country_df.iloc[3]["Area"] == 3287263
assert not country_df.isna().any(axis=None)
```

### Exercise 2: Working with DataFrames

*Hint*: For solving this exercise, the contents of the on-site session 05 from 28/05/2025 can be helpful (see ILIAS).

This exercise concerns very similar data to exercise 1. This time, the data concerns the German states:

| State | Population | Area |
| --- | ---: | ---: |
| Baden-Württemberg	| 11280257 | 35747,85 |
| Bayern	| 13369393 | 70541,58 |
| Berlin	| 3755251	| 891,12 |
| Brandenburg	| 2573135 | 29654,38 |
| Bremen	| 684864 | 419,61 |
| Hamburg	| 1892122 | 755,09 |
| Hessen	| 6391360	| 21115,62 |
| Mecklenburg-Vorpommern | 1628378	| 23294,90 |
| Niedersachsen	| 8140242	| 47709,90 |
| Nordrhein-Westfalen	| 18139116 | 34112,72 |
| Rheinland-Pfalz	| 4159150 | 19857,97 |
| Saarland	| 992666 | 2571,52 |
| Sachsen	| 4086152 | 18449,86 |
| Sachsen-Anhalt	| 2186643 | 20467,20 |
| Schleswig-Holstein	| 2953270 | 15804,30 |
| Thüringen	| 2126846	| 16202,37 |

This exact data can also be found in csv format in the file `german_states.csv` in the subdirectory `data`.

**a)**

Use methods provided by pandas to calculate the population number and area of Germany (by summing all 16 states).
Save the results in two variables `population_sum` and `area_sum`.

Also, calculate the average population per German state and save the results in the variable `population_mean`.

<!-- grade_id: cell-7c85f544be5d362c, points: 0, grade: false, solution: true -->
```python
import pandas as pd

state_df = pd.read_csv("data/german_states.csv", sep=",")

### BEGIN SOLUTION
population_sum = state_df["Population"].sum()
area_sum = state_df["Area"].sum()
population_mean = state_df["Population"].mean()
### END SOLUTION
```

<!-- grade_id: cell-b6de4749f77dc40f, points: 0, grade: true, solution: false -->
```python
import math

assert population_sum == 84358845
assert math.isclose(area_sum, 357595.99)
assert math.isclose(population_mean, 5272427.8125)
```

**b)**

In the on-site exercise, we extracted single columns of a pandas DataFrame as instances of pandas' class `Series` (1D data objects) using `[]`.

The basic math operators (`+`, `-`, `*`, `/`) can be applied to `Series` instances of equal dimensions to perform elementwise calculations.
Use this behavior to calculate the population densities for the German states.
Save the result (type `Series`) in the variable `population_density`.

Then, add a new column `Density` to the existing DataFrame containing the calculated values for the population density.
To do this, you can use the same syntax as with Python dictionaries:

`<dataframe>[<column name>] = <values>`

<!-- grade_id: cell-a89db5bd6ff303e9, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
population_density = state_df["Population"] / state_df["Area"]
population_density

state_df["Density"] = population_density
### END SOLUTION
```

<!-- grade_id: cell-ac844ba8fad224b2, points: 0, grade: true, solution: false -->
```python
import math

assert type(population_density) is pd.Series
assert "Density" in state_df.columns

assert math.isclose(population_density.max(), 4214.080034114373)
assert math.isclose(population_density.mean(), 703.7969306103763)

assert math.isclose(state_df["Density"].max(), 4214.080034114373)
assert math.isclose(state_df["Density"].mean(), 703.7969306103763)
```

**c)**

Use a filter operation (using `[]`, as shown in the on-site exercise) to extract a DataFrame from `state_df` which contains only one row.
This row is the German state with the greatest population.

Save the result (type `DataFrame`) in the variable `max_population_state`.

<!-- grade_id: cell-8377f6a6cbc0d44f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
max_population_state = state_df[state_df["Population"] == state_df["Population"].max()]
### END SOLUTION
display(max_population_state)
```

<!-- grade_id: cell-87c15255558bad3c, points: 0, grade: true, solution: false -->
```python
assert type(max_population_state) is pd.DataFrame
assert len(max_population_state) == 1
assert max_population_state["State"].iloc[0] == "Nordrhein-Westfalen"
```

### Exercise 3: Decision Tree

Consider the following decision tree which is applied in a production's quality control:

<img align="center" width=1200 src="figures/decision_tree.png"/>

Use `if`-conditions to implement the logic of the above decision tree within a function `decision_tree()`.

The function's arguments are the attributes `material_hardness_bhn` (material hardness in Brinell) and `surface_roughness_ra_um` (surface roughness in micrometers) of a component.
The return value is the result of the quality control:
* `True`, if the quality control is passed
* `False`, if the quality control is failed

<!-- grade_id: cell-2a02663f723846d3, points: 0, grade: false, solution: true -->
```python
def decision_tree(material_hardness_bhn , surface_roughness_ra_um):
    ### BEGIN SOLUTION
    if material_hardness_bhn <= 0.8575 and surface_roughness_ra_um <= 0.5392 and material_hardness_bhn <= -1.6170:
        return False
    if material_hardness_bhn <= 0.8575 and surface_roughness_ra_um <= 0.5392 and material_hardness_bhn > -1.6170:
        return True
    if material_hardness_bhn <= 0.8575 and surface_roughness_ra_um > 0.5392:
        return False
    if material_hardness_bhn > 0.8575:
        return True
    ### END SOLUTION
```

<!-- grade_id: cell-afb1fcc190c1b1b7, points: 0, grade: true, solution: false -->
```python
assert decision_tree(-2, 0.5) == False
assert decision_tree(-2, 0.6) == False
assert decision_tree(0.8, 0.5) == True
assert decision_tree(0.8, 0.6) == False
assert decision_tree(2, 0.5) == True
assert decision_tree(2, 0.6) == True
```