## Assignment 05: Data Science

### Aufgabe 1: Erzeugen von DataFrames

*Hinweis*: Bei der Bearbeitung dieser Aufgabe hilft Ihnen die Einführung in die Bibliothek pandas der Saalübung vom 28.05.2025 (siehe ILIAS). 

In dieser Aufgabe betrachten wir folgende tabellarische Daten:

| Country | Population | Area |
| --- | ---: | ---: |
| Germany	| 84.552.242 | 357.022 |
| France	| 66.548.530 | 551.500 |
| China | 1.419.321.278 | 9.596.960 |
| India	| 1.450.935.791	| 3.287.263 |

Die Tabelle enthält zu verschiedenen Ländern die Informationen Bevölkerungszahl und Fläche in $km^2$.

**a)**

Erzeugen Sie eine Instanz der Klasse pandas DataFrame mit Namen `country_df`, die die Daten der obigen Tabelle enthält.
Dazu soll zunächst ein entsprechendes Dictionary (`country_dict`) erzeugt werden, welches anschließend zur Initialisierung des DataFrame verwendet wird.

Verwenden Sie die Funktion `display()`, um den erzeugten DataFrame ausgeben zu lassen.

<!-- grade_id: cell-01940a3cfd34615e, points: 0, grade: false, solution: true -->
```python
import pandas as pd

### BEGIN SOLUTION
country_dict = {
    "Country": ["Germany", "France", "China", "India"],
    "Population": [84552242, 66548530, 1419321278, 1450935791],
    "Area": [357022, 551500, 9596960, 3287263]
}

country_df = pd.DataFrame(country_dict)
display(country_df)
### END SOLUTION
```

<!-- grade_id: cell-933c08e46fa7fec4, points: 0, grade: true, solution: false -->
```python
assert type(country_dict) is dict
assert type(country_df) is pd.DataFrame

assert country_df.shape == (4,3)
assert set(country_df.columns) == {"Country", "Population", "Area"}
assert country_df.iloc[0]["Population"] == 84552242
assert country_df.iloc[1]["Population"] == 66548530
assert country_df.iloc[2]["Population"] == 1419321278
assert country_df.iloc[3]["Population"] == 1450935791
assert country_df.iloc[0]["Area"] == 357022
assert country_df.iloc[1]["Area"] == 551500
assert country_df.iloc[2]["Area"] == 9596960
assert country_df.iloc[3]["Area"] == 3287263
assert not country_df.isna().any(axis=None)
```

**b)**

Erzeugen Sie nun eine csv-Datei mit Namen `country_data.csv`, welche die Daten der obigen Tabelle enthält.

Anschließend sollten Sie die Daten mit Hilfe der folgenden Test-Zelle einlesen und anzeigen können.

*Hinweis:* In dieser Aufgabe muss kein Code geschrieben werden, sondern lediglich manuell die csv-Datei in dem Ordner, in dem sich auch dieses Jupyter Notebook befindet, erzeugt werden. Es gibt daher unte keine Lösungszelle, sondern nur eine Testzelle, die versucht die erzeugte Datei einzulesen und anzuzeigen.

<!-- grade_id: cell-41fec27d71f0d2bf, points: 0, grade: true, solution: false -->
```python
import pandas as pd

# Read csv file and display contents
country_csv_df = pd.read_csv("country_data.csv", sep=",")
display(country_csv_df)

# Tests concerning contained data
assert type(country_csv_df) is pd.DataFrame

assert country_df.shape == (4,3)
assert set(country_df.columns) == {"Country", "Population", "Area"}
assert country_df.iloc[0]["Population"] == 84552242
assert country_df.iloc[1]["Population"] == 66548530
assert country_df.iloc[2]["Population"] == 1419321278
assert country_df.iloc[3]["Population"] == 1450935791
assert country_df.iloc[0]["Area"] == 357022
assert country_df.iloc[1]["Area"] == 551500
assert country_df.iloc[2]["Area"] == 9596960
assert country_df.iloc[3]["Area"] == 3287263
assert not country_df.isna().any(axis=None)
```

### Aufgabe 2: Arbeiten mit DataFrames

*Hinweis*: Bei der Bearbeitung dieser Aufgabe hilft Ihnen die Einführung in die Bibliothek pandas der Saalübung vom 28.05.2025 (siehe ILIAS). 

In dieser Aufgabe arbeiten wir mit sehr ähnlichen Daten wie zuvor in Aufgabe 1, dieses Mal zu den deutschen Bundesländern:

| State | Population | Area |
| --- | ---: | ---: |
| Baden-Württemberg	| 11280257 | 35747,85 |
| Bayern	| 13369393 | 70541,58 |
| Berlin	| 3755251	| 891,12 |
| Brandenburg	| 2573135 | 29654,38 |
| Bremen	| 684864 | 419,61 |
| Hamburg	| 1892122 | 755,09 |
| Hessen	| 6391360	| 21115,62 |
| Mecklenburg-Vorpommern | 1628378	| 23294,90 |
| Niedersachsen	| 8140242	| 47709,90 |
| Nordrhein-Westfalen	| 18139116 | 34112,72 |
| Rheinland-Pfalz	| 4159150 | 19857,97 |
| Saarland	| 992666 | 2571,52 |
| Sachsen	| 4086152 | 18449,86 |
| Sachsen-Anhalt	| 2186643 | 20467,20 |
| Schleswig-Holstein	| 2953270 | 15804,30 |
| Thüringen	| 2126846	| 16202,37 |

Die entsprechenden Daten finden Sie bereits im csv-Format in der Datei `german_states.csv` im Unterordner `data`.

**a)**

Verwenden Sie von pandas bereitgestellte Methoden um die Einwohnerzahl und Fläche von Deutschland (Summe der 16 Bundesländer) zu berechnen.
Speichern Sie die Ergebnisse in den Variablen `population_sum` and `area_sum`.

Berechnen Sie außerdem die durchschnittliche Einwohnerzahl der Bundesländer und speichern Sie das Ergebnis in `population_mean`.

<!-- grade_id: cell-7c85f544be5d362c, points: 0, grade: false, solution: true -->
```python
import pandas as pd

state_df = pd.read_csv("data/german_states.csv", sep=",")

### BEGIN SOLUTION
population_sum = state_df["Population"].sum()
area_sum = state_df["Area"].sum()
population_mean = state_df["Population"].mean()
### END SOLUTION
```

<!-- grade_id: cell-b6de4749f77dc40f, points: 0, grade: true, solution: false -->
```python
import math

assert population_sum == 84358845
assert math.isclose(area_sum, 357595.99)
assert math.isclose(population_mean, 5272427.8125)
```

**b)**

In der Übung wurde gezeigt, dass mit Hilfe von `[]` einzelne Spalten eines DataFrame als Instanzen von pandas `Series` (1D-Datenobjekt) extrahiert werden können.

Auf solchen `Series`-Instanzen gleicher Dimension können einfache Rechenoperationen mit Hilfe der gewohnten Oparatoren (`+`, `-`, `*`, `/`) durchgeführt werden.
Nutzen Sie dieses Verhalten, um die Bevölkerungsdichte der Bundesländer effizient zu berechnen.
Speichern Sie das Ergebnis (Typ `Series`) in der Variablen `population_density`.

Fügen Sie dem existierenden DataFrame anschließend eine neue Spalte `Density` hinzu, welche die berechneten Werte der Bevölkerungsdichte enthält.
Dazu können Sie wie beim Hinzufügen in ein Dictionary vorgehen:

`<dataframe>[<spaltenname>] = <werte>`

<!-- grade_id: cell-a89db5bd6ff303e9, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
population_density = state_df["Population"] / state_df["Area"]
population_density

state_df["Density"] = population_density
### END SOLUTION
```

<!-- grade_id: cell-ac844ba8fad224b2, points: 0, grade: true, solution: false -->
```python
import math

assert type(population_density) is pd.Series
assert "Density" in state_df.columns

assert math.isclose(population_density.max(), 4214.080034114373)
assert math.isclose(population_density.mean(), 703.7969306103763)

assert math.isclose(state_df["Density"].max(), 4214.080034114373)
assert math.isclose(state_df["Density"].mean(), 703.7969306103763)
```

**c)**

Benutzen Sie eine Filterung (mit `[]`, wie in der Übung gezeigt), um einen DataFrame zu extrahieren, der nur noch ein Zeile enthält:
Diese Zeile soll dem Bundesland mit der größten Einwohnerzahl entsprechen.
Speichern Sie das Ergebnis in der Variablen `max_population_state`.

<!-- grade_id: cell-8377f6a6cbc0d44f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
max_population_state = state_df[state_df["Population"] == state_df["Population"].max()]
### END SOLUTION
display(max_population_state)
```

<!-- grade_id: cell-87c15255558bad3c, points: 0, grade: true, solution: false -->
```python
assert type(max_population_state) is pd.DataFrame
assert len(max_population_state) == 1
assert max_population_state["State"].iloc[0] == "Nordrhein-Westfalen"
```

### Aufgabe 3: Entscheidungsbaum

Betrachten Sie folgenden Entscheidungsbaum, welcher der Qualitätskontrolle einer Fertigung dient:

<img align="center" width=1200 src="figures/decision_tree.png"/>

Verwenden Sie eine Reihe von Bedingungen, um die Logik des abgebildeten Entscheidungsbaums in der Funktion `decision_tree` zu implementieren.

Argumente der Funktion sind dabei die Attribute `material_hardness_bhn` (Härte des Materials in Brinellhärte) und `surface_roughness_ra_um` (Rauigkeit der Oberfläche in Mikrometern) einer Komponente.
Rückgabe der Funktion ist das Ergebnis der Qualitätskontrolle:
* `True`, falls die Qualitätskontrolle bestanden ist ("Pass QC")
* `False`, falls die Qualitätskontrolle nicht bestanden ist ("Fail QC")

<!-- grade_id: cell-2a02663f723846d3, points: 0, grade: false, solution: true -->
```python
def decision_tree(material_hardness_bhn , surface_roughness_ra_um):
    ### BEGIN SOLUTION
    if material_hardness_bhn <= 0.8575 and surface_roughness_ra_um <= 0.5392 and material_hardness_bhn <= -1.6170:
        return False
    if material_hardness_bhn <= 0.8575 and surface_roughness_ra_um <= 0.5392 and material_hardness_bhn > -1.6170:
        return True
    if material_hardness_bhn <= 0.8575 and surface_roughness_ra_um > 0.5392:
        return False
    if material_hardness_bhn > 0.8575:
        return True
    ### END SOLUTION
```

<!-- grade_id: cell-afb1fcc190c1b1b7, points: 0, grade: true, solution: false -->
```python
assert decision_tree(-2, 0.5) == False
assert decision_tree(-2, 0.6) == False
assert decision_tree(0.8, 0.5) == True
assert decision_tree(0.8, 0.6) == False
assert decision_tree(2, 0.5) == True
assert decision_tree(2, 0.6) == True
```