<img align="right" width=160 src="https://upload.wikimedia.org/wikipedia/commons/thumb/3/3a/Logo_KIT.svg/1920px-Logo_KIT.svg.png"/>

# IDS Assignments

**Allgemeine Hinweise**:

* Im Rahmen der Tutorien haben Sie 60 min Zeit für die Bearbeitung der Aufgaben. Bitte wenden Sie sich umgehend an Ihren Tutor, sobald die Bearbeitung abgeschlossen ist.
* Bitte überprüfen Sie, dass die enthaltenen Tests fehlerfrei ausgeführt werden können.
* Bei Unklarheiten, Fragen oder Problemen, wenden Sie sich bitte **unmittelbar** an Ihren Tutor (nicht erst zum Ende des Tutoriums) und holen Sie sich Unterstützung ein.
* Das Arbeiten in 2-er Teams ist erlaubt und erwünscht! :)


**General Notes**:

* During the tutorials, you have 60 minutes to work on the tasks. Please contact your tutor immediately once you have completed the work.
* Make sure that the included tests can run without errors.
* If anything is unclear, or if you have questions or problems, please contact your tutor immediately (not just at the end of the tutorial) and ask for support.
* Working in pairs (2-person teams) is allowed and encouraged! :)

## Assignment 02

### Exercise 1: Working with files and data

#### 1 A)

The file `temperatures_daily.csv` in the subfolder `data` contains the daily maximum temperatures recorded in Rheinstetten for the period 02.08.2023 - 01.02.2025.

Each line has the format `<date>;<temperature>`.

Read in the file and save all the data it contains as a date-temperature tuple (`(date, temperature)`) in the `single_temps` list.
The second entry of the tuple `temperature` should have the data type `float`.

Opening the file and iterating over the individual lines is already implemented.
To solve the task, you only need to extract the date and temperature from the line strings, convert the data type if necessary and write it to the target array `single_temps`.

*Note*: It is recommended that you first open the file in a text editor, for example, and view it manually. Note in particular that the file contains a header line.

*Note*: The function `datetime.date.fromisoformat(<str>)` can be used to convert the date string into an object of the class `datetime.date`. Alternatively, you can continue working with the strings.

**Attention**: The file is not error-free. For example, individual values may be missing or incorrect. However, formatting the values to `float` is possible for all entries.

<!-- grade_id: cell-695438bf79cc75c0, points: 0, grade: false, solution: true -->
```python
import datetime

single_temps = []

input_file = "./data/temperatures_daily.csv"

# Open temperature file
with open(input_file, "r") as file_opened:
    # Iterate over lines (Str type) in the file
    for i, line in enumerate(file_opened):
        # YOUR CODE HERE
        
        # END OF YOUR CODE
    # Close file
    file_opened.close()
```

<!-- grade_id: cell-45aa9d5504a453b4, points: 1, grade: true, solution: false -->
```python
assert len(single_temps) == 547
for date, temp in single_temps:
    assert type(temp) == float
assert single_temps[0][1] == 18.2
assert single_temps[-1][1] == 1.4
```

#### 1 B)

Complete the following implementation, which should calculate the average daily maximum temperature for all contained months (08/2023 - 02/2025).

The aim is to fill the dictionary `average_high` so that the key corresponds to the respective month (format `mm/YYYYY`, data type `str`) and the value of the calculated average temperature (`float`) corresponds to the values of the month.

In order to calculate the average maximum values by month, it can be helpful to first group all temperature values by month.
Therefore, we first fill a dictionary `temps_by_month`, which should contain a list of all values measured in this month under the month key (`mm/YYYYY`).

**Attention**: The file is not error-free. For example, individual values may be missing or incorrect. Please proceed as follows:
* Due to missing values, the number of values actually present must be taken into account when calculating the average.
* Incorrect values are more difficult to identify and can initially be ignored in the calculation.

<!-- grade_id: cell-479353e3a94b8137, points: 0, grade: false, solution: true -->
```python
temps_by_month = {}
average_high = {}

for date, temp in single_temps:
    # Build month-Str as instructed (mm/YYYY)
    # YOUR CODE HERE
    
    # END OF YOUR CODE
    if month not in temps_by_month:
        temps_by_month[month] = []
    temps_by_month[month].append(temp)

# Iterate over temps_by_month, compute the monthly average high for each month and write it into the Dict average_high
# YOUR CODE HERE

# END OF YOUR CODE
```

<!-- grade_id: cell-2315c38b77b04863, points: 1, grade: true, solution: false -->
```python
import math

assert len(average_high) == 19
for key, value in average_high.items():
    assert type(key) is str
    assert type(value) is float
assert "10/2023" in average_high
assert math.isclose(average_high["11/2023"], 7.08, rel_tol=1e-5)
assert math.isclose(average_high["05/2024"], 18.2903, rel_tol=1e-5)
assert math.isclose(average_high["02/2025"], 1.4, rel_tol=1e-5)
```

#### 1 C) *BONUS*

* *This task is optional. If it is solved successfully, it can replace any other subtask of this task sheet.*

Now write the previously calculated monthly average temperatures into a file named ``temperatures_monthly_averages.csv`` in the subfolder ``data``.

The contents of the file should correspond to the csv format with the separator `;` and contain the two columns "Month" and "Average Temperature".

* A header line must therefore first be written with the two column names (note the separator!)
* This is followed by one line per month in the correct format
* If in doubt, use the format of the file imported in task A) as a guide

<!-- grade_id: cell-f7f05da090a0da2f, points: 0, grade: false, solution: true -->
```python
output_file = "./data/temperatures_monthly_averages.csv"

with open(output_file, "w") as file_opened:
    # YOUR CODE HERE
    
    # END OF YOUR CODE
    file_opened.close()
```

<!-- grade_id: cell-9e442d447aa30579, points: 1, grade: true, solution: false -->
```python
import os
assert os.path.isfile(output_file)
with open(output_file, "r") as file_reader:
    content = file_reader.read()
    file_reader.close()
lines = content.split("\n")
assert "Month" in lines[0]
assert "Average Temperature" in lines[0]

### Check value consistency
input_temp_sum = sum([entry for entry in average_high.values()])

temp_sum = 0
for line in lines[1:]:
    if len(line) > 0:
        temp_sum += float(line.rsplit(";")[1])
assert math.isclose(temp_sum, input_temp_sum, rel_tol=1e-3)
```

### Exercise 2: Object-oriented programming

#### 2 A)

Consider the following, still very simple implementation of a class `Car`.

Complete the magic method `__eq__` so that it returns `True` exactly when the two compared instances of the class are indistinguishable, i.e. the values of all attributes are identical.

Three instances of the class `Car` are then to be created as follows:
* `car_1` and `car_2` are not “identical” (comparison with `__eq__` results in `False`)
* `car_1` and `car_3` are "identical" (comparison with `__eq__` results in `True`)

<!-- grade_id: cell-7d02de331f43302f, points: 0, grade: false, solution: true -->
```python
class Car:
    
    def __init__(self, color):
        self.color = color
        self.mileage = 0
        
    def __eq__(self, other):
        # YOUR CODE HERE
        
        # END OF YOUR CODE
    
    def drive(self, miles):
        self.mileage += miles
    
car_1 = None
car_2 = None
car_3 = None

# YOUR CODE HERE

# END OF YOUR CODE
```

<!-- grade_id: cell-4724e0fd3c4b4ff4, points: 1, grade: true, solution: false -->
```python
assert type(car_1) == Car
assert type(car_2) == Car
assert type(car_3) == Car
assert car_1 == car_3
assert car_1 != car_2
```

#### 2 B)

In this task, fuel consumption logic is to be added to the ‘Car’ class.

The fuel consumption of the model per 100 km driven (`fuel_consumption_liter_per_100km`) is stored in each instance.
In addition to a total odometer (`mileage`), each instance also has an attribute containing the current fuel level (`fuel_liters`).

The methods `drive()`, `fuel_up()` and `get_remaining_distance()` are now to be implemented.

When `drive` is called, a corresponding distance is traveled, i.e. the odometer must be incremented and the current fuel level updated according to consumption.

A call to `fuel_up` means a refill of fuel, here the current fuel level is to be adjusted accordingly.

The method `get_remaining_distance` returns the distance in kilometers that can be covered with the current fuel level.

*Note*: A constant consumption of the vehicle is assumed here, which only depends on the distance to be covered.

<!-- grade_id: cell-6f85adabf0914287, points: 0, grade: false, solution: true -->
```python
class Car:
    
    def __init__(self, color, fuel_consumption):
        self.color = color
        self.fuel_consumption_liter_per_100km = fuel_consumption
        self.fuel_liters = 0
        self.mileage = 0
    
    def drive(self, distance_km):
        # YOUR CODE HERE
        
        # END OF YOUR CODE
        
    def fuel_up(self, liters_added):
        # YOUR CODE HERE
        
        # END OF YOUR CODE
        
    def get_remaining_distance(self):
        # YOUR CODE HERE
        
        # END OF YOUR CODE
```

<!-- grade_id: cell-728105f777ae311f, points: 1, grade: true, solution: false -->
```python
my_car = Car("red", 5.2)
my_car.fuel_up(40)
my_car.drive(120)
assert my_car.mileage == 120
assert my_car.fuel_liters == 33.76
assert math.isclose(my_car.get_remaining_distance(), 649.2307, abs_tol=0.01)
```