<img align="right" width=160 src="https://upload.wikimedia.org/wikipedia/commons/thumb/3/3a/Logo_KIT.svg/1920px-Logo_KIT.svg.png"/>

# IDS Assignments

**Allgemeine Hinweise**:

* Im Rahmen der Tutorien haben Sie 60 min Zeit für die Bearbeitung der Aufgaben. Bitte wenden Sie sich umgehend an Ihren Tutor, sobald die Bearbeitung abgeschlossen ist.
* Bitte überprüfen Sie, dass die enthaltenen Tests fehlerfrei ausgeführt werden können.
* Bei Unklarheiten, Fragen oder Problemen, wenden Sie sich bitte **unmittelbar** an Ihren Tutor (nicht erst zum Ende des Tutoriums) und holen Sie sich Unterstützung ein.
* Das Arbeiten in 2-er Teams ist erlaubt und erwünscht! :)


**General Notes**:

* During the tutorials, you have 60 minutes to work on the tasks. Please contact your tutor immediately once you have completed the work.
* Make sure that the included tests can run without errors.
* If anything is unclear, or if you have questions or problems, please contact your tutor immediately (not just at the end of the tutorial) and ask for support.
* Working in pairs (2-person teams) is allowed and encouraged! :)

## Assignment 04: Algorithms

### Exercise 1: Conversion Between Number Systems

#### 1 A) Binary to Decimal

Implement an algorithm to convert a binary string to the decimal system. Continue implementing the function `binary_to_decimal`

The formula is:

$(b_{n-1}\ldots b_1 b_0)_2 = \left(\sum_{i=0}^{n-1} b_i 2^i\right)_{10}$

**Hint:** The output of the algorithm is an integer number


<!-- grade_id: cell-42a04ad1b4d179e1, points: 0, grade: false, solution: true -->
```python
def binary_to_decimal(binary_string):
    decimal_value = 0
    # YOUR CODE HERE
    
    # END OF YOUR CODE
    return decimal_value
```

<!-- grade_id: cell-f0c46b0dfe143394, points: 1, grade: true, solution: false -->
```python
assert binary_to_decimal

assert binary_to_decimal("101") == 5
assert binary_to_decimal("10110") == 22
print("TEST PASSED!")
```

### 1 B) Horner's Method

Write a function called `horners_method`. This function will take a regular decimal number and convert it into a different number system (like binary, octal, or hexadecimal) using a technique called Horner's method. The function should then give you back the converted number in that new system.

The function needs two pieces of information:

1.  `number`: This is the decimal number you want to convert.
2.  `base`: This is the base of the new number system you want to convert to (e.g., 2 for binary, 8 for octal, 16 for hexadecimal). If you don't specify a `base`, the function should automatically assume you want to convert to binary (base 2).

**Special note for Hexadecimal:** If you're converting to hexadecimal (base 16), you'll sometimes get remainders that are 10 or greater. For these, you'll need to use a special dictionary called `dec_hex_dict` to change those numbers into their correct hexadecimal letters (like 'A' for 10, 'B' for 11, etc.).

<!-- grade_id: cell-c5e514a0a6ad9a94, points: 0, grade: false, solution: true -->
```python
dec_hex_dict = {
    '0': '0',
     '1': '1',
     '2': '2',
     '3': '3',
     '4': '4',
     '5': '5',
     '6': '6',
     '7': '7',
     '8': '8',
     '9': '9',
     '10': 'A',
     '11': 'B',
     '12': 'C',
     '13': 'D',
     '14': 'E',
     '15': 'F'
}

# YOUR CODE HERE

# END OF YOUR CODE
```

<!-- grade_id: cell-4a7f9a595efd4a86, points: 0, grade: true, solution: false -->
```python
assert horners_method
assert horners_method(22) == "10110"
assert horners_method(31) == "11111"
assert horners_method(22, 8) == "26"
assert horners_method(31, 8) == "37"
assert horners_method(10, 10) == "10"
assert horners_method(511, 16) == "1FF"
assert horners_method(4090, 16) == "FFA"
print("TEST PASSED!")
```

#### Exercise 2: Swap Sort

The sorting algorithm *Swap Sort* can also be used to sort lists with comparable elements.
However, the list must not contain any duplicates.

The following implementation of the algorithm is given:

```python
def swap_sort(list_to_sort):
    i = 0
    while i < len(list_to_sort):
        m = 0
        for other_entry in list_to_sort:
            if other_entry < list_to_sort[i]:
                m += 1
        if i == m:
            i += 1
        else:
            h = list_to_sort[m]
            list_to_sort[m] = list_to_sort[i]
            list_to_sort[i] = h
    return list_to_sort
```

Extend the existing implementation so that the number of iterations performed is returned in addition to the sorted list.
To do this, copy the above implementation into the following code cell.

**No lines of code of the algorithm should be changed, only new lines should be added.**

*Note*: By **an iteration** of the algorithm is meant a run consisting of counting the smaller elements of the list and then swapping two elements or increasing the index ``i``.

Then fill the ``unsorted`` list with the numbers ``{5,4,1,3,2}`` so that the maximum number of iterations is reached. How many iterations is that?

<!-- grade_id: cell-7223e0f31a61643f, points: 0, grade: false, solution: true -->
```python
def swap_sort(list_to_sort):
    # YOUR CODE HERE
    
    # END OF YOUR CODE
    return list_to_sort, iteration_counter

# YOUR CODE HERE

# END OF YOUR CODE
sorted_list, iteration_number = swap_sort(unsorted)
print(iteration_number)
```

<!-- grade_id: cell-62f3175b4df8e270, points: 1, grade: true, solution: false -->
```python
import random
unsorted = list(range(10))
assert swap_sort(unsorted)[1] == 10
random.shuffle(unsorted)
assert swap_sort(unsorted)[1] >= 10
for i in range(len(unsorted) - 1):
    assert unsorted[i] <= unsorted[i+1]
print("TEST PASSED!")
```

### Exercise 3: Sieve of Eratosthenes

In this task, an algorithm for determining prime numbers is to be implemented, the so-called *Sieve of Eratosthenes*.

The algorithm works as follows:
* To determine all prime numbers up to an integer ``n > 1``, start with a list of all integers from ``2`` to ``n``.
* Starting with ``i=2`` delete all multiples of ``i`` from the list
* Repeat the previous step with the next largest entry in the list
* If ``i`` is the largest element in the list, the algorithm is complete. The list now contains all prime numbers between ``2`` and ``n``.
  

<!-- grade_id: cell-7b127a8ea604e223, points: 0, grade: false, solution: true -->
```python
def sieve_of_eratosthenes(n):
    prime_numbers = list(range(2, n+1))
    # YOUR CODE HERE
    
    # END OF YOUR CODE
    return prime_numbers
```

<!-- grade_id: cell-73f72fff1cd7a1cf, points: 2, grade: true, solution: false -->
```python
assert len(sieve_of_eratosthenes(10)) == 4
assert 2 in sieve_of_eratosthenes(10)
assert 3 in sieve_of_eratosthenes(10)
assert 5 in sieve_of_eratosthenes(10)
assert 7 in sieve_of_eratosthenes(10)
assert len(sieve_of_eratosthenes(29)) == 10
print("TEST PASSED!")
```