# Assignment 01

### Aufgabe 1: Funktionen

Implementiere die Funktion ``euclidean_distance_2d`` welche die euklidische Distanz zwischen zwei Punkten ``(x_1, y_1)`` und ``(x_2, y_2)`` im zweidimensionalen Raum berechnet.
Die Funktion nimmt dabei die Koordinaten ``x_1, y_1, x_2, y_2`` als vier Parameter mit numerischen Werten (``float``) entgegen.

Berechne anschließend die Distanz zwischen den beiden bereits definierten 2-D Punkten ``point_1`` und ``point_2`` unter Verwendung der zuvor implementierten Funktion. 

*Hinweis*: Die euklidische Distanz ``d`` kann mit folgender Formel berechnet werden: $d = \sqrt{|x_2-x_1|^2 + |y_2-y_1|^2}$

<!-- grade_id: cell-68371f066db3bc37, points: 0, grade: false, solution: true -->
```python
import math

def euclidean_distance_2d(x_1, y_1, x_2, y_2):
    ### BEGIN SOLUTION

    ### END SOLUTION
    return distance

point_1 = [0.23, -1.98]
point_2 = [-10.2548876, 4]

dist = None
# Compute Euclidean distance between point_1 and point_2 and save as dist.
### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-1083788a70e7a1fb, points: 1, grade: true, solution: false -->
```python
assert dist
assert euclidean_distance_2d(0, 0, 0, 0) == 0
assert math.isclose(euclidean_distance_2d(0.123456, -0.987654, -0.123456, 0.987654), 1.9906800924829686, rel_tol=1e-5)
assert math.isclose(dist, 12.070346638959206, rel_tol=1e-5)
print("TESTS PASSED!")
```

### Aufgabe 2: Dictionaries

#### Aufgabe 2 a)

Gegeben sind eine Liste mit Schlüsseln `keys_list`, sowie eine Liste mit Werten `values_list`.

Befülle das leere Dictionary `my_dict` sodass der i-te Eintrag der Liste `keys_list` ein Schlüssel im Dictionary ist, dessen Wert dem i-ten Eintrag der Liste `values_list` entspricht.

Sonstige Schlüssel oder Werte dürfen nicht in `my_dict` enthalten sein.

Benutze die existierenden Listen, um das Dictionary zu befüllen und gib die enthaltenen Werte nicht erneut ein!

<!-- grade_id: cell-e3ea4658fa9266cd, points: 0, grade: false, solution: true -->
```python
my_dict = {}

keys_list = [1, 2, 3, 4, 5]
values_list = ["monday", "tuesday", "wednesday", "thursday", "friday"]

### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-26cb72fbcbf64dd0, points: 1, grade: true, solution: false -->
```python
assert my_dict[1] == "monday"
assert my_dict[3] == "wednesday"
assert len(my_dict) == 5
assert 0 not in my_dict
assert 6 not in my_dict
print("TESTS PASSED!")
```

#### Aufgabe 2 b)

Gegeben ist ein Dictionary `measurements_dict`, dessen Werte gemessenen Temperaturwerten in °C entsprechen.

Die Messwerte sollen innerhalb des Dictionaries in Fahrenheit konvertiert werden. 
Die zugehörigen Schlüssel sollen erhalten bleiben.
(Sie dürfen die Messwerte im Dictionary `measurements_dict` direkt entsprechend ändern.)

**Achtung**: Die Schlüssel und Werte müssen dabei dem ursprünglichen Dictionary `measurements_dict` entnommen werden und dürfen nicht manuell im Code eingesetzt werden.

*Hinweis*: Ist $c$ ein Temperaturwert in Celsius, so lässt sich der Wert mit Hilfe der folgenden Formel in einen Wert $f$ der Einheiten Fahrenheit umwandeln:
<p style="text-align: center;">
$f = \frac{9}{5} \cdot c + 32$
</p>

<!-- grade_id: cell-0ed48280a015f3e3, points: 0, grade: false, solution: true -->
```python
measurements_dict = {
    "1": 23.5,
    "2": 23.9,
    "3": 24.7,
    "4": 25.1,
    "5": 25.0,
    "6": 24.9,
    "7": 22.3,
    "8": 19.6,
    "9": 12.5,
    "10": 5.0
}

### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-b16c3535e364ff1c, points: 1, grade: true, solution: false -->
```python
assert len(measurements_dict) == 10
assert sum(measurements_dict.values()) == 691.7
print("TESTS PASSED!")
```

#### Aufgabe 2 c)

Implementiere eine Funktion `add_entry_for_key`, welche einem existierenden Dictionary `my_dict` einer Liste unter dem Schlüssel `key` den Wert `new_entry` hinzufügt.
Ist der Schlüssel `key` noch nicht in `my_dict` enthalten, so soll als Wert zu `key` eine neue Liste mit dem Eintrag `new_entry` als Wert eingefügt werden.
Somit sind ausnahmslos alle mit Hilfe von `add_entry_for_key` hinzugefügten Werte in `my_dict` Listen.

Die Funktion nimmt die Parameter `my_dict`, `key` und `new_entry` in dieser Reihenfolge entgegen und erzeugt keine Rückgabe.

Nutze die Funktion `add_entry_for_key` anschließend, um entsprechend der unten stehenden Tabelle
* die Note von Tom im Dictionary `exam_results` zu ergänzen.
* dem Dictionary `exam_results` unter dem Schlüssel `seat_number` die Sitzplatznummern der Studierenden im Klausurhörsaal hinzuzufügen.

| Name | Note | Sitzplatz |
| ---- | ---- | --------- |
| Lisa | 2,0 | 11 |
| Lena | 2,7 | 21 |
| Luisa | 1,3 | 31 |
| Tom | 1,7 | 41 |

*Hinweis*: Beachten Sie jeweils die Reihenfolge der Listen!

<!-- grade_id: cell-5ba830463fba67c0, points: 0, grade: false, solution: true -->
```python
def add_entry_for_key(my_dict, key, new_entry):
    ### BEGIN SOLUTION

    ### END SOLUTION
    
exam_results = {
    "students": ["Lisa", "Lena", "Luisa"],
    "grades": [2.0, 2.7, 1.3]
}

add_entry_for_key(exam_results, "students", "Tom")

### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-4f94ae01b36b616e, points: 1, grade: true, solution: false -->
```python
assert "students" in exam_results
assert "grades" in exam_results
assert "seat_number" in exam_results
assert exam_results["seat_number"][3] == 41
assert len(exam_results) == 3
for key, value in exam_results.items():
    assert len(value) == 4
print("TESTS PASSED!")
```

### Aufgabe 3: Bedingungen und Schleifen

Gegeben sind die drei folgenden Dictionaries:

* Das Dictionary `inventory` enthält unter der entsprechenden Teilnummer als Schlüssel die Anzahl der im Inventar vorhandenen Teile.
* Das Dictionary `requirements` enthält eine Liste benötigter Teile, um ein Produkt zu fertigen (Teilenummer ist der Schlüssel, der Wert gibt die Anzahl der benötigten Teile des Typs an.
* Das Dictionary `price_list` enthält unter einer Teilnummer als Schlüssel den Preis für ein Teil des Typs.

Implementiere einen Codeblock, der die Kosten berechnet, um die fehlenden Teile zur Anfertigung eines Produkts entsprechend `requirements` zu fertigen.
Prüfe dazu, welche Teile in welcher Anzahl beschafft werden müssen, damit in Summe (aus Inventar und Neubeschaffung) genügend Teile zur Fertigung vorhanden sind.
Berechne anschließend den Preis der neu zu beschaffenden Teile und speichere diesen in der Variablen `missing_parts_costs`.

**Achtung**: Entnehmen Sie die jeweiligen Werte unbedingt den existierenden Dictionaries. Sie dürfen keine Werte (`str` oder numerisch) in Ihrem Lösungscode verwenden.

<!-- grade_id: cell-cb23657f6d285409, points: 0, grade: false, solution: true -->
```python
inventory = {
    "12343": 2,
    "12344": 7,
    "12347": 4,
    "12348": 1,
    "12349": 12,
    "12350": 4,
    "12351": 8,
    "12352": 2,
    "12355": 2,
    "12357": 4,
    "12358": 6
}

requirements = {
    "12344": 4,
    "12346": 4,
    "12347": 4,
    "12348": 4,
    "12349": 4,
    "12352": 6,
    "12353": 6,
    "12354": 6,
    "12358": 6
}

price_list = {
    "12343": 2.99,
    "12344": 2.99,
    "12345": 3.87,
    "12346": 3.87,
    "12347": 1.72,
    "12348": 1.79,
    "12349": 4.14,
    "12350": 4.36,
    "12351": 5.12,
    "12352": 2.19,
    "12353": 2.88,
    "12354": 6.48,
    "12355": 2.78,
    "12356": 5.07,
    "12357": 4.88,
    "12358": 3.49
}

missing_parts_costs = 0

### BEGIN SOLUTION


### END SOLUTION
    
print("Summed costs for missing parts: {}".format(missing_parts_costs))
```

<!-- grade_id: cell-eb42148c379a5022, points: 1, grade: true, solution: false -->
```python
import math
assert math.isclose(missing_parts_costs, 85.77)
print("TESTS PASSED!")
```