## Algorithms

### Exercise 1: Insertion Sort

In this exercise, we will implement a sorting algorithm called Insertion Sort. 

This algorithm sorts a list of numbers (`my_list`) from smallest to largest. It does this by taking each number from the original list (`my_list`) and inserting it into its correct position within a new list (`sorted_list`) that is being built up to be sorted.

Steps: 

- Define a new empty list called `sorted_list`, which will be returned when the function terminates. 

- Loop through each element in the `my_list` (using a for loop). Do the following in each loop iteration:
    - Initialize an index variable called `i` and set its value to 0
    - Then, keep incrementing the `i` by one to find the correct position for the current element from `my_list` using a while loop with the following conditions:
        - `i` is less than the length of `sorted_list`
        - and the current element is greater than or equal to the current element of `sorted_list` (`sorted_list[i]`)
          
          
    - Once the while loop terminates, insert the current element from `my_list` into `sorted_list` at the `i`-th index

<!-- grade_id: cell-5dadede7ab35fe0e, points: 0, grade: false, solution: true -->
```python
def insertion_sort(my_list):
    ### BEGIN SOLUTION

    sorted_list = []

    for element in my_list:
        i = 0
        while i < len(sorted_list) and element >= sorted_list[i]:
            i += 1
        sorted_list.insert(i, element)
    ### END SOLUTION
    return sorted_list
```

<!-- grade_id: cell-d9c082c9ccd289e3, points: 2, grade: true, solution: false -->
```python
my_list = [2,3,1,5,4,3,2]
sorted_list = insertion_sort(my_list)
assert len(sorted_list) == 7
for i in range(len(sorted_list)-1):
    assert sorted_list[i] <= sorted_list[i+1]

### BEGIN HIDDEN TESTS
import numpy as np
np.random.seed(89)
liste_r = list([np.random.randint(500) for i in range(100)])
sortierte_liste_r = insertion_sort(liste_r)
assert len(sortierte_liste_r) == 100
for i in range(len(sortierte_liste_r)-1):
    assert sortierte_liste_r[i] <= sortierte_liste_r[i+1]
    

### END HIDDEN TESTS
```

### Exercise 2: Rotation in $\mathbb{R}^3$

This exercise deals with the rotation of vectors in $\mathbb{R}^3$.

Below are various algorithms for calculating vector rotations around specific axes and angles.

The rotation matrices $R \in \mathbb{R}^{3x3}$ for rotations around the three coordinate axes in $\mathbb{R}^3$ can be calculated as:
* Rotation around the x-axis:

${\displaystyle R_{x}(\alpha )={\begin{pmatrix}1&0&0\\0&\cos \alpha &-\sin \alpha \\0&\sin \alpha &\cos \alpha \end{pmatrix}}}$

* Rotation around the y-axis:

${\displaystyle R_{y}(\alpha )={\begin{pmatrix}\cos \alpha &0&\sin \alpha \\0&1&0\\-\sin \alpha &0&\cos \alpha \end{pmatrix}}}$

* Rotation around the z-axis:

${\displaystyle R_{z}(\alpha )={\begin{pmatrix}\cos \alpha &-\sin \alpha &0\\\sin \alpha &\cos \alpha &0\\0&0&1\end{pmatrix}}}$



Let $v \in \mathbb{R}^3$ be a 3-dimensional vector.
The matrix-vector product $v_r = R_{x}(\alpha) \cdot v \in \mathbb{R}^3$, corresponds to the vector $v$ rotated by angle $\alpha$ around the coordinate axis $x=\begin{pmatrix}1\\0\\0\end{pmatrix}$. (Analog for the y and z axes.)


Implement the function `rotate_matrix()`, which calculates the corresponding rotation matrix for a rotation around the x-, y- or z-axis of the $\mathbb{R}^3$ by a certain angle `angular_rad`. The function takes an integer argument `axis`, which corresponds to the index of the coordinate axis around which the rotation is to take place (0: x-axis, 1: y-axis, 2: z-axis).

Hint: The matrix is represented as a nested list

<!-- grade_id: cell-7afaae70fbdfc508, points: 0, grade: false, solution: true -->
```python
from math import cos, sin

def rotate_matrix(axis, angular_rad):
    if axis not in [0, 1, 2]:
        raise RuntimeError("Undefined rotation axis! (Choose an integer 0: x-axis, 1: y-axis, 2: z-axis)")

    r_matrix = [[0, 0, 0], [0, 0, 0], [0, 0, 0]]

    ### BEGIN SOLUTION
    
    cos_r = cos(angular_rad)
    sin_r = sin(angular_rad)

    r_matrix[axis][axis] = 1
    
    if axis == 0:  # X-axis rotation
        r_matrix[1][1] = cos_r
        r_matrix[1][2] = -sin_r
        r_matrix[2][1] = sin_r
        r_matrix[2][2] = cos_r
    elif axis == 1:  # Y-axis rotation
        r_matrix[0][0] = cos_r
        r_matrix[0][2] = sin_r
        r_matrix[2][0] = -sin_r
        r_matrix[2][2] = cos_r
    elif axis == 2:  # Z-axis rotation
        r_matrix[0][0] = cos_r
        r_matrix[0][1] = -sin_r
        r_matrix[1][0] = sin_r
        r_matrix[1][1] = cos_r
        
    ### END SOLUTION
    return r_matrix
```

<!-- grade_id: cell-ebbcf3446ce5220f, points: 2, grade: true, solution: false -->
```python
 import math

# Test unit matrix with rotation angle 0
for i in range(0, 3):
    r_matrix = rotate_matrix(i, 0)
    assert math.isclose(r_matrix[0][0], 1, abs_tol=1e-6) and math.isclose(r_matrix[1][1], 1, abs_tol=1e-6) and math.isclose(r_matrix[2][2], 1, abs_tol=1e-6)
    assert math.isclose(sum(r_matrix[0]), 1, abs_tol=1e-6) and math.isclose(sum(r_matrix[1]), 1, abs_tol=1e-6) and math.isclose(sum(r_matrix[2]), 1, abs_tol=1e-6)
    
# Test Case: Rotation 90°
r_matrix = rotate_matrix(2, 0.5 * math.pi)
assert math.isclose(sum(r_matrix[0]), -1, abs_tol=1e-6)
assert math.isclose(sum(r_matrix[1]), 1, abs_tol=1e-6)
assert math.isclose(sum(r_matrix[2]), 1, abs_tol=1e-6)
assert math.isclose(r_matrix[0][1], -1, abs_tol=1e-6)
assert math.isclose(r_matrix[1][0], 1, abs_tol=1e-6)
assert math.isclose(r_matrix[2][2], 1, abs_tol=1e-6)

### BEGIN HIDDEN TESTS
import numpy as np
v = np.array([0.5,0.25,1.5])
r = np.array(rotate_matrix(1, 72 / 180 * math.pi))
res = np.matmul(r, v)
assert np.abs(res[0] - 1.581093) < 1e-5
assert np.abs(res[1] - 0.25) < 1e-5
assert np.abs(res[2] - (-0.012003)) < 1e-5

r = np.array(rotate_matrix(0, -30 / 180 * math.pi))
res_2 = np.matmul(r, res)
assert np.abs(res_2[0] - 1.581093) < 1e-5
assert np.abs(res_2[1] - 0.210505) < 1e-5
assert np.abs(res_2[2] - (-0.135395)) < 1e-5

r_1 = np.array(rotate_matrix(2, 134 / 180 * math.pi))
r_2 = np.array(rotate_matrix(2, -134 / 180 * math.pi))
res_3 = np.matmul(r_1, r_2)

assert np.linalg.norm(res_3 - np.identity(3)) < 1e-5
### END HIDDEN TESTS
```

### Exercise 3: Matrix-Vector Multiplication


**a)**

In part of the exercise, the matrix-vector product is to be implemented in $\mathbb{R}^n$.

For a matrix `matrix` from $\mathbb{R}^{m x n}$ and a vector `vector` from $\mathbb{R}^{n}$, implement 
a function `matrix_vector_product(matrix, vector)` that calculates and returns the result of the matrix-vector multiplication `result_vector` $\in \mathbb{R}^m$.

Matrices are again represented as nested lists (for $\mathbb{R}^{m x n}$: a list with m entries, each entry is a list with n numerical entries) and vectors as lists.


The matrix-vector product is computed as follows:
Given an $m \times n$ matrix $A$ and an $n \times 1$ column vector $\mathbf{x}$:

$A = \begin{pmatrix} a_{11} & a_{12} & \cdots & a_{1n} \\ a_{21} & a_{22} & \cdots & a_{2n} \\ \vdots & \vdots & \ddots & \vdots \\ a_{m1} & a_{m2} & \cdots & a_{mn} \end{pmatrix}$

$\mathbf{x} = \begin{pmatrix} x_1 \\ x_2 \\ \vdots \\ x_n \end{pmatrix}$

The product $A\mathbf{x}$ is an $m \times 1$ column vector $\mathbf{b}$:

$\mathbf{b} = A\mathbf{x} = \begin{pmatrix} b_1 \\ b_2 \\ \vdots \\ b_m \end{pmatrix}$

where each element $b_i$ of the resulting vector is calculated as the dot product of the $i$-th row of matrix $A$ and the vector $\mathbf{x}$:

$b_i = a_{i1}x_1 + a_{i2}x_2 + \cdots + a_{in}x_n = \sum_{j=1}^{n} a_{ij}x_j$



<!-- grade_id: cell-51277670c0fb5bc5, points: 0, grade: false, solution: true -->
```python
def matrix_vector_product(matrix, vector):
    ### BEGIN SOLUTION
    result_vector = []
    for i in range(len(matrix)):
        entry = 0
        for j in range(len(matrix[i])):
            entry += matrix[i][j] * vector[j]
        result_vector.append(entry)
    ### END SOLUTION
    return result_vector
```

<!-- grade_id: cell-61f842d333badb70, points: 2, grade: true, solution: false -->
```python
assert matrix_vector_product([[1]], [1]) == [1]
assert len(matrix_vector_product([[1,0,0]],[1,1,1])) == 1
assert len(matrix_vector_product([[1,0,0],[0,1,0],[0,0,1]],[1,1,1])) == 3
assert len(matrix_vector_product([[1,0,0],[0,1,0],[0,0,1],[1,0,0],[0,1,0],[0,0,1]],[1,1,1])) == 6

### BEGIN HIDDEN TESTS
vector = [2.37, -3.41, -0.01]
matrix = [[1.01, -0.29, 12.91], [3.98, 2.77, -5.81]]
erg = matrix_vector_product(matrix, vector)
assert len(erg) == 2
assert abs(erg[0] - 3.2535) < 1e-5
assert abs(erg[1] - 0.045) < 1e-5
### END HIDDEN TESTS
```

**b)**

What is the complexity (Big-O Notation) of the `multiply_matrices()` for matrices of dimension $\mathbb{R}^{m x n}$?

1.) $O(n^2)$

2.) $O(m^2)$

3.) $O(m \cdot n)$

4.) $O(2^{m+n})$

5.) $O(1)$

<!-- grade_id: cell-60cf8613b7494458, points: 1, grade: true, solution: false -->
```python
# Right answer is 3 O(m . n)
### BEGIN HIDDEN TESTS
### END HIDDEN TESTS
```

### Exercise 4: Matrix Multiplication

**a)**
In this Exercise, you must implement the function `multiply_matrices`, which takes two matrices as arguments `matrix1` and `matrix2` and computes their multiplication. Here, we assume that the matrices can have arbitrary sizes. 

The result of the multiplication should be assigned to the variable `result_matrix`. Note again that the matrices are nested lists. 

```python
def multiply_matrices(matrix1, matrix2):


    rows1 = len(matrix1)
    cols1 = len(matrix1[0]) if rows1 > 0 else 0
    rows2 = len(matrix2)
    cols2 = len(matrix2[0]) if rows2 > 0 else 0

    # Check if multiplication is possible
    if cols1 != rows2:
        return ValueError("Error: The number of columns in the first matrix must equal the number of rows in the second matrix.")

    ### BEGIN SOLUTION
    result_matrix = []
    for i in range(rows1):
        temp = []
        for j in range(cols2):
          temp.append(0)
        result_matrix.append(temp)
        
    # or we use list comprehension
    # result_matrix = [[0 for _ in range(cols2)] for _ in range(rows1)]

    for i in range(rows1):  
        for j in range(cols2): 
            for k in range(cols1): 
                result_matrix[i][j] += matrix1[i][k] * matrix2[k][j]
                
    ### END SOLUTION

    return result_matrix

```

```python
### BEGIN HIDDEN TESTS
A = [
    [1, 2, 3],
    [4, 5, 6]
]

B = [
    [7, 8],
    [9, 10],
    [11, 12]
]

C = multiply_matrices(A, B)
print(C)

assert C[0] == [58, 64]
assert C[1] == [139, 154]


D = [
    [1, 2],
    [3, 4]
]

E = [
    [5, 6],
    [7, 8]
]

F = multiply_matrices(D, E)

assert F[0] == [19, 22]
assert F[1] == [43, 50]


G = [
    [1, 2, 3],
    [4, 5, 6]
]

H = [
    [7, 8]
]

I = multiply_matrices(G, H)
assert isinstance(I, ValueError)
### END HIDDEN TESTS
```

**b)**

What is the complexity (Big-O Notation) of the `multiply_matrices()` function?

1.) $O(n^2)$

2.) $O(m^2)$

3.) $O(m \cdot n)$

4.) $O(2^{m+n})$

5.) $O(n^3)$

```python
# Right answer is 5 O(n**3)
### BEGIN HIDDEN TESTS
### END HIDDEN TESTS
```