<img align="right" width=160 src="https://upload.wikimedia.org/wikipedia/commons/thumb/3/3a/Logo_KIT.svg/1920px-Logo_KIT.svg.png"/>

# IDS Assignments

**Allgemeine Hinweise**:

* Im Rahmen der Tutorien haben Sie 60 min Zeit für die Bearbeitung der Aufgaben. Bitte wenden Sie sich umgehend an Ihren Tutor, sobald die Bearbeitung abgeschlossen ist.
* Bitte überprüfen Sie, dass die enthaltenen Tests fehlerfrei ausgeführt werden können.
* Bei Unklarheiten, Fragen oder Problemen, wenden Sie sich bitte **unmittelbar** an Ihren Tutor (nicht erst zum Ende des Tutoriums) und holen Sie sich Unterstützung ein.
* Das Arbeiten in 2-er Teams ist erlaubt und erwünscht! :)


**General Notes**:

* During the tutorials, you have 60 minutes to work on the tasks. Please contact your tutor immediately once you have completed the work.
* Make sure that the included tests can run without errors.
* If anything is unclear, or if you have questions or problems, please contact your tutor immediately (not just at the end of the tutorial) and ask for support.
* Working in pairs (2-person teams) is allowed and encouraged! :)

# Wiederholungen und Auffrischungen

### Aufgabe 1: If-Bedingungen

**a)**

Schreiben Sie die verschachtelten if-Bedingungen so um, dass nur noch eine einzelne if-Bedingung verwendet wird.

```python
temperature = 23.99999
day = "Mo"
working_days = ["Mo", "Di", "Mi", "Do", "Fr"]

if temperature >= 24:
    if day in working_days:
        print("Hitzefrei!")
else:
    print("Work hard!")
```

<!-- grade_id: cell-291b2405ee25ba0e, points: 0, grade: false, solution: true -->
```python
temperature = 23.99999
day = "Mo"
working_days = ["Mo", "Di", "Mi", "Do", "Fr"]

# YOUR CODE HERE

# END OF YOUR CODE
```

**b)**

Schreiben Sie die folgende if-Bedingung so um, dass keine `and`- oder `or`-Verknüpfungen mehr verwendet werden, der Code aber funktional identisch bleibt. 
Stattdessen sollen mehrere if-Bedingungen verwendet werden.

```python
working_days = ["Mo", "Di", "Mi", "Do", "Fr"]
is_public_holiday = False
on_vacation = False
day = "Mi"

if day in working_days and not (is_public_holiday or on_vacation):
    print("Work hard!")
else:
    print("Relax and enjoy life!")
```

<!-- grade_id: cell-09cfba54e54f5fd3, points: 0, grade: false, solution: true -->
```python
working_days = ["Mo", "Di", "Mi", "Do", "Fr"]
is_public_holiday = False
on_vacation = False
day = "Mi"

# YOUR CODE HERE

# END OF YOUR CODE
```

### Aufgabe 2: Typumwandlung

Ein Roboter ist mit einem Entfernungssensor ausgestattet, der Messwerte als Zeichenkette (String) im Format `XX.XX m` liefert (z.B. `12.34 m`).
Der folgende Python-Code soll diese Messwerte in Gleitkommazahlen (Float) umwandeln, damit sie für Berechnungen verwendet werden können.
Allerdings fehlt eine entscheidende Zeile im Code. Ergänzen Sie die Zeile im Codeblock.

*Hinweis:* Die Stelle, an der Code fehlt, ist wie folgt unterstrichen ________. Sie dürfen eine oder zwei Zeilen verwenden.

```python
1 def convert_distance_to_float(distance_str):
2   _________________________
3   _________________________
4   return distance_in_meters
5
6 distance_str = "25.87 m"
7 distance_float = convert_distance_to_float(distance_str)
8 print(f"Gemessene Entfernung: {distance_float} Meter")
```

<!-- grade_id: cell-b2a34fbf6e4f0c82, points: 0, grade: false, solution: true -->
```python
def convert_distance_to_float(distance_str):
    # YOUR CODE HERE
    
    # END OF YOUR CODE
    return distance_in_meters

distance_str = "25.87 m"
distance_float = convert_distance_to_float(distance_str)
print(f"Gemessene Entfernung: {distance_float} Meter")
```

<!-- grade_id: cell-d5cae997a444d3fb, points: 0, grade: true, solution: false -->
```python
assert distance_float == 25.87
assert convert_distance_to_float("100.00 m") == 100
assert convert_distance_to_float("8.12 m") == 8.12
```

### Aufgabe 3: Objekte und Vererbung

Gegeben sind drei Klassen, die voneinander erben, sowie der Beginn eines zugehörigen Hauptprogramms (siehe Code-Zelle unten).
Aktuell wird von diesem Programm bereits die folgende Ausgabe erzeugt:

```python
Creating Jin's motorized vehicle.
Creating car.
Creating Aya's motorized vehicle.
Creating motorbike.
Creating Tom's motorized vehicle.
Creating motorbike.
```

Ihre Aufgabe besteht darin, zusätzliche Zeilen zu implementieren, welche exakt die zusätzliche nachfolgende Ausgabe erzeugen. 
Dabei dürfen keine neuen Variablen deklariert werden und die bestehende Liste vehicles soll nach Abschluss des Programms alle erzeugten Instanzen eines MotorizedVehicle enthalten.

```python
Jane drives 10km in Jin's car.
Kilometer counter increased to 10km.
Tom sells the vehicle to Raj.
Raj rides 1km on Raj's motorbike.
Kilometer counter increased to 3913km.
Kilometer counter increased to 3914km.
Creating Tom's motorized vehicle.
Creating car.
Tom drives 1000km in Tom's car.
Kilometer counter increased to 1000km.
```

<!-- grade_id: cell-ec1ce6e641563768, points: 0, grade: false, solution: true -->
```python
class MotorizedVehicle:
    def __init__(self, seats, kW, kilometers, owner):
        self.seats = seats
        self.kW = kW
        self.kilometers = kilometers
        self.owner = owner
        print(f"Creating {self.owner}'s motorized vehicle.")

    def sell(self, new_owner):
        print(f"{self.owner} sells the vehicle to {new_owner}.")
        self.owner = new_owner

    def is_driven(self, distance):
        self.kilometers += distance
        print(f"Kilometer counter increased to {self.kilometers}km.")


class Car(MotorizedVehicle):
    def __init__(self, kW, kilometers, owner):
        super().__init__(5, kW, kilometers, owner)
        print(f"Creating car.")

    def drive(self, distance, driver):
        print(f"{driver} drives {distance}km in {self.owner}'s car.")
        self.is_driven(distance)


class Motorbike(MotorizedVehicle):
    def __init__(self, kW, kilometers, owner):
        super().__init__(2, kW, kilometers, owner)
        print(f"Creating motorbike.")

    def ride(self, distance, rider):
        print(f"{rider} rides {distance}km on {self.owner}'s motorbike.")
        self.is_driven(distance)

vehicles = [Car(100, 0, "Jin"),
            Motorbike(40, 2743, "Aya"),
            Motorbike(50, 3912, "Tom")]

# YOUR CODE HERE

# END OF YOUR CODE
```

<!-- grade_id: cell-b496652f0d6eccf6, points: 0, grade: true, solution: false -->
```python
assert len(vehicles) == 4
assert vehicles[0].owner == "Jin"
assert vehicles[1].owner == "Aya"
assert vehicles[2].owner == "Raj"
assert vehicles[3].owner == "Tom"
```