## Dateien und Daten

### Aufgabe 1

Im Python Crash-Kurs haben wir bereits die Grundlagen im Umgang mit Dateien in Python gelernt.

In dieser Aufgabe werden wir tabellarische Daten in eine Datei schreiben und die Datei sukzessive erweitern.
Zuletzt werden wir die gesamten Daten einlesen und weiterverwenden. 

#### A) In eine Datei schreiben

Die Daten einer Produktionsanlage sollen in einer Datei `production_data.csv` gespeichert werden.

Darin soll jeden Tag die Menge produzierter Teile (`amount`), sowie die Menge an Ausschussteilen (`scrap`) erfasst werden.
In dieser Aufgabe initialisieren wir die Datei und schreiben zunächst eine entsprechende Header-Zeile hinein.

Öffnen Sie die Datei im Schreib-Modus (`"w"`) und schreiben Sie die Textzeile `"day;amount;scrap"` in die Datei (am Ende der Zeile soll ein Zeilenumbruch mit `\n` erfolgen).

<!-- grade_id: cell-2cf4b1cdd6ffcd33, points: 0, grade: false, solution: true -->
```python
from pathlib import Path

notebook_dir = Path().resolve()
datafile_path = notebook_dir / "production_data.csv"

### BEGIN SOLUTION
with open(datafile_path, "w") as file:
    file.write("day;amount;scrap\n")
    file.close()
### END SOLUTION
```

<!-- grade_id: cell-ee4f7b4506a489bf, points: 0, grade: true, solution: false -->
```python
assert Path.is_file(datafile_path)
```

#### B) Daten zu einer existierenden Datei hinzufügen

Nun soll eine Funktion implementiert werden, die der Datei die Daten eines Produktionstags hinzufügt.
Die Funktion `append_production_data(day, amount, scrap)` nimmt also als Argumente einen Tag (Index), die an diesem Tag produzierte Menge, sowie die Ausschussmenge des Tages als Argumente entgegen.
Diese Daten sollen dann entsprechend dem Format der Kopfzeile (siehe Aufgabenteil **A)**) in die Datei eingefügt werden.

*Hinweis:* Um die Datei dabei nicht vollständig zu überschreiben muss die Datei im Modus entsprechenden Modus `a` (append) geöffnet werden.

<!-- grade_id: cell-9b287e8b4ed2d40a, points: 0, grade: false, solution: true -->
```python
def append_production_data(day, amount, scrap):
### BEGIN SOLUTION
    with open(datafile_path, "a") as file:
        file.write(f"{day};{amount};{scrap}\n")
        file.close()
### END SOLUTION

append_production_data(1, 1400, 19)
append_production_data(2, 1400, 21)
append_production_data(3, 1600, 22)
append_production_data(4, 1600, 24)
append_production_data(5, 1600, 20)
append_production_data(6, 1600, 36)
append_production_data(7, 1600, 39)
append_production_data(8, 1400, 19)
append_production_data(9, 1400, 22)
append_production_data(10, 1600, 22)
append_production_data(11, 1600, 26)
append_production_data(12, 1600, 24)
append_production_data(13, 1600, 32)
append_production_data(14, 1600, 37)
```

#### C) Daten aus einer csv Datei auslesen

Die geschriebene Datei entspricht dem sogenannten csv-Format ("comma separated values").
Mithilfe des Python-Standard-Moduls `csv` können Daten in diesem Format etwas bequemer eingelesen werden.

Im folgenden Code-Block ist ein entsprechendes Einlesen der Daten bereits implementiert.
Die Variable `line` enthält dabei in jedem Schleifendurchlauf eine Liste der Werte einer Zeile (Datentyp `str`).
Konvertieren Sie die eingelesenen Werte in numerische Werte und fügen Sie diese sinngemäß den beiden Listen `amounts` und `scraps` hinzu.

Berechnen Sie anschließend die Summer der produzierten Teile, sowie die Summe der Ausschuss-Teile über den gesamten enthaltenen Zeitraum (14 Tage).

*Hinweise*:
  * Bitte beachten Sie, dass zunächst die Kopfzeile der Datei eingelesen wird und diese keine Daten enthält.
  * Sie können die zuvor geschriebene Datei manuell öffnen und anschauen, um ein besseres Verständnis für die Datei und die Programmabläufe zu gewinnen.

<!-- grade_id: cell-27a47ffbfbbfb866, points: 0, grade: false, solution: true -->
```python
import csv

amounts = []
scraps = []

with open(datafile_path, "r") as file:
    csv_data_reader = csv.reader(file, delimiter=";")
    for line in csv_data_reader:
        ### Convert numeric data and append lists amounts and scraps
        ### BEGIN SOLUTION
        if line[0] == "day":
            continue
        amounts.append(float(line[1]))
        scraps.append(float(line[2]))
        ### END SOLUTION
        
amount_sum = None
scraps_sum = None
### Compute sums
### BEGIN SOLUTION
amount_sum = sum(amounts)
scraps_sum = sum(scraps)
### END SOLUTION
```

<!-- grade_id: cell-482e19064dca6a63, points: 0, grade: true, solution: false -->
```python
assert len(amounts) == 14
assert amounts[0] == 1400
assert len(scraps) == 14
assert amount_sum == 21600
assert scraps_sum == 363
```