## Objekte, Klassen und Instanzen

***
### Aufgabe 1

Im Unterordner `programme` zu dieser Übungseinheit liegt das Modul `coffee_machine.py`, in welchem die Klasse `CoffeeMachine` implementiert ist.

Der bereits vorhandene Code fügt den Unterordner deinem Interpreter-Pfad hinzu.

Importiere die Klasse `CoffeeMachine`.
Erzeuge eine Instanz mit Namen `my_coffee_machine`.
Nutze anschließend die Methoden der Instanz, um Kaffee und Wasser nachzufüllen und einen Kaffee zu brühen!

<!-- grade_id: cell-e647e7bfbca1e9b8, points: 0, grade: false, solution: true -->
```python
from pathlib import Path
import sys

notebook_dir = str(Path().resolve())
import_dir = notebook_dir + "/programme"

sys.path.insert(0, import_dir)

### BEGIN SOLUTION
from coffee_machine import CoffeeMachine

my_coffee_machine = CoffeeMachine()

my_coffee_machine.add_coffee(100)
my_coffee_machine.add_water(500)
my_coffee_machine.brew_coffee()
### END SOLUTION
```

<!-- grade_id: cell-30edd9af18414682, points: 1, grade: true, solution: false -->
```python
assert isinstance(my_coffee_machine, CoffeeMachine)
assert my_coffee_machine.brew_counter > 0
```

***
### Aufgabe 2

Die Klasse `CoffeeMachine` hat bereits ein Klassenvariable `waste_storage_size` sowie eine Instanz-Variable `waste_storage`, die aktuell im Code der Klasse nicht verwendet werden.

Ändere den Code der Klasse in `coffee_machine.py` wie folgt ab:

* Bei jedem Brühvorgang (`brew_coffee()`) wird der verwendete Kaffee dem Abfallbehälter (`waste_storage`) der Kaffeemaschine hinzugefügt. (Gehe dabei davon aus, dass 1 Einheit verwendeter Kaffee zu 1 Einheit Abfall führt.)
* Ein Brühvorgang kann nur erfolgen, wenn der Abfallbehälter noch nicht voll ist (`waste_storage` und `waste_storage_size` vergleichen). Ändere hierzu die Methode `is_ready_to_brew` entsprechend ab.
* Ergänze eine Methode `empty_waste_storage`, um den Abfallbehälter der Kaffeemaschine zu leeren.

Importiere anschließend die Klasse `CoffeeMachine` erneut und erzeuge eine Instanz mit Namen `my_cofee_machine`.
Brühe damit Kaffee bis der Abfallbehälter der Maschine voll ist und kein Brühvorgang mehr möglich ist. (Bei Bedarf Wasser und Kaffee nachfüllen, bei aktueller Einstellung der Maschine müssen dazu 10 Tassen Kaffee gebrüht werden.)

*Hinweis:* Nach Ändern des Codes in `coffee_machine.py` muss ein erneuter Import durchgeführt werden. Dazu den Interpreter neu starten oder das geladene Modul explizit aus dem laufenden Interpreter löschen (siehe bereits vorhandener Code-Schnippsel).

<!-- grade_id: cell-dd469a73bfc73bfe, points: 0, grade: false, solution: true -->
```python
from pathlib import Path
import sys

notebook_dir = str(Path().resolve())
import_dir = notebook_dir + "/programme"

sys.path.insert(0, import_dir)

# Entferne ggf. das bereits importierte Modul, um einen Reload zu erzwingen und Änderungen im Modul zu berücksichtigen.
if "coffee_machine" in sys.modules.keys():
    del sys.modules["coffee_machine"]

### BEGIN SOLUTION
from coffee_machine import CoffeeMachine

my_coffee_machine = CoffeeMachine()
for i in range(10):
    my_coffee_machine.add_coffee(12)
    my_coffee_machine.add_water(150)
    my_coffee_machine.brew_coffee()
### END SOLUTION
```

<!-- grade_id: cell-e1f0f971caa63fc6, points: 1, grade: true, solution: false -->
```python
assert isinstance(my_coffee_machine, CoffeeMachine)
assert my_coffee_machine.waste_storage > my_coffee_machine.waste_storage_size
my_coffee_machine.waste_storage = 0
my_coffee_machine.add_water(150)
my_coffee_machine.add_coffee(12)
assert my_coffee_machine.is_ready_to_brew()
```