## Algorithmen

### Aufgabe 1: Insertion Sort

In dieser Übung werden wir einen Sortieralgorithmus namens Insertion Sort implementieren.

Dieser Algorithmus sortiert eine Liste von Zahlen (`my_list`) von der kleinsten zur größten. Dazu wird jede Zahl aus der ursprünglichen Liste (`my_list`) genommen und an der richtigen Stelle in eine neue Liste (`sorted_list`) eingefügt, die für die Sortierung aufgebaut wird.

Schritte:

- Definieren Sie eine neue leere Liste namens `sorted_list`, die zurückgegeben wird, wenn die Funktion beendet ist.

- Schleife durch jedes Element in `my_list` (mit einer for-Schleife). Führen Sie in jeder Schleifeniteration Folgendes aus:
    - Initialisiere eine Indexvariable namens „i“ und setze ihren Wert auf 0.
    - Dann wird `i` immer um eins erhöht, um die richtige Position für das aktuelle Element aus `my_list` zu finden, indem eine while-Schleife mit folgenden Bedingungen verwendet wird:
        - `i` ist kleiner als die Länge von `sorted_list`
        - und das aktuelle Element ist größer als oder gleich dem aktuellen Element von `sorted_list` (`sorted_list[i]`)
          
          
    - Sobald die while-Schleife beendet ist, wird das aktuelle Element aus `my_list` in `sorted_list` am `i`-ten Index eingefügt

<!-- grade_id: cell-5dadede7ab35fe0e, points: 0, grade: false, solution: true -->
```python
def insertion_sort(my_list):
    ### BEGIN SOLUTION

    sorted_list = []

    for element in my_list:
        i = 0
        while i < len(sorted_list) and element >= sorted_list[i]:
            i += 1
        sorted_list.insert(i, element)
    ### END SOLUTION
    return sorted_list
```

<!-- grade_id: cell-d9c082c9ccd289e3, points: 2, grade: true, solution: false -->
```python
my_list = [2,3,1,5,4,3,2]
sorted_list = insertion_sort(my_list)
assert len(sorted_list) == 7
for i in range(len(sorted_list)-1):
    assert sorted_list[i] <= sorted_list[i+1]

### BEGIN HIDDEN TESTS
import numpy as np
np.random.seed(89)
liste_r = list([np.random.randint(500) for i in range(100)])
sortierte_liste_r = insertion_sort(liste_r)
assert len(sortierte_liste_r) == 100
for i in range(len(sortierte_liste_r)-1):
    assert sortierte_liste_r[i] <= sortierte_liste_r[i+1]
    

### END HIDDEN TESTS
```

### Aufgabe 2: Rotation in $\mathbb{R}^3$

Diese Übung befasst sich mit der Rotation von Vektoren in $\mathbb{R}^3$.

Im Folgenden werden verschiedene Algorithmen zur Berechnung von Vektorrotationen um bestimmte Achsen und Winkel vorgestellt.

Die Rotationsmatrizen $R \in \mathbb{R}^{3x3}$ für Rotationen um die drei Koordinatenachsen in $\mathbb{R}^3$ können wie folgt berechnet werden:
* Drehung um die x-Achse:

${\displaystyle R_{x}(\alpha )={\begin{pmatrix}1&0&0\\0&\cos \alpha &-\sin \alpha \\0&\sin \alpha &\cos \alpha \end{pmatrix}}}$

* Drehung um die y-Achse:

${\displaystyle R_{y}(\alpha )={\begin{pmatrix}\cos \alpha &0&\sin \alpha \\0&1&0\\-\sin \alpha &0&\cos \alpha \end{pmatrix}}}$

* Drehung um die z-Achse:

${\displaystyle R_{z}(\alpha )={\begin{pmatrix}\cos \alpha &-\sin \alpha &0\\\sin \alpha &\cos \alpha &0\\0&0&1\end{pmatrix}}}$



Sei $v \in \mathbb{R}^3$ ein 3-dimensionaler Vektor.
Das Matrix-Vektorprodukt $v_r = R_{x}(\alpha) \cdot v \in \mathbb{R}^3$, entspricht dem um den Winkel $\alpha$ gedrehten Vektor $v$ um die Koordinatenachse $x=\begin{pmatrix}1\\0\\0\end{pmatrix}$. (Analoges gilt für die y- und z-Achse.)


Implementieren Sie die Funktion `rotate_matrix()`, die die entsprechende Rotationsmatrix für eine Drehung um die x-, y- oder z-Achse des $\mathbb{R}^3$ um einen bestimmten Winkel `angular_rad` berechnet. Die Funktion nimmt ein ganzzahliges Argument `achse` entgegen, das dem Index der Koordinatenachse entspricht, um die die Drehung erfolgen soll (0: x-Achse, 1: y-Achse, 2: z-Achse).

Hinweis: Die Matrix ist

<!-- grade_id: cell-7afaae70fbdfc508, points: 0, grade: false, solution: true -->
```python
from math import cos, sin

def rotate_matrix(axis, angular_rad):
    if axis not in [0, 1, 2]:
        raise RuntimeError("Undefined rotation axis! (Choose an integer 0: x-axis, 1: y-axis, 2: z-axis)")

    r_matrix = [[0, 0, 0], [0, 0, 0], [0, 0, 0]]

    ### BEGIN SOLUTION
    
    cos_r = cos(angular_rad)
    sin_r = sin(angular_rad)

    r_matrix[axis][axis] = 1
    
    if axis == 0:  # X-axis rotation
        r_matrix[1][1] = cos_r
        r_matrix[1][2] = -sin_r
        r_matrix[2][1] = sin_r
        r_matrix[2][2] = cos_r
    elif axis == 1:  # Y-axis rotation
        r_matrix[0][0] = cos_r
        r_matrix[0][2] = sin_r
        r_matrix[2][0] = -sin_r
        r_matrix[2][2] = cos_r
    elif axis == 2:  # Z-axis rotation
        r_matrix[0][0] = cos_r
        r_matrix[0][1] = -sin_r
        r_matrix[1][0] = sin_r
        r_matrix[1][1] = cos_r
        
    ### END SOLUTION
    return r_matrix
```

<!-- grade_id: cell-ebbcf3446ce5220f, points: 2, grade: true, solution: false -->
```python
 import math

# Test unit matrix with rotation angle 0
for i in range(0, 3):
    r_matrix = rotate_matrix(i, 0)
    assert math.isclose(r_matrix[0][0], 1, abs_tol=1e-6) and math.isclose(r_matrix[1][1], 1, abs_tol=1e-6) and math.isclose(r_matrix[2][2], 1, abs_tol=1e-6)
    assert math.isclose(sum(r_matrix[0]), 1, abs_tol=1e-6) and math.isclose(sum(r_matrix[1]), 1, abs_tol=1e-6) and math.isclose(sum(r_matrix[2]), 1, abs_tol=1e-6)
    
# Test Case: Rotation 90°
r_matrix = rotate_matrix(2, 0.5 * math.pi)
assert math.isclose(sum(r_matrix[0]), -1, abs_tol=1e-6)
assert math.isclose(sum(r_matrix[1]), 1, abs_tol=1e-6)
assert math.isclose(sum(r_matrix[2]), 1, abs_tol=1e-6)
assert math.isclose(r_matrix[0][1], -1, abs_tol=1e-6)
assert math.isclose(r_matrix[1][0], 1, abs_tol=1e-6)
assert math.isclose(r_matrix[2][2], 1, abs_tol=1e-6)

### BEGIN HIDDEN TESTS
import numpy as np
v = np.array([0.5,0.25,1.5])
r = np.array(rotate_matrix(1, 72 / 180 * math.pi))
res = np.matmul(r, v)
assert np.abs(res[0] - 1.581093) < 1e-5
assert np.abs(res[1] - 0.25) < 1e-5
assert np.abs(res[2] - (-0.012003)) < 1e-5

r = np.array(rotate_matrix(0, -30 / 180 * math.pi))
res_2 = np.matmul(r, res)
assert np.abs(res_2[0] - 1.581093) < 1e-5
assert np.abs(res_2[1] - 0.210505) < 1e-5
assert np.abs(res_2[2] - (-0.135395)) < 1e-5

r_1 = np.array(rotate_matrix(2, 134 / 180 * math.pi))
r_2 = np.array(rotate_matrix(2, -134 / 180 * math.pi))
res_3 = np.matmul(r_1, r_2)

assert np.linalg.norm(res_3 - np.identity(3)) < 1e-5
### END HIDDEN TESTS
```

### Aufgabe 3: Matrix-Vektor-Multiplikation


**a)**

In einem Teil der Übung soll das Matrix-Vektorprodukt in $\mathbb{R}^n$ implementiert werden.

Implementieren Sie für eine Matrix `matrix` aus $\mathbb{R}^{m x n}$ und einen Vektor `vector` aus $\mathbb{R}^{n}$ unter 
eine Funktion `matrix_vector_product(matrix, vector)`, die das Ergebnis der Matrix-Vektor-Multiplikation `result_vector` $\in \mathbb{R}^m$ berechnet und zurückgibt.

Matrizen werden wieder als verschachtelte Listen dargestellt (für $\mathbb{R}^{m x n}$: eine Liste mit m Einträgen, jeder Eintrag ist eine Liste mit n numerischen Einträgen) und Vektoren als Listen.


Das Matrix-Vektor-Produkt wird wie folgt berechnet:
Gegeben sei eine $m \times n$-Matrix $A$ und ein $n \times 1$-Spaltenvektor $\mathbf{x}$:

$A = \begin{pmatrix} a_{11} & a_{12} & \cdots & a_{1n} \\ a_{21} & a_{22} & \cdots & a_{2n} \\ \vdots & \ddots & \vdots \\ a_{m1} & a_{m2} & \cdots & a_{mn} \end{pmatrix}$

$\mathbf{x} = \begin{pmatrix} x_1 \\\ x_2 \\\ \vdots \\ x_n \end{pmatrix}$

Das Produkt $A\mathbf{x}$ ist ein $m \times 1$ Spaltenvektor $\mathbf{b}$:

$\mathbf{b} = A\mathbf{x} = \begin{pmatrix} b_1 \\\ b_2 \\\ \vdots \\\ b_m \end{pmatrix}$

wobei jedes Element $b_i$ des resultierenden Vektors als das Punktprodukt der $i$-ten Zeile der Matrix $A$ und des Vektors $\mathbf{x}$ berechnet wird:

$b_i = a_{i1}x_1 + a_{i2}x_2 + \cdots + a_{in}x_n = \sum_{j=1}^{n} a_{ij}x_j$

<!-- grade_id: cell-51277670c0fb5bc5, points: 0, grade: false, solution: true -->
```python
def matrix_vector_product(matrix, vector):
    ### BEGIN SOLUTION
    result_vector = []
    for i in range(len(matrix)):
        entry = 0
        for j in range(len(matrix[i])):
            entry += matrix[i][j] * vector[j]
        result_vector.append(entry)
    ### END SOLUTION
    return result_vector
```

<!-- grade_id: cell-61f842d333badb70, points: 2, grade: true, solution: false -->
```python
assert matrix_vector_product([[1]], [1]) == [1]
assert len(matrix_vector_product([[1,0,0]],[1,1,1])) == 1
assert len(matrix_vector_product([[1,0,0],[0,1,0],[0,0,1]],[1,1,1])) == 3
assert len(matrix_vector_product([[1,0,0],[0,1,0],[0,0,1],[1,0,0],[0,1,0],[0,0,1]],[1,1,1])) == 6

### BEGIN HIDDEN TESTS
vector = [2.37, -3.41, -0.01]
matrix = [[1.01, -0.29, 12.91], [3.98, 2.77, -5.81]]
erg = matrix_vector_product(matrix, vector)
assert len(erg) == 2
assert abs(erg[0] - 3.2535) < 1e-5
assert abs(erg[1] - 0.045) < 1e-5
### END HIDDEN TESTS
```

**b)**

Wie hoch ist die Komplexität (Big-O-Notation) der  `matrix_vector_product()`-Funktion für Matrizen der Dimension $\mathbb{R}^{m x n}$?

1.) $O(n^2)$

2.) $O(m^2)$

3.) $O(m \cdot n)$

4.) $O(2^{m+n})$

5.) $O(1)$

<!-- grade_id: cell-9654eba846efd188, points: 0, grade: false, solution: true -->
```python
# Right answer is 3 O(m . n)
### BEGIN HIDDEN TESTS
### END HIDDEN TESTS
```

### Aufgabe 4: Matrixmultiplikation

**a)**

In dieser Übung müssen Sie die Funktion `multiply_matrices` implementieren, die zwei Matrizen als Argumente `matrix1` und `matrix2` annimmt und deren Multiplikation berechnet. Wir nehmen hier an, dass die Matrizen beliebige Größen haben können. 

Das Ergebnis der Multiplikation soll der Variablen `result_matrix` zugewiesen werden. Beachten Sie wiederum, dass die Matrizen verschachtelte Listen sind. 

<!-- grade_id: cell-f2852a9ce73a06b5, points: 0, grade: false, solution: true -->
```python
def multiply_matrices(matrix1, matrix2):


    rows1 = len(matrix1)
    cols1 = len(matrix1[0]) if rows1 > 0 else 0
    rows2 = len(matrix2)
    cols2 = len(matrix2[0]) if rows2 > 0 else 0

    # Check if multiplication is possible
    if cols1 != rows2:
        return ValueError("Error: The number of columns in the first matrix must equal the number of rows in the second matrix.")

    ### BEGIN SOLUTION
    result_matrix = []
    for i in range(rows1):
        temp = []
        for j in range(cols2):
          temp.append(0)
        result_matrix.append(temp)
        
    # or we use list comprehension
    # result_matrix = [[0 for _ in range(cols2)] for _ in range(rows1)]

    for i in range(rows1):  
        for j in range(cols2): 
            for k in range(cols1): 
                result_matrix[i][j] += matrix1[i][k] * matrix2[k][j]
                
    ### END SOLUTION

    return result_matrix

```

<!-- grade_id: cell-f294bed37885fb57, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
A = [
    [1, 2, 3],
    [4, 5, 6]
]

B = [
    [7, 8],
    [9, 10],
    [11, 12]
]

C = multiply_matrices(A, B)
print(C)

assert C[0] == [58, 64]
assert C[1] == [139, 154]


D = [
    [1, 2],
    [3, 4]
]

E = [
    [5, 6],
    [7, 8]
]

F = multiply_matrices(D, E)

assert F[0] == [19, 22]
assert F[1] == [43, 50]


G = [
    [1, 2, 3],
    [4, 5, 6]
]

H = [
    [7, 8]
]

I = multiply_matrices(G, H)
assert isinstance(I, ValueError)
### END HIDDEN TESTS
```

**b)**

Wie hoch ist die Komplexität (Big-O Notation) der Funktion `multiply_matrices()`?

1.) $O(n^2)$

2.) $O(m^2)$

3.) $O(m \cdot n)$

4.) $O(2^{m+n})$

5.) $O(n^3)$

```python
# Right answer is 5 O(n**3)
### BEGIN HIDDEN TESTS
### END HIDDEN TESTS
```