# Error Messages and Stack Traces - Part II

***
### Exercise 1

This task involves calculating the investment returns of a deposit with a customer-dependent interest rate.
The base interest rate is 5%, and depending on the customer group, an additional bonus interest rate is granted:

| Customer Group | Bonus Interest Rate |
| --- | --- |
| Private |	0% |
| Enterprise |	1% |
| Charitable |	2% |

The returns should be calculated based on two deposit periods.
The deposit from the first period earns compound interest (interest on interest), while the deposit from the second period earns simple interest only.

After that, an example calculation should be made for an investment of €1,000 in period 1 and €500 in period 2 for all three customer groups, and the respective returns should be displayed.


```python
interest_rate = 1.05

interest_correction_by_customer = {
    "private": 0.0
    "enterprise": 0.01
    "charitable": 0.02
}

def revenue(cap_1, cap_1, customer_type)
     interest_rate = interest_rate + interes_correction_by_customer[customer_type]
    result = interest_rate * (interest_rate * cap_1 + cap_2)

cap_1 = 1000
cap_2 = 500

for customer_type in interest_correction_by_customer.items():
return_on_invest = revenue(cap_1, cap_2, customer_type)
print(return_on_invest)
```

However, a few mistakes have crept in here...
Correct all the errors to solve the task correctly. An iterative approach makes sense here: run the code – analyze the error message – fix the error.

<!-- grade_id: cell-cda2160980c1f537, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
interest_rate = 1.05

interest_correction_by_customer = {
    "private": 0.0,
    "enterprise": 0.01,
    "charitable": 0.02
}

def revenue(cap_1, cap_2, customer_type):
    interest_rate_customer = interest_rate + interest_correction_by_customer[customer_type]
    result = interest_rate_customer * (interest_rate_customer * cap_1 + cap_2)
    return result

cap_1 = 1000
cap_2 = 500

for customer_type in interest_correction_by_customer:
    return_on_invest = revenue(cap_1, cap_2, customer_type)
    print(return_on_invest)
### END SOLUTION
```

<!-- grade_id: cell-43f7b4b587c11f7f, points: 0, grade: true, solution: false -->
```python
assert revenue(cap_1, cap_2, "private") == 1627.5
### BEGIN HIDDEN TESTS
assert revenue
assert revenue(cap_1, cap_2, "private") == 1627.5
assert revenue(-50, 100, "charitable") == 49.755
assert revenue(0, 0, "enterprise") == 0

try:
    revenue(0, 0, "")
except:
    pass
else:
    raise RuntimeError("Calling revenue with invalid customer type should raise an Exception!")
    
assert "private" in interest_correction_by_customer.keys()
assert len(interest_correction_by_customer) == 3
### END HIDDEN TESTS
```