## Programming

### Exercise 1: Multiple Choice (3 Punkte)

In this task, nothing needs to be implemented. General questions are asked, each with several answer options.
From these options, the correct one must be selected.

To submit your answer, you must use the provided code cell and edit it exactly as instructed.
Since grading is automated, please pay close attention to following the given instructions precisely.

**a)**

Which of the following statements about variable names in Python is correct?
(Exactly one answer is correct)

1) ASCII characters can be used freely in variable names.
2) Variable names may begin with the character `_`.
3) Variable names must always begin with a letter.
4) Only letters and numbers are allowed in variable names.

Please answer this question by creating a variable named `answer_1a` in the code cell below and assigning it the numeric value of your chosen answer.

(Specifically: If you believe statement *1) ASCII characters can be used freely in variable names.* is correct, then write `answer_1a = 1`.)

<!-- grade_id: cell-82af01e7d9fbe68a, points: 0, grade: false, solution: true -->
```python
answer_1a = None
### BEGIN SOLUTION
answer_1a = 2
### END SOLUTION
```

<!-- grade_id: cell-b8c62f5fc0df8d61, points: 1, grade: true, solution: false -->
```python
assert answer_1a
assert answer_1a in [1, 2, 3, 4]
### BEGIN HIDDEN TESTS
assert answer_1a == 2
### END HIDDEN TESTS
```

**b)**

What is the correct file extension for a Python module? (Exactly one answer is correct)

1) `.py`
2) `.ipynb`
3) `.ipy`
4) `.python`

Please answer this question by creating a variable named `answer_1b` in the code cell below and assigning it the numeric value of your chosen answer.

<!-- grade_id: cell-520e95d7cd94dc95, points: 0, grade: false, solution: true -->
```python
answer_1b = None
### BEGIN SOLUTION
answer_1b = 1
### END SOLUTION
```

<!-- grade_id: cell-26610294288486e9, points: 1, grade: true, solution: false -->
```python
assert answer_1b
assert answer_1b in [1, 2, 3, 4]
### BEGIN HIDDEN TESTS
assert answer_1b == 1
### END HIDDEN TESTS
```

**c)**

How can the programming language Python be classified?
(Exactly one answer is correct)

1) Interpreted language, strongly typed.
2) Interpreted language, dynamically typed.
3) Compiled language, strongly typed.
4) Compiled language, dynamically typed.

Please answer this question by creating a variable named `answer_1c` in the code cell below and assigning it the numeric value of your chosen answer.

<!-- grade_id: cell-594036f80265e49f, points: 0, grade: false, solution: true -->
```python
answer_1c = None
### BEGIN SOLUTION
answer_1c = 2
### END SOLUTION
```

<!-- grade_id: cell-ca6574137b1fe825, points: 1, grade: true, solution: false -->
```python
assert answer_1c
assert answer_1c in [1, 2, 3, 4]
### BEGIN HIDDEN TESTS
assert answer_1c == 2
### END HIDDEN TESTS
```

### Exercise 2: Variables & Assignments (2 Punkte)
 **a)**
 
Create the following variables:
* Variable `a` of type `int`  	with value `5`
* Variable `b` of type `float`	with value `2.5`
* Variable `c` of type `bool` 	with value `True`
* Variable `d` of type `str`  	with value `Hallo IDS`

<!-- grade_id: cell-b85218bec9869cc7, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
a = 5
b = 2.5
c = True
d = "Hallo IDS"
### END SOLUTION
```

<!-- grade_id: cell-2ad53ccf36211ae9, points: 1, grade: true, solution: false -->
```python
assert a
assert b
assert c
assert d
### BEGIN HIDDEN TESTS
assert a == 5
assert b == 2.5
assert c == True
assert d == "Hallo IDS"
### END HIDDEN TESTS
```

**b)**

Create the following variables:
* Variable `value_1` with value `12.3`
* Variable `value_2` with value `4.8`

Calculate:
* The sum of both variables and store the result in a new variable `add`
* The difference of both variables and store the result in a new variable `diff`
* The product of both variables and store the result in a new variable `prod`
* The quotient of both variables and store the result in a new variable `quot`

<!-- grade_id: cell-ce001117fcb94861, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
value_1 = 12.3
value_2 = 4.8

add = value_1 + value_2
diff = value_2 - value_1
prod = value_1 * value_2
quot = value_1 / value_2
value_2 / value_1
### END SOLUTION
```

<!-- grade_id: cell-0c0581f4bd2b8580, points: 1, grade: true, solution: false -->
```python
assert add
assert diff
assert prod
assert quot
### BEGIN HIDDEN TESTS
assert abs(add - 17.1) < 1e-6
assert abs(diff - 7.5) < 1e-6 or abs(diff + 7.5) < 1e-6
assert abs(prod - 59.04) < 1e-6
assert abs(quot - 2.5625) < 1e-6 or abs(quot - 0.3902439) < 1e-6
### END HIDDEN TESTS
```

### Aufgabe 3: Mixed Topics (4 Punkte)


Write Python code in the following notebook cell that, for different combinations of side lengths `a` and `b` of a right triangle, calculates the length of the hypotenuse `c` and stores it in a dictionary.

**a)** First, implement a function `pythagoras` that takes the lengths of the two legs of a right triangle, `a` and `b`, as arguments and computes the corresponding hypotenuse length. This value should be returned using `return`.

**b)** The possible side lengths of the two legs `a` and `b` of the right triangle are from the set `{1, 2, 3, 4}`.
For each possible combination of side lengths `a` and `b`, use the Pythagorean theorem to calculate the corresponding hypotenuse length `c`.
Store all resulting values in a single nested dictionary called `hypothenuses`.
Afterward, the length of the hypotenuse for a triangle with sides `a` and `b` should be accessible using `hypothenuses[a][b]`.

The final code block already provided prints the entries of `hypothenuses` in a formatted way, so that each line shows one of the possible triangles.
This code should work without modification and output the triangle dimensions for the 10 possible combinations of side lengths from `sides`.

**Hinweise**: 
* The Pythagorean theorem for calculating the hypotenuse length `c` is: $c = \sqrt{a^2 + b^2}$
* Assuming that `a` and `b` are not distinguishable (i.e., the order doesn’t matter), you only need to compute each triangle permutation once.
(If `a=1` and `b=2` have already been computed, `a=2` and `b=1` do not need to be considered again. Thus, a total of 10 values must be computed.)

<!-- grade_id: cell-467bee93838670f0, points: 0, grade: false, solution: true -->
```python
from math import sqrt

def pythagoras(a, b):
# A: Implement the function pythagoras
### BEGIN SOLUTION
    c_squared = a**2 + b**2
    c = sqrt(c_squared)
    return c
### END SOLUTION

sides = range(1, 5, 1)

# B: Calculate the lengths of the hypothenuses as described
### BEGIN SOLUTION
hypothenuses = {}
for a in sides:
    hypothenuses[a] = {}
    for b in sides:
        if b <= a:
            hyp = pythagoras(a, b)
            hypothenuses[a][b] = hyp
### END SOLUTION

# Print the results
print(f'{"a":>10s}{"b":>10s}{"c":>10s}')
print("-"*30)

for a in hypothenuses:
    for b in hypothenuses[a]:
        print(f"{a:10d}{b:10d}{hypothenuses[a][b]:10.3f}")
```

<!-- grade_id: cell-9bd4aeb96fcbe509, points: 2, grade: true, solution: false -->
```python
assert pythagoras
### BEGIN HIDDEN TESTS
assert abs(pythagoras(1, 2) - 2.23606797749979) < 1e-6
assert abs(pythagoras(0, 0) - 0) < 1e-6
assert abs(pythagoras(12, 12) - 16.97056274847714) < 1e-6
### END HIDDEN TESTS
```

<!-- grade_id: cell-4df55c5fb16520a4, points: 2, grade: true, solution: false -->
```python
assert type(hypothenuses) == dict
for side_length in sides:
    assert side_length in hypothenuses.keys()
    assert type(hypothenuses[side_length]) == dict
### BEGIN HIDDEN TESTS
assert len(hypothenuses.keys()) == 4

for side_length in sides:
    for other_side_length in sides:
        assert other_side_length in hypothenuses[side_length].keys() or side_length in hypothenuses[other_side_length].keys()

number_of_entries = 0
for key, item in hypothenuses.items():
    number_of_entries += len(item)
assert number_of_entries >= 10

assert abs(hypothenuses[1][1] - 1.4142135623730951) < 1e-6
assert abs(hypothenuses[2][2] - 2.8284271247461903) < 1e-6
assert abs(hypothenuses[3][3] - 4.242640687119285) < 1e-6
assert abs(hypothenuses[4][4] - 5.656854249492381) < 1e-6
### END HIDDEN TESTS
```