## Objektoriente Programmierung

### Aufgabe 1: Arbeiten mit Objekt-Instanzen

**a)**

Gegeben ist eine Implementierung einer Klasse `Dog`.

Erstelle eine Instanz der Klasse `Dog` mit dem Variablennamen `my_dog`, deren Attribut `name` den Wert `Bello` hat und deren Attribut `Alter` den Wert 4 hat.

<!-- grade_id: cell-fd51b547f7e21ddc, points: 0, grade: false, solution: true -->
```python
class Dog:
    def __init__(self, name, age):
        self.name = name
        self.age=age

    def __str__(self):
        return "I am a Dog and my name is " + self.name + "and I'm " + str(self.age) + " years old!" 

    def interact(self):
        print("Wuff wuff!")
        
### BEGIN SOLUTION
my_dog = Dog("Bello", 4)
### END SOLUTION
```

<!-- grade_id: cell-abad16bbf308de4b, points: 1, grade: true, solution: false -->
```python
assert my_dog

### BEGIN HIDDEN TESTS
assert my_dog.name == "Bello"
assert my_dog.age == 4
assert type(my_dog) == Dog
### END HIDDEN TESTS
```

**b)**

Konvertiere die Instanz `my_dog` in `str` und speichere das Ergebnis in einer Variablen `dog_str`. Dann drucke die Variable `dog_str` aus. Rufe schließlich die Methode `interact` der Instanz `my_dog` auf.

<!-- grade_id: cell-4123097186d05241, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
dog_str = str(my_dog)
print(dog_str)
my_dog.interact()
### END SOLUTION
```

<!-- grade_id: cell-7942f66cd67a4f46, points: 1, grade: true, solution: false -->
```python
assert dog_str
### BEGIN HIDDEN TESTS
assert "I am a Dog and my name is" in dog_str
### END HIDDEN TESTS
```

**c)**

Implementieren Sie eine Klasse `SausageDog`, die von `Dog` erbt, und erzeugen Sie eine Instanz der Klasse namens `my_sausage_dog`. Der Konstruktor sollte die gleichen Argumente nehmen und sie mit der Funktion `super().__init__()` initialisieren. Schließlich fügen Sie dem Konstruktor ein neues Argument mit dem Namen `vaccinated` hinzu.

<!-- grade_id: cell-400b7b6f5ac210b6, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class SausageDog(Dog):
    def __init__(self, name, age, vaccinated):
        super().__init__(name, age)
        self.vaccinated = vaccinated
        
my_sausage_dog = SausageDog("Heinz", 0, True)
### END SOLUTION
```

<!-- grade_id: cell-0bf7739a493e5e59, points: 1, grade: true, solution: false -->
```python
assert type(my_sausage_dog) == SausageDog

### BEGIN HIDDEN TESTS
assert isinstance(my_sausage_dog, Dog)
### END HIDDEN TESTS
```

### Aufgabe 2: Klassen & Objekte I

**a)**

Implementieren Sie eine Klasse `Kontakt` mit den Attributen `name`, `first_name`, and `phone`.
Alle diese Attribute sollten im Konstruktor der Klasse übergeben werden können.

Implementieren Sie außerdem die Vergleichsfunktion (`__eq__`) für zwei Instanzen der Klasse, die prüft, ob die Inhalte der drei Attribute der Klasse genau übereinstimmen.

**b)**

Implementieren Sie weiterhin die Funktionen `__repr__` und `__str__` der Klasse Contact.

Die String-Repräsentation, die von `__repr__` zurückgegeben wird, sollte ein String sein, der `name`, `first_name`, and `phone` durch ein Komma getrennt enthält. Z.B., „Mustermann, Max, 12345234“

Die Methode `__str__` sollte nur den vollständigen Namen (bestehend aus den Attributen `name` und `first_name` ) der zur Instanz gehörenden Person zurückgeben. Beispiel: `Mayer, Bernd`.

<!-- grade_id: cell-16be56280702e53b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Contact:
    
    def __init__(self, name, first_name, phone):
        self.name = name
        self.first_name = first_name
        self.phone = phone
        
    def __eq__(self, other):
        if self.name != other.name:
            return False
        if self.first_name != other.first_name:
            return False
        if self.phone != other.phone:
            return False
        return True
    
    def __str__(self):
        return self.name + ", " + self.first_name
    
    def __repr__(self):
        return ", ".join([self.name, self.first_name, self.phone])
        
### END SOLUTION
```

<!-- grade_id: cell-93b74041d99b8430, points: 2, grade: true, solution: false -->
```python
# Tests for part a)
assert Contact
assert type(Contact("Mustermann", "Max", "0721 12345678")) == Contact
assert type(Contact("Mustermann", "Max", "0721 12345678")) == Contact
contact_arne = Contact("A.", "Arne", "0721 12345678")
contact_anne = Contact("A.", "Anne", "0721 12345678")
contact_anne_2 = Contact("A.", "Anne", "0721 12345678")
assert contact_arne != contact_anne
assert contact_anne == contact_anne_2

### BEGIN HIDDEN TESTS
assert Contact("", "Anne", "") == Contact("", "Anne", "")
assert Contact("", "Anne", "a") != Contact("", "Anne", "")
assert Contact("a", "Anne", "") != Contact("", "Anne", "")
### END HIDDEN TESTS
```

<!-- grade_id: cell-861425ea80e12a67, points: 2, grade: true, solution: false -->
```python
# Tests for part b)
contact_arne = Contact("A.", "Arne", "1234")
assert contact_arne.__repr__()
assert contact_arne.__str__()
assert contact_arne.__str__() != contact_arne.__repr__()

assert contact_arne.first_name in str(contact_arne)
assert contact_arne.name in str(contact_arne)

### BEGIN HIDDEN TESTS
contact_arne_2 = Contact("A.", "Arne", "12345")
assert str(contact_arne) == str(contact_arne_2)
assert contact_arne.__repr__() != contact_arne_2.__repr__()
assert Contact("A.", "Arne", "1234").__repr__() == contact_arne.__repr__()
### END HIDDEN TESTS
```

### Aufgabe 3: Klassen & Objekte II

Gegeben ist eine Implementierung der Klasse `Point`, die einen Punkt im dreidimensionalen Raum darstellt.

**a)**

Implementiere die Methode `euclidean_norm`, die den euklidischen Abstand vom Ursprung im dreidimensionalen Raum berechnet. Verwenden Sie diese Gleichung ist $||(x,y,z)|| = \sqrt{ x^2 + y^2 + z^2 }$



**b)**

Implementiere nun die magische Methode `__eq__`, die eine zweite Punktinstanz nimmt und prüft, ob sie gleich sind. Zwei Punkte sind identisch, wenn alle drei Koordinaten (`x`, `y` und `z`) identisch sind.

Hinweis: Verwenden Sie die Funktion `isclose` aus dem Modul `math`, um zu prüfen, ob zwei `float`-Werte ausreichend nahe beieinander liegen (z.B. `isclose(x, y, abs_tol=1e-6)`). Die erforderliche Genauigkeit wird durch das Argument `abs_tol` bestimmt und kann auf `1e-6=0.000001` gesetzt werden.


**c)**

Schließlich implementiere die Methoden `__lt__` und `__gt__`, die prüfen, ob eine Instanz kleiner bzw. größer als eine andere Instanz ist. Die Regel hier ist einfach: wenn der euklidische Abstand einer Instanz kleiner ist, dann ist die Instanz kleiner. Wenn zwei Instanzen (Punkte) den gleichen euklidischen Abstand haben, wird die Instanz mit der kleineren x-Koordinate (Attribut `x`) als kleiner angesehen. Im Falle der Gleichheit ist wiederum die Instanz mit der kleineren y-Koordinate (Attribut `y`) die kleinere.

<!-- grade_id: cell-718616833edba891, points: 0, grade: false, solution: true -->
```python
from math import isclose

class Point:
    def __init__(self, x, y, z):
        self.x = float (x)
        self.y = float (y)
        self.z = float (z)
        
    ### BEGIN SOLUTION
    def euclidean_norm(self):
        norm_sq = self.x**2 + self.y**2 + self.z**2
        return norm_sq**0.5

    def __eq__(self, other):
        return isclose(self.x, other.x, abs_tol=1e-6) and isclose(self.y, other.y, abs_tol=1e-6) and isclose(self.z, other.z, abs_tol=1e-6)
    
    def __lt__(self, other):
        norm_self = self.euclidean_norm()
        norm_other = other.euclidean_norm()
        if norm_self < norm_other:
            return True
        elif norm_self > norm_other:
            return False
        else:
            if self.x < other.x:
                return True
            elif self.x > other.x:
                return False
            else:
                return self.y < other.y
    
    def __gt__(self, other):
        return other < self and other != self
    ### END SOLUTION
```

<!-- grade_id: cell-e4191da70eddaf07, points: 1, grade: true, solution: false -->
```python
import math 
p1 = Point(1, 2, 3)
p2 = Point(1, 2, 3)

# Tests for part a)
### BEGIN HIDDEN TESTS
assert math.isclose(p1.euclidean_norm() , 3.7416573867739413, abs_tol=1e-6)
### END HIDDEN TESTS
```

<!-- grade_id: cell-f8a076f37fa7ed0f, points: 2, grade: true, solution: false -->
```python
# Tests for part b)

### BEGIN HIDDEN TESTS
assert Point(1, 2, 3) == Point(1, 2, 3)
assert Point(0, 0, 0) == Point(0, 0, 0)
assert Point(-1, -2, -3) == Point(-1, -2, -3)
assert Point(1.05, 2.9, 3) != Point(2.9, 1.05, 3)
### END HIDDEN TESTS
```

```python
# Tests for part c)

a = Point(1, 1, 2)
b = Point(1, 2, 1)
c = Point(1, 3, 1)
d = Point(1, 1, 2)
e = Point(2, 1, 1)
assert a < c
assert a < b
assert b > a
assert a < e
assert e > a
assert a == d

### BEGIN HIDDEN TESTS
assert Point(1, 2, 3) == Point(1, 2, 3)
assert Point(2, 1, 1) > Point(1, 2, 1)
assert not Point(2, 1, 1) < Point(1, 2, 1)
assert not Point(2, 1, 1) == Point(1, 2, 1)
assert not Point(1, 2, 3) < Point(1, 2, 3)
assert not Point(1, 2, 3) > Point(1, 2, 3)
assert Point(3, 1, 1) > Point(1, 2, 1)
assert not Point(3, 1, 1) < Point(1, 2, 1)
assert Point(1.99, 1, 1) > Point(1.98, 1, 1)
assert not Point(1.99, 1, 1) < Point(1.98, 1, 1)
assert Point(1.99, 1, 1) < Point(1.98, 1.1, 1)
assert not Point(1.99, 1, 1) > Point(1.98, 1.1, 1)
assert Point(1.99, 1, 1) < Point(1.98, 1, 1.1)
assert not Point(1.99, 1, 1.1) < Point(1.98, 1, 1)
assert not Point(1.99, 1, 1) < Point(1.98, 1, 1)
assert not Point(1.99, 1, 1) < Point(1.99, 1, 1)
assert not Point(3.5, 1.09, 2) < Point(2, 3.5, 1.09)
### END HIDDEN TESTS
```