## Objects, Classes and Instances

***
### Exercise 1

In the subdirectory `programs` of this exercise, you can find the module `coffee_machine.py`. This module contains the implementation of the class `CoffeeMachine`.

The already implemented code adds the subdirectory to your interpreter's path.

Import the class `CoffeeMachine` from the custom module.
Create an instance of the class with name `my_coffee_machine`.
Then, use the class' methods to add coffee and water to your coffee machine instance and to brew a coffee!

<!-- grade_id: cell-e647e7bfbca1e9b8, points: 0, grade: false, solution: true -->
```python
from pathlib import Path
import sys

notebook_dir = str(Path().resolve())
import_dir = notebook_dir + "/programs"

sys.path.insert(0, import_dir)

### BEGIN SOLUTION
from coffee_machine import CoffeeMachine

my_coffee_machine = CoffeeMachine()

my_coffee_machine.add_coffee(100)
my_coffee_machine.add_water(500)
my_coffee_machine.brew_coffee()
### END SOLUTION
```

<!-- grade_id: cell-30edd9af18414682, points: 1, grade: true, solution: false -->
```python
assert isinstance(my_coffee_machine, CoffeeMachine)
assert my_coffee_machine.brew_counter > 0
```

***
### Exercise 2

The class `CoffeeMachine` already has a class variable called `waste_storage_size` and an instance variable `waste_storage`.
Both are not used in the class' current implementation.

Extend the code of the class in the file `coffee_machine.py` as follows:

* For each brewing of coffee (`brew_coffee()`), the used coffee grounds are added to the machine's waste storage (`waste_storage`). (You can assume 1 unit of used coffee corresponds to 1 unit of waste.)
* A brewing can only be performed as long as the waste storage is not full (compare `waste_storage` to `waste_storage_size`). Make changes to the method `is_ready_to_brew` accordingly.
* Add a method `empty_waste_storage` to empty the coffee machine's waste storage.

Then, import the changed class `CoffeeMachine` again and create a new instance called `my_coffee_machine`.
Brew coffee until the waste storage is filled to its capacity limit and the brewing of coffee is no longer possible.
(You will have to add water and coffee to do so. According to the current configuration, you have to brew 10 cups of coffee to fill the waste storage.)

*Hint:* After making changes to the implementation in `coffee_machine.py`, the module has to be imported anew. To bring your interpreter to re-import the changed module, you either have to restart the interpreter or to explicitly delete the module from the running interpreter (this is what the code below does).

<!-- grade_id: cell-dd469a73bfc73bfe, points: 0, grade: false, solution: true -->
```python
from pathlib import Path
import sys

notebook_dir = str(Path().resolve())
import_dir = notebook_dir + "/programs"

sys.path.insert(0, import_dir)

# Entferne ggf. das bereits importierte Modul, um einen Reload zu erzwingen und Änderungen im Modul zu berücksichtigen.
if "coffee_machine" in sys.modules.keys():
    del sys.modules["coffee_machine"]

### BEGIN SOLUTION
from coffee_machine import CoffeeMachine

my_coffee_machine = CoffeeMachine()
for i in range(10):
    my_coffee_machine.add_coffee(12)
    my_coffee_machine.add_water(150)
    my_coffee_machine.brew_coffee()
### END SOLUTION
```

<!-- grade_id: cell-e1f0f971caa63fc6, points: 1, grade: true, solution: false -->
```python
assert isinstance(my_coffee_machine, CoffeeMachine)
assert my_coffee_machine.waste_storage > my_coffee_machine.waste_storage_size
my_coffee_machine.waste_storage = 0
my_coffee_machine.add_water(150)
my_coffee_machine.add_coffee(12)
assert my_coffee_machine.is_ready_to_brew()
```