## Datenverarbeitung mit Pandas

<img align="center" width=300 src="figures/pandas.svg"/>

Pandas ist eine Python-Bibliothek, die auf numpy aufbaut und für die Verarbeitung, Analyse und Darstellung von Tabellendaten verwendet wird.

In vielen Fällen könnt Ihr pandas auch verwendet, um Daten für KI Anwendungen vorzubereiten.

Pandas ist sehr mächtig. Wir können hier daher nur einen ersten Einstieg in die wichtigsten Konzepte geben.
Für Pandas gibt es sehr gute weiterführende Tutorials.

Hier findet Ihr die offizielle Dokumenation und Tutorials: [pandas Dokumentation](https://pandas.pydata.org/docs/). 

Die wichtigsten Quelle für dieses Tutorial ist enfalls von dort: [Getting Started](https://pandas.pydata.org/docs/getting_started/intro_tutorials/index.html)


## Welche Daten kann man mit Pandas verarbeiten?

Tabellarische Daten, d.h. Daten, die aus einem Spreadsheet (z.B. einer Exceldatei) oder einer Datenbank (z.B. einer SQL Datenbank) kommen, können mit Pandas eingelesen, analysiert und weiterverarbeitet werden.
Die zentrale Datenstruktur in Pandas ist der ``DataFrame``.

<img align="center" src="figures/df.png" width="800" >

Wichtige Eigenschaften von DataFrames:
* Ein DataFrame ist 2-dimensional, hat also Zeilen und Spalten und kann unterschiedliche Datentypen wie int, float oder strings enthalten.
* Der hier dargestellte DataFrame hat 5 Spalten (``matriculation_number``, ``surname``, ``course_of_study``, ``exam_grade``, ``pcc``) und 6 Zeilen. Die Spaltennamen nennt man auch Spaltenindex und die Zeilennamen auch Zeilenindex.
* Alle Daten einer Spalte sind vom gleichen Datentyp oder werden beim Einlesen in den gleichen Datentyp umgewandelt. Im Beispiel ist die Klausurnote vom Typ float während der Name vom Typ String ist. 

### Teil 1: Ein erster DataFrame  

Wir können den im Beispiel gezeigten DataFrame sehr leicht als entsprechende Instanz erzeugen.

Als Ausgang dient uns ein Dictionary:
Die Keys des Dictionary werden zu Spaltennamen.
Die Values sind jeweils Listen, die alle die gleiche Länge (in unserem Beispiel 6) haben.

```python
# Let's install pandas first :)
!pip install pandas
```

```python
import pandas as pd 

# Dictionary with the outputs
d = {
    "matriculation_number" : [12345, 90890, 89750, 90876, 90873, 54365],
    "surname" : ["Hummels", "Witsel", "Reus", "Matthäus", "Kroos", "Dörr"],
    "course_of_study" : ["IE", "ME", "EE", "ME", "ME", "IE"],
    "exam_grade" : [1.0, 1.3, 1.7, 1.0, 2.0, 4.0],
    "pcc" : [True, False, True, False, True, True]
}

# We pass the dictionary as an argument
s = pd.DataFrame(d)

# In Jupyter, the tables are displayed very nicely with the help of display()
display(s)
```

```python
# We can also use the print() function to show the DataFrame.
print(s)
```

Die DataFrame-Methode``.describe()`` gibt einen ersten Überblick über die numerischen Werte der Tabelle.

```python
s.describe()
```

```python
# This returns a tuple (<n rowes>, <n columns>)
s.shape
```

```python
# Names of the columns
s.columns
```

```python
# index 
s.index
```

### Aufgabe 1

**a)**

Erstellen Sie ein leeres Dictionary names `machine_dict`, welches Ihnen als Ausgang zur Erstellung eines DataFrames dient.

Das Dictionary soll die Informationen aus der unten dargestellten Tabelle enthalten.

Danach verwenden Sie dieses Dictionary, um einen DataFrame namens `machine_df` zu erstellen. Schließlich, verwenden Sie die **display()**-Funktion, um den erstellten DataFrame anzuzeigen. 

| machine | speed |  tempreture | status |
|---|---|---| --- |
| 1  | 24  | 12.  | 1|
| 2  |  30 |  16.5 | 1 |
| 2  |  90 |  39.4 | 1 |
| 3  | 40  | 22.2  |1 |
| 3  | 9  |  3.92 | 1 |


<!-- grade_id: cell-dd5f71940e1efef7, points: 0, grade: false, solution: true -->
```python
import pandas as pd 

machine_dict = None
machine_df = None 

### BEGIN SOLUTION
machine_dict = {
    "machine" : [1,2,2,3,3],
    "speed" : [24,30,90,40,9],
    "temperature" : [12.0, 16.5, 39.4, 22.2, 3.92],
    "status": [1,1,1,1,1]
}
machine_df = pd.DataFrame(machine_dict)
### END SOLUTION
display(machine_df)
```

<!-- grade_id: cell-b325560770f08782, points: 1, grade: true, solution: false -->
```python
assert isinstance(machine_dict, dict)
assert isinstance(machine_df, pd.DataFrame)
assert {'speed', 'machine', 'status', 'temperature'} == set(machine_df.columns)
assert machine_df.shape == (5,4)
```

**b)**

Verwenden Sie die **describe()**-Methode, um deskriptive statistische Werte über die numerischen Felder im DataFrame zu erhalten. 

<!-- grade_id: cell-73c674e1f6d2187d, points: 0, grade: false, solution: true -->
```python
machine_df.describe()
```

<!-- grade_id: cell-b916ab277ae6c188, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert machine_df.describe().shape == (8, 4)
assert set(machine_df.describe().columns) == {'machine', 'speed', 'temperature', 'status'}
### END HIDDEN TESTS
```

### Teil 2: DataFrames Einlesen 

Oft werden Dataframes aus einer .csv-Datei gelesen, die Tabelleneinträge enthält. 
csv-Dateien könnt Ihr zum Beispiel mit einem Texteditor (z.B. Wordpad, notepadd++) oder mit Excel öffnen.

Die Methode `pandas.read_csv()` zum Erstellen eines Dataframes aus einer csv-Datei hat eine ganze Reihe von Parametern, von denen die wichtigsten sind 
* der Dateiname (inkl. Dateipfad)
* `sep`: das Zeichen, das zur Trennung der Zeileneinträge verwendet wird (z.B. `,` oder `;`)

Lesen wir beispielhaft die Datei ``students.csv`` (Trennzeichen `,`) die im Ordner ``data`` liegt:

```python
students = pd.read_csv("data/students.csv", sep = ",")
display(students)
```

### Aufgabe 2

**a)**

Lesen Sie die Datei `machine_data_small.csv` (mit Komma getrennt) die im Ordner ``data`` liegt ein und weisen Sie diesen DataFrame der Variable `machine_df` zu. Schließlich können Sie den eingelsenen DataFrame mit der `display()` Funktion anzeigen. 

<!-- grade_id: cell-57935cb85a0eacc5, points: 0, grade: false, solution: true -->
```python
machine_df = None

### BEGIN SOLUTION
machine_df = pd.read_csv("data/machine_data_small.csv", sep = ",")
display(machine_df)
### END SOLUTION
```

<!-- grade_id: cell-bd1b09605ed9a483, points: 1, grade: true, solution: false -->
```python
import math
assert set(machine_df.columns) == {'machine', 'speed', 'temperature', 'status'}
assert machine_df.shape == (5,4)
assert machine_df["status"].mean() == 1
```

### Teil 3: Zugriff auf Spalten und Zeilen und einfache Operationen

<img align="center" width=600 src="figures/data_frame_select_col.png"/>


Mit ``[]``, d.h, konkret mit  ``students[‘exam_grade’]``, können wir auf die ganze Spalte `exam_grade` zugreifen.

```python
# We can access a single column
exam_grades = students['exam_grade']
exam_grades
```


Um auf mehrere Spalten zugleich zuzugreifen, z.B. `matriculation_number` und `exam_grade`, kann eine Liste von Spaltennamen in ``[]`` eingefügt werden.
Das Ergebnis ist dann erneut vom Typ `DataFrame`.

```python
# We also access multiple columns at once
students[['exam_grade', 'surname', 'pcc']]
```

```python
# Maximum of a column
exam_grades.max()
```

```python
# Minimum of a column
exam_grades.min()
```

```python
# Standard deviation of a column
exam_grades.std()
```

```python
# Sum of a column
exam_grades.sum()
```

### Unnötige Zeilen oder Spalten Löschen

Mit der Methode **drop()** können wir Spalten oder Zeilen aus unserem DataFrame entfernen. 
Diese Methode nimmt eine Liste von Zeilen-/Spaltennamen zusammen mit einem ``axis``-Parameter entgegen. 

* Der Parameter ``axis=0`` oder ``axis='index'`` zeigt an, dass wir ***Zeilen*** löschen.
* Der Parameter ``axis=1`` oder ``axis='columns'`` zeigt an, dass wir ***Spalten*** löschen.

```python
students = students.drop([4], axis = "index")
students
```

```python
students = students.drop(["pcc"], axis = "columns")
display(students)
```

### Aufgabe 3

**a)**

Greifen Sie auf die ``tempreture``-Spalte im DataFrame ``machine_df`` zu und weisen Sie sie der Variablen ``tempreture`` zu. Danach geben Sie das Minimum, das Maximum und den Durchnitt der Werte der Spalte aus. 

Tipp: Verwenden Sie die Methode **.mean()**, um den Durchnitt einer Spalte zu berechnen.

<!-- grade_id: cell-d7398c4268326099, points: 0, grade: false, solution: true -->
```python
temperature = None 

### BEGIN SOLUTION
temperature = machine_df["temperature"]
print(temperature.max())
print(temperature.min())
print(temperature.mean())
### END SOLUTION
```

<!-- grade_id: cell-0c92ee5f59a59d5c, points: 0, grade: true, solution: false -->
```python
assert isinstance(temperature, pd.Series)
assert temperature.max() == 39.4
assert temperature.min() == 3.92
assert (temperature.mean() > 18) and (temperature.mean() < 19)
```

**b)**

Löschen Sie die Spalte ``status`` aus dem DataFrame ``machine_df`` und weisen Sie den neuen erzegten DataFrame (Ohne der gelöschten Spalte) der Variablen ``machine_df`` zu.  

<!-- grade_id: cell-f2aff4def0369ec0, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
machine_df = machine_df.drop(["status"], axis = "columns")
### END SOLUTION
```

<!-- grade_id: cell-62af690e68e9f910, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert "status" not in machine_df.columns
### END HIDDEN TESTS
```

### Teil 4: Zeilen Filtern

<img align="center" width=600 src="figures/data_frame_filter_row.png"/>

Spannend ist es auch, nur bestimmte Zeilen eines DataFrames zu betrachten.
Zum Beispiel, wenn Sie die Studierenden betrachten möchten, deren Note besser oder gleich 1.3 ist, d.h. für die gilt, dass die Spalte `klausurnote <= 1.3` ist.
Das kann über eine Filterung erreicht werden, indem die Filter-Bedingung innerhalb einer eckigen Klammer ``[]`` wie folgt beschrieben wird:

```python
best_students = students[students["exam_grade"] >= 1.3]
best_students
```

Um mehreren Bedingungen zu implementieren, kann man die Bedingungen innerhalb Klammern implementieren und sie mit einander durch logische Operatoren verknüpfen. 

WICHTIG: um Bedingung in Pandas zu verküpfen, verwenden Sie das **&**-Zeichen anstatt des **and**-Operators, und das **|**-Zeichen (Senkrechter Strich) anstatt des **or**-Operators.

```python
best_me_students = students[
    (students["exam_grade"] >= 1.3) & 
    (students["course_of_study"] == "ME")
]

best_me_students
```

```python
students[
    (students["exam_grade"] >= 1.3) | 
    (students["course_of_study"] == "MACH")
]
```

### Aufgabe 4

**a)**

Filtern Sie die Zeilen des DataFrames ``machine_df`` nach den Kriterien ``machine == 2`` und ``speed > 35`` und weisen Sie das Ergebnis der Filterung der Variablen `m2` zu.
Geben Sie danach ``m2`` aus und prüfen Sie das Ergebnis.

<!-- grade_id: cell-fb671dd713b0c5c4, points: 0, grade: false, solution: true -->
```python
m2 = None 

### BEGIN SOLUTION
m2 = machine_df[
    (machine_df["machine"] == 2) & 
    (machine_df["speed"] > 35)
]
m2
### END SOLUTION
```

<!-- grade_id: cell-04b22c65ab47fde7, points: 0, grade: true, solution: false -->
```python
assert m2.shape == (1,3)
assert m2["machine"].values[0] == 2
assert m2["speed"].values[0] == 90
assert m2["temperature"].values[0] == 39.4
```