## Object-oriented Programming

### Exercise 1: Working with Object-instances

**a)**

Given is an implementation of a class `Dog`.

Create an instance of the class `Dog` with variable name `my_dog`, whose attribute `name` has the value `Bello` and whose attribute `age` is 4.

<!-- grade_id: cell-fd51b547f7e21ddc, points: 0, grade: false, solution: true -->
```python
class Dog:
    def __init__(self, name, age):
        self.name = name
        self.age=age

    def __str__(self):
        return "I am a Dog and my name is " + self.name + "and I'm " + str(self.age) + " years old!" 

    def interact(self):
        print("Wuff wuff!")
        
### BEGIN SOLUTION
my_dog = Dog("Bello", 4)
### END SOLUTION
```

<!-- grade_id: cell-abad16bbf308de4b, points: 1, grade: true, solution: false -->
```python
assert my_dog

### BEGIN HIDDEN TESTS
assert my_dog.name == "Bello"
assert my_dog.age == 4

assert type(my_dog) == Dog
### END HIDDEN TESTS
```

**b)**

Convert the instance `my_dog` to `str` and save the result in a variable `dog_str`. Then print out the variable `dog_str`. Finally, call the method `interact` of the instance `my_dog`. 

<!-- grade_id: cell-4123097186d05241, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
dog_str = str(my_dog)
print(dog_str)
my_dog.interact()
### END SOLUTION
```

<!-- grade_id: cell-7942f66cd67a4f46, points: 1, grade: true, solution: false -->
```python
assert dog_str
### BEGIN HIDDEN TESTS
assert "I am a Dog and my name is" in dog_str
### END HIDDEN TESTS
```

**c)**

Implement a class `SausageDog` which inherits from `Dog` and create an instance of the class named `my_sausage_dog`. The constructor should take the same arguments and initialize them using `super().__init__()` function. Finally, add a new argument to the constructor named `vaccinated`. 

<!-- grade_id: cell-400b7b6f5ac210b6, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class SausageDog(Dog):
    def __init__(self, name, age, vaccinated):
        super().__init__(name, age)
        self.vaccinated = vaccinated
        
        

my_sausage_dog = SausageDog("Heinz", 0, True)
### END SOLUTION
```

<!-- grade_id: cell-0bf7739a493e5e59, points: 1, grade: true, solution: false -->
```python
assert type(my_sausage_dog) == SausageDog

### BEGIN HIDDEN TESTS
assert isinstance(my_sausage_dog, Dog)
### END HIDDEN TESTS
```

### Task 2: Classes & Objects I

**a)**

Implement a class `Contact` with the attributes `name`, `first_name`, and `phone`.
All these attributes should be able to be passed in the constructor of the class. 

Also, implement the comparison function (`__eq__`) for two instances of the class, which checks whether the contents of the three attributes of the class match exactly.

**b)**

Continue to implement functions `__repr__` and `__str__` of the Contact class.

The string representation returned by `__repr__` should be a string that includes the `name`, `first_name`, and `phone` separated by a comma. E.g., "Mustermann, Max, 12345234"

The method `__str__` should only return the full name (consisting of the attributes `name` and `first_name` ) of the person belonging to the instance. Example: `Mayer, Bernd`.

<!-- grade_id: cell-16be56280702e53b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Contact:
    
    def __init__(self, name, first_name, phone):
        self.name = name
        self.first_name = first_name
        self.phone = phone
        
    def __eq__(self, other):
        if self.name != other.name:
            return False
        if self.first_name != other.first_name:
            return False
        if self.phone != other.phone:
            return False
        return True
    
    def __str__(self):
        return self.name + ", " + self.first_name
    
    def __repr__(self):
        return ", ".join([self.name, self.first_name, self.phone])
        
### END SOLUTION
```

<!-- grade_id: cell-93b74041d99b8430, points: 2, grade: true, solution: false -->
```python
# Tests for part a)
assert Contact
assert type(Contact("Mustermann", "Max", "0721 12345678")) == Contact
assert type(Contact("Mustermann", "Max", "0721 12345678")) == Contact
contact_arne = Contact("A.", "Arne", "0721 12345678")
contact_anne = Contact("A.", "Anne", "0721 12345678")
contact_anne_2 = Contact("A.", "Anne", "0721 12345678")
assert contact_arne != contact_anne
assert contact_anne == contact_anne_2

### BEGIN HIDDEN TESTS
assert Contact("", "Anne", "") == Contact("", "Anne", "")
assert Contact("", "Anne", "a") != Contact("", "Anne", "")
assert Contact("a", "Anne", "") != Contact("", "Anne", "")
### END HIDDEN TESTS
```

<!-- grade_id: cell-861425ea80e12a67, points: 2, grade: true, solution: false -->
```python
# Tests for part b)
contact_arne = Contact("A.", "Arne", "1234")
assert contact_arne.__repr__()
assert contact_arne.__str__()
assert contact_arne.__str__() != contact_arne.__repr__()

assert contact_arne.first_name in str(contact_arne)
assert contact_arne.name in str(contact_arne)

### BEGIN HIDDEN TESTS
contact_arne_2 = Contact("A.", "Arne", "12345")
assert str(contact_arne) == str(contact_arne_2)
assert contact_arne.__repr__() != contact_arne_2.__repr__()
assert Contact("A.", "Arne", "1234").__repr__() == contact_arne.__repr__()
### END HIDDEN TESTS
```

### Task 3: Classes & Objects II

Given is an implementation of the class `Point`, which represents a point in three-dimensional space.

**a)**

Implement the method `euclidean_norm` that computes Euclidean distance from the origin in 3-dimensional space. Use this equation is $||(x,y,z)|| = \sqrt{ x^2 + y^2 + z^2 }$



**b)**

Now implement the magic method `__eq__` that takes a second point instance and checks if they are equal. Two points are identical if all three coordinates (`x`, `y` and `z`) are identical.

Note: Use the function `isclose` from the module `math` to check whether two `float` values are sufficiently close to each other (e.g. `isclose(x, y, abs_tol=1e-6)`). The required precision is determined by the argument `abs_tol` and can be set to `1e-6=0.000001`.


**c)**

Finally implement the methods `__lt__` and `__gt__`, which check if an instance is less than or greater than another instance, respectively. The rule here is simple: if the Euclidean distance of an instance is smaller, then the instance is smaller. If two instances (points) have the same Euclidean distance, the instance with the smaller x-coordinate (attribute `x`) is considered smaller. In the case of equality again, the instance with the smaller y-coordinate (attribute `y`) is the smaller one.



<!-- grade_id: cell-718616833edba891, points: 0, grade: false, solution: true -->
```python
from math import isclose

class Point:
    def __init__(self, x, y, z):
        self.x = float (x)
        self.y = float (y)
        self.z = float (z)
        
    ### BEGIN SOLUTION
    def euclidean_norm(self):
        norm_sq = self.x**2 + self.y**2 + self.z**2
        return norm_sq**0.5

    def __eq__(self, other):
        return isclose(self.x, other.x, abs_tol=1e-6) and isclose(self.y, other.y, abs_tol=1e-6) and isclose(self.z, other.z, abs_tol=1e-6)
    
    def __lt__(self, other):
        norm_self = self.euclidean_norm()
        norm_other = other.euclidean_norm()
        if norm_self < norm_other:
            return True
        elif norm_self > norm_other:
            return False
        else:
            if self.x < other.x:
                return True
            elif self.x > other.x:
                return False
            else:
                return self.y < other.y
    
    def __gt__(self, other):
        return other < self and other != self
    ### END SOLUTION
```

```python
import math 
p1 = Point(1, 2, 3)
p2 = Point(1, 2, 3)

# Tests for part a)
### BEGIN HIDDEN TESTS
assert math.isclose(p1.euclidean_norm() , 3.7416573867739413, abs_tol=1e-6)
### END HIDDEN TESTS
```

```python
# Tests for part b)

### BEGIN HIDDEN TESTS
assert Point(1, 2, 3) == Point(1, 2, 3)
assert Point(0, 0, 0) == Point(0, 0, 0)
assert Point(-1, -2, -3) == Point(-1, -2, -3)
assert Point(1.05, 2.9, 3) != Point(2.9, 1.05, 3)
### END HIDDEN TESTS
```

<!-- grade_id: cell-f8a076f37fa7ed0f, points: 2, grade: true, solution: false -->
```python
# Tests for part c)

a = Point(1, 1, 2)
b = Point(1, 2, 1)
c = Point(1, 3, 1)
d = Point(1, 1, 2)
e = Point(2, 1, 1)
assert a < c
assert a < b
assert b > a
assert a < e
assert e > a
assert a == d

### BEGIN HIDDEN TESTS
assert Point(1, 2, 3) == Point(1, 2, 3)
assert Point(2, 1, 1) > Point(1, 2, 1)
assert not Point(2, 1, 1) < Point(1, 2, 1)
assert not Point(2, 1, 1) == Point(1, 2, 1)
assert not Point(1, 2, 3) < Point(1, 2, 3)
assert not Point(1, 2, 3) > Point(1, 2, 3)
assert Point(3, 1, 1) > Point(1, 2, 1)
assert not Point(3, 1, 1) < Point(1, 2, 1)
assert Point(1.99, 1, 1) > Point(1.98, 1, 1)
assert not Point(1.99, 1, 1) < Point(1.98, 1, 1)
assert Point(1.99, 1, 1) < Point(1.98, 1.1, 1)
assert not Point(1.99, 1, 1) > Point(1.98, 1.1, 1)
assert Point(1.99, 1, 1) < Point(1.98, 1, 1.1)
assert not Point(1.99, 1, 1.1) < Point(1.98, 1, 1)
assert not Point(1.99, 1, 1) < Point(1.98, 1, 1)
assert not Point(1.99, 1, 1) < Point(1.99, 1, 1)
assert not Point(3.5, 1.09, 2) < Point(2, 3.5, 1.09)
### END HIDDEN TESTS
```