## Files and Data

### Exercise 1

In the Python crash course, we already learned the basics of working with files in Python.

In this task, we will write tabular data to a file and gradually expand the file.
Finally, we will read in all the data and continue working with it.

#### A) Writing data to a file

The data of a production facility is to be stored in a file named `production_data.csv`.

Each day, the number of produced parts (`amount`) and the number of defective parts (`scrap`) should be recorded in it.
In this task, we will initialize the file and first write the corresponding header line.

Open the file in write mode (`"w"`) and write the text line `"day;amount;scrap"` to the file (the line should end with a line break using `\n`).

<!-- grade_id: cell-2cf4b1cdd6ffcd33, points: 0, grade: false, solution: true -->
```python
from pathlib import Path

notebook_dir = Path().resolve()
datafile_path = notebook_dir / "production_data.csv"

### BEGIN SOLUTION
with open(datafile_path, "w") as file:
    file.write("day;amount;scrap\n")
    file.close()
### END SOLUTION
```

<!-- grade_id: cell-ee4f7b4506a489bf, points: 0, grade: true, solution: false -->
```python
assert Path.is_file(datafile_path)
```

#### B) Adding data to an existing file

Now a function should be implemented that adds the data of a production day to the file.
The function `append_production_data(day, amount, scrap)` takes as arguments a day (index), the quantity produced on that day, and the number of defective parts for that day.
This data should then be appended to the file in the same format as the header line (see part **A)**).

*Note*: To avoid overwriting the entire file, the file must be opened in append mode (`a`).

<!-- grade_id: cell-9b287e8b4ed2d40a, points: 0, grade: false, solution: true -->
```python
def append_production_data(day, amount, scrap):
### BEGIN SOLUTION
    with open(datafile_path, "a") as file:
        file.write(f"{day};{amount};{scrap}\n")
        file.close()
### END SOLUTION

append_production_data(1, 1400, 19)
append_production_data(2, 1400, 21)
append_production_data(3, 1600, 22)
append_production_data(4, 1600, 24)
append_production_data(5, 1600, 20)
append_production_data(6, 1600, 36)
append_production_data(7, 1600, 39)
append_production_data(8, 1400, 19)
append_production_data(9, 1400, 22)
append_production_data(10, 1600, 22)
append_production_data(11, 1600, 26)
append_production_data(12, 1600, 24)
append_production_data(13, 1600, 32)
append_production_data(14, 1600, 37)
```

#### C) Reading data from a (csv) file

The written file follows the so-called CSV format ("comma separated values").
Using Python’s built-in `csv` module, data in this format can be read more conveniently.

In the following code block, reading the data has already been implemented.
The variable `line` contains, in each loop iteration, a list of the values from one line (data type `str`).
Convert the read-in values into numeric values and add them accordingly to the two lists `amounts` and `scraps`.

Then calculate the total number of produced parts as well as the total number of defective parts over the entire recorded period (14 days).

*Notes*:
* Keep in mind that the first row of the file is the header and does not contain any data.
* You can open the previously written file manually to get a better understanding of the data and how the program works.

<!-- grade_id: cell-27a47ffbfbbfb866, points: 0, grade: false, solution: true -->
```python
import csv

amounts = []
scraps = []

with open(datafile_path, "r") as file:
    csv_data_reader = csv.reader(file, delimiter=";")
    for line in csv_data_reader:
        ### Convert numeric data and append lists amounts and scraps
        ### BEGIN SOLUTION
        if line[0] == "day":
            continue
        amounts.append(float(line[1]))
        scraps.append(float(line[2]))
        ### END SOLUTION
        
amount_sum = None
scraps_sum = None
### Compute sums
### BEGIN SOLUTION
amount_sum = sum(amounts)
scraps_sum = sum(scraps)
### END SOLUTION
```

<!-- grade_id: cell-482e19064dca6a63, points: 0, grade: true, solution: false -->
```python
assert len(amounts) == 14
assert amounts[0] == 1400
assert len(scraps) == 14
assert amount_sum == 21600
assert scraps_sum == 363
```