## Data processing with Pandas

<img align="center" width=300 src="figures/pandas.svg"/>

Pandas is a Python library based on numpy that processes, analyzes, and displays tabular data.

In many cases, you can also use pandas to prepare data for AI applications.

Pandas is very powerful. Therefore, We can only briefly introduce the most important concepts here.
There are very good advanced tutorials for Pandas.

You can find the official documentation and tutorials here: [pandas documentation](https://pandas.pydata.org/docs/). 

The most important source for this tutorial is also from there: [Getting Started](https://pandas.pydata.org/docs/getting_started/intro_tutorials/index.html)



## Which data can be processed with Pandas?

Tabular data, meaning data that comes from a spreadsheet (e.g., an Excel file) or a database (e.g., an SQL database), can be read in, analyzed, and further processed with Pandas.
The central data structure in Pandas is the ``DataFrame``.

<img align="center" src="figures/df.png" width="800" >

Important properties of DataFrames:
* A DataFrame is 2-dimensional, such that it has rows and columns and can contain different data types such as int, float, or strings.
* The DataFrame shown here has 5 columns (``matriculation_number``, ``surname``, ``course_of_study``, ``exam_grade``, ``pcc``) and 6 rows (0, ..., 4). The column names are also called column indices, and the row names are also called row indices.
* All data in a column is of the same data type or is converted to the same data type when read in. In the example, the exam grade is of type float while the name is of type string.



### Task 1: Your first DataFrame  

We can easily create the DataFrame shown in the example. 
We use a dictionary as the output:
* The keys of the dictionary become column names.
* The values must be lists that all have the same length - in our example 6!

```python
# Let's install pandas first :)
!pip install pandas
```

<!-- grade_id: cell-52127398526d5745, points: 0, grade: true, solution: false -->
```python
import pandas as pd 


# Dictionary with the outputs
d = {
    "matriculation_number" : [12345, 90890, 89750, 90876, 90873, 54365],
    "surname" : ["Hummels", "Witsel", "Reus", "Matthäus", "Kroos", "Dörr"],
    "course_of_study" : ["IE", "ME", "EE", "ME", "ME", "IE"],
    "exam_grade" : [1.0, 1.3, 1.7, 1.0, 2.0, 4.0],
    "pcc" : [True, False, True, False, True, True]
}

# We pass the dictionary as an argument
s = pd.DataFrame(d)

# In Jupyter, the tables are displayed very nicely with the help of display()

display(s)

```

```python
# We can also use the print() function to show the dataframe. But it the display() function shows 
# the table in a nicer way 
print(s)
```

As shown below, you can use the ``.describe()`` method to get an initial overview of the numerical values in the table.

```python
s.describe()
```

```python
# This returns a tuple (<n rowes>, <n columns>)
s.shape
```

```python
# names of the columns
s.columns
```

```python
# index 
s.index
```

**a)**

Create a dictionary called `machine_dict`, which serves as the output for creating a DataFrame.

The Dictionary should contain the information from the one shown below. Then use this dictionary to create a DataFrame called `machine_df`. 

Finally, use the **display()** function to display the created DataFrame.

| machine | speed |  tempreture | status |
|---|---|---| --- |
| 1  | 24  | 12.  | 1|
| 2  |  30 |  16.5 | 1 |
| 2  |  90 |  39.4 | 1 |
| 3  | 40  | 22.2  |1 |
| 3  | 9  |  3.92 | 1 |


<!-- grade_id: cell-dd5f71940e1efef7, points: 0, grade: false, solution: true -->
```python
import pandas as pd 

machine_dict = None
machine_df = None 

### BEGIN SOLUTION
machine_dict = {
    "machine" : [1,2,2,3,3],
    "speed" : [24,30,90,40,9],
    "temperature" : [12.0, 16.5, 39.4, 22.2, 3.92],
    "status": [1,1,1,1,1]
}
m = pd.DataFrame(machine_dict)
### END SOLUTION
display(m)
```

<!-- grade_id: cell-b325560770f08782, points: 1, grade: true, solution: false -->
```python
assert isinstance(machine_dict, dict)
assert isinstance(m, pd.DataFrame)
### BEGIN HIDDEN TESTS
assert {'speed', 'machine', 'status', 'temperature'} == set(m.columns)
assert m.shape == (5,4) 
### END HIDDEN TESTS
```

**b)**

Use the **describe()** method to show descriptive statistical values about the numerical fields in the DataFrame.

<!-- grade_id: cell-73c674e1f6d2187d, points: 0, grade: false, solution: true -->
```python
m.describe()
```

<!-- grade_id: cell-b916ab277ae6c188, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert m.describe().shape == (8, 4)
assert set(m.describe().columns) == {'machine', 'speed', 'temperature', 'status'}
### END HIDDEN TESTS
```

### Task 2: Reading DataFrames

Dataframes are often read from a .csv file that contains table entries. 
You can open .csv files with a text editor (e.g., Wordpad, Notepad++) or with Excel, for example.

The method `pandas.read_csv()` for creating a data frame from a .csv (text) file has a whole series of parameters, the most important of which are
* the (required) file name
* sep: the character used to separate the line entries

For example, let's read the file ``students.csv`` (separated by a comma), which is located in the ``data`` folder:

```python
students = pd.read_csv("data/students.csv", sep = ",")
display(students)
```

**a)**

Read the file `machine_data_small.csv` (comma separated), which is located in the folder ``data`` and assign this DataFrame to the variable `machine_df`. Finally, you can display the inserted DataFrame with the `Display()` function. 

<!-- grade_id: cell-57935cb85a0eacc5, points: 0, grade: false, solution: true -->
```python
machine_df = None

### BEGIN SOLUTION
machine_df = pd.read_csv("data/machine_data_small.csv", sep = ",")
display(machine_df)
### END SOLUTION
```

<!-- grade_id: cell-bd1b09605ed9a483, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import math
assert set(machine_df.columns) == {'machine', 'speed', 'temperature', 'status'}
assert machine_df.shape == (5,4)
assert machine_df["status"].mean() == 1
### END HIDDEN TESTS
```

### Task 3: Accessing columns and rows and simple operations

<img align="center" width=600 src="figures/data_frame_select_col.png"/>


With ``[]``, i.e. specifically with ``students[‘exam_grade’]``, we can access the whole column `exam_grade`. 
If you are interested in more than one column, e.g., ``matriculation_number``, you can also simply select these via ``[]``.
You simply enter all the column names you are interested in as a list. 
You will then receive another DataFrame as output of this operation. 

```python
# We can access a single column
exam_grades = students['exam_grade']
exam_grades
```

```python
# We also access multiple columns at once
students[['exam_grade', 'surname', 'pcc']]
```

```python
# Maximum of a column
exam_grades.max()
```

```python
# Minimum of a column
exam_grades.min()
```

```python
# Standard deviation of a column
exam_grades.std()
```

```python
# Sum of a column
exam_grades.sum()
```

### Delete unnecessary rows or columns

We can use the **drop()** method to remove columns or rows from our DataFrame. 
This method takes a list of row/column names together with an ``axis`` parameter.

* The parameter ``axis=0`` or ``axis=‘index’`` indicates that we are deleting ***rows***.
* The parameter ``axis=1`` or ``axis=‘columns’`` indicates that we are deleting ***columns***.

```python
students = students.drop([4], axis = "index")
students


```

```python
students = students.drop(["pcc"], axis = "columns")
display(students)
```

**a)**

Access the ``tempreture`` column in the DataFrame ``machine_df`` and assign it to the variable ``tempreture``. Then, print out the minimum, maximum, and average the column `tempreture`

Tip: Use the **.mean()** method to get the average of a column.

<!-- grade_id: cell-d7398c4268326099, points: 0, grade: false, solution: true -->
```python
temperature = None 

### BEGIN SOLUTION
temperature = machine_df["temperature"]
print(temperature.max())
print(temperature.min())
print(temperature.mean())
### END SOLUTION
```

<!-- grade_id: cell-0c92ee5f59a59d5c, points: 0, grade: true, solution: false -->
```python
assert isinstance(temperature, pd.Series)
### BEGIN HIDDEN TESTS
assert temperature.max() == 39.4
assert temperature.min() == 3.92
assert (temperature.mean() > 18) and (temperature.mean() < 19)
### END HIDDEN TESTS
```

**b)**

Delete the column ``status`` from the DataFrame ``machine_df`` and assign the newly generated DataFrame (without the deleted column) to the variable ``machine_df``.

<!-- grade_id: cell-f2aff4def0369ec0, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
machine_df = machine_df.drop(["status"], axis = "columns")
### END SOLUTION
```

<!-- grade_id: cell-62af690e68e9f910, points: 0, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert "status" not in machine_df.columns
### END HIDDEN TESTS
```

### Aufgabe 4: Filtering rows

<img align="center" width=600 src="figures/data_frame_filter_row.png"/>

It is also interesting to only look at certain rows of a DataFrame.
Perhaps you only want to look at those students whose grade is better than or equal to 1.3
You can achieve this via filtering by describing the condition within the square bracket ``[]`` as follows:

```python
best_students = students[students["exam_grade"] >= 1.3]
best_students
```

To implement multiple conditions, you can implement the conditions inside parentheses and link them to each other using logical operators.

IMPORTANT: To concatenate conditions in pandas, use the **&** character instead of the **and** operator, and the **|** character (vertical bar) instead of the **or** operator.

```python
best_me_students = students[
    (students["exam_grade"] >= 1.3) & 
    (students["course_of_study"] == "ME")
]

best_me_students
```

```python
students[
    (students["exam_grade"] >= 1.3) | 
    (students["course_of_study"] == "MACH")
]
```

**a)**

Filter the lines of the DataFrame ``machine_df`` and assign the lines with ``machine == 2`` and ``speed > 35`` to the variable m2. Then output ``m2`` and check whether the solution is correct.

<!-- grade_id: cell-fb671dd713b0c5c4, points: 0, grade: false, solution: true -->
```python
m2 = None 

### BEGIN SOLUTION
m2 = machine_df[
    (machine_df["machine"] == 2) & 
    (machine_df["speed"] > 35)
]
m2
### END SOLUTION
```

<!-- grade_id: cell-04b22c65ab47fde7, points: 0, grade: true, solution: false -->
```python
assert m2.shape == (1,3)
### BEGIN HIDDEN TESTS
assert m2["machine"].values[0] == 2
assert m2["speed"].values[0] == 90
assert m2["temperature"].values[0] == 39.4
### END HIDDEN TESTS
```