## Programmierung I: Grundlagen & Allgemeines

### Aufgabe 1: Multiple Choice (3 Punkte)

In dieser Aufgabe muss nichts implementiert werden. Es werden allgemeine Fragen gestellt und jeweils mehrere Antwortmöglichkeiten präsentiert.
Aus diesen Antwortmöglichkeiten ist die richtige auszuwählen.

Zur Eingabe der Antwort muss zwingend die dazu bereitgestellte Code-Zelle entsprechend der Anweisung bearbeitet werden.
Da die Punktevergabe automatisiert erfolgt, achten Sie bitte sehr genau auf die Einhaltung entsprechender Voragaben.

**a)**

Welche der folgenden Aussagen zu Variablennamen in Python ist zutreffend (genau eine Antwort ist richtig):

1) In Variablennamen können die ASCII-Zeichen beliebig eingesetzt werden.
2) Variablennamen dürfen mit dem Zeichen `_` beginnen.
3) Variablennamen müssen immer mit einem Buchstaben beginnen.
4) In Variablennamen sind ausschließlich Buchstaben und Zahlen zulässig.

Bitte beantworten Sie diese Frage, indem sie in der folgenden Code-Zelle eine Variable mit Namen `answer_1a` erzeugen und ihr den numerischen Wert der gewählten Antwort zuweisen.

(Konkret: Sind Sie der Meinung, dass Aussage *1) In Variablennamen können die ASCII-Zeichen beliebig eingesetzt werden.* korrekt ist, so schreiben Sie bitte `answer_1a = 1`.)

<!-- grade_id: cell-82af01e7d9fbe68a, points: 0, grade: false, solution: true -->
```python
answer_1a = None
### BEGIN SOLUTION
answer_1a = 2
### END SOLUTION
```

<!-- grade_id: cell-b8c62f5fc0df8d61, points: 1, grade: true, solution: false -->
```python
assert answer_1a
assert answer_1a in [1, 2, 3, 4]
### BEGIN HIDDEN TESTS
assert answer_1a == 2
### END HIDDEN TESTS
```

**b)**

Welche ist die korrekte Dateiendung für ein Python-Modul (genau eine Antwort ist richtig):

1) `.py`
2) `.ipynb`
3) `.ipy`
4) `.python`

Bitte beantworten Sie diese Frage, indem sie in der folgenden Code-Zelle eine Variable mit Namen `answer_1b` erzeugen und ihr den numerischen Wert der gewählten Antwort zuweisen.

<!-- grade_id: cell-520e95d7cd94dc95, points: 0, grade: false, solution: true -->
```python
answer_1b = None
### BEGIN SOLUTION
answer_1b = 1
### END SOLUTION
```

<!-- grade_id: cell-26610294288486e9, points: 1, grade: true, solution: false -->
```python
assert answer_1b
assert answer_1b in [1, 2, 3, 4]
### BEGIN HIDDEN TESTS
assert answer_1b == 1
### END HIDDEN TESTS
```

**c)**

Wie ist die Programmiersprache Python einzuordnen? (Genau eine Antwort ist richtig)

1) Interpretierte Sprache, stark typisiert.
2) Interpretierte Sprache, dynamisch typisiert.
3) Kompilierte Sprache, stark typisiert.
4) Kompilierte Sprache, dynamisch typisiert.

Bitte beantworten Sie diese Frage, indem sie in der folgenden Code-Zelle eine Variable mit Namen `answer_1c` erzeugen und ihr den numerischen Wert der gewählten Antwort zuweisen.

<!-- grade_id: cell-594036f80265e49f, points: 0, grade: false, solution: true -->
```python
answer_1c = None
### BEGIN SOLUTION
answer_1c = 2
### END SOLUTION
```

<!-- grade_id: cell-ca6574137b1fe825, points: 1, grade: true, solution: false -->
```python
assert answer_1c
assert answer_1c in [1, 2, 3, 4]
### BEGIN HIDDEN TESTS
assert answer_1c == 2
### END HIDDEN TESTS
```

### Aufgabe 2: Variablen & Zuweisungen (2 Punkte)
 **a)**
 
Erzeugen Sie folgende Variablen:
* Variable `a` vom Typ `int`  	mit Wert `5`
* Variable `b` vom Typ `float`	mit Wert `2.5`
* Variable `c` vom Typ `bool` 	mit Wert `True`
* Variable `d` vom Typ `str`  	mit Wert `Hallo IDS`

<!-- grade_id: cell-b85218bec9869cc7, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
a = 5
b = 2.5
c = True
d = "Hallo IDS"
### END SOLUTION
```

<!-- grade_id: cell-2ad53ccf36211ae9, points: 1, grade: true, solution: false -->
```python
assert a
assert b
assert c
assert d
### BEGIN HIDDEN TESTS
assert a == 5
assert b == 2.5
assert c == True
assert d == "Hallo IDS"
### END HIDDEN TESTS
```

**b)**

Erzeugen Sie folgende Variablen:
* Variable `value_1` vom mit Wert `12.3`
* Variable `value_2` vom mit Wert `4.8`

Berechnen Sie:
* Die Summe beider Variablen und speichern Sie das Ergebnis in einer neuen Variablen `add`
* Die Differenz beider Variablen und speichern Sie das Ergebnis in einer neuen Variablen `diff`
* Das Produkt beider Variablen und speichern Sie das Ergebnis in einer neuen Variablen `prod`
* Den Quotient beider Variablen und speichern Sie das Ergebnis in einer neuen Variablen `quot`

<!-- grade_id: cell-ce001117fcb94861, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
value_1 = 12.3
value_2 = 4.8

add = value_1 + value_2
diff = value_2 - value_1
prod = value_1 * value_2
quot = value_1 / value_2
value_2 / value_1
### END SOLUTION
```

<!-- grade_id: cell-0c0581f4bd2b8580, points: 1, grade: true, solution: false -->
```python
assert add
assert diff
assert prod
assert quot
### BEGIN HIDDEN TESTS
assert abs(add - 17.1) < 1e-6
assert abs(diff - 7.5) < 1e-6 or abs(diff + 7.5) < 1e-6
assert abs(prod - 59.04) < 1e-6
assert abs(quot - 2.5625) < 1e-6 or abs(quot - 0.3902439) < 1e-6
### END HIDDEN TESTS
```

### Aufgabe 3: Gemischtes (4 Punkte)

Schreiben Sie in der folgenden Notebook-Zelle Python-Code der, für verschiedene Kombinationen von Seitenlängen `a` und `b` eines rechtwinkligen Dreiecks, die Länge der Hypotenuse `c` berechnet und in einem Dictionary speichert.

**a)** Dazu soll zunächst eine Funktion `pythagoras` implementiert werden, welche die Längen der Katheten eines rechtwinkligen Dreiecks `a` und `b` als Argumente entgegennimmt und die zugehörige Hypotenusen-Länge berechnet. Diese muss mit `return` zurückgegeben werden.

**b)** Die Möglichen Seitenlängen der beiden Katheten `a` und `b` des rechtwinkligen Dreiecks seien aus der Menge `{1, 2, 3, 4}`.
Für jede mögliche Kombination von Seitenlängen `a` und `b`, berechnen Sie mit Hilfe des Satzes des Pythagoras die Seitenlänge der zugehörigen Hypotenuse `c`.
Speichern Sie alle zugehörigen Ergebnisse in einem einzigen geschachteteln Dictionary `hypothenuses`.
Mit `hypothenuses[a][b]` soll anschließend die Länge der Hypotenuse eines Dreiecks mit Seiten `a` und `b` abrufbar sein.

Der letzte bereits vorhandene Code-Block gibt die Einträge von `hypothenuses` formattiert auf dem Bildschirm aus, so dass jede Zeile eines der möglichen Dreiecke enthält.
Dieser sollte ohne Anpassungen funktionieren und für die 10 möglichen Kombinationen von Seitenlängen aus `sides` die zugehörigen Dreiecksmaße ausgeben.

**Hinweise**: 
* Der Satz des Pythagoras zur Berechnung der Länge der Hypotenuse `c` lautet: $c = \sqrt{a^2 + b^2}$
* Unter der Annahme, dass `a` und `b` nicht weiter unterscheidbar sind, berechnen Sie jede mögliche Permutation des Dreiecks nur ein einziges Mal. (Ist a=1, b=2 bereits berechnet, so muss a=2, b=1 nicht erneut betrachtet werden. Somit müssen insgesamt 10 Werte berechnet werden.)

<!-- grade_id: cell-467bee93838670f0, points: 0, grade: false, solution: true -->
```python
from math import sqrt

def pythagoras(a, b):
# A: Implementieren Sie hier die Funktion pythagoras
### BEGIN SOLUTION
    c_squared = a**2 + b**2
    c = sqrt(c_squared)
    return c
### END SOLUTION

sides = range(1, 5, 1)

# B: Berechnen Sie hier die Hypothenusen-Längen wie in der Aufgabenstellung beschrieben
### BEGIN SOLUTION
hypothenuses = {}
for a in sides:
    hypothenuses[a] = {}
    for b in sides:
        if b <= a:
            hyp = pythagoras(a, b)
            hypothenuses[a][b] = hyp
### END SOLUTION

# Ausgabe der Ergebnisse
print(f'{"a":>10s}{"b":>10s}{"c":>10s}')
print("-"*30)

for a in hypothenuses:
    for b in hypothenuses[a]:
        print(f"{a:10d}{b:10d}{hypothenuses[a][b]:10.3f}")
```

<!-- grade_id: cell-9bd4aeb96fcbe509, points: 2, grade: true, solution: false -->
```python
assert pythagoras
### BEGIN HIDDEN TESTS
assert abs(pythagoras(1, 2) - 2.23606797749979) < 1e-6
assert abs(pythagoras(0, 0) - 0) < 1e-6
assert abs(pythagoras(12, 12) - 16.97056274847714) < 1e-6
### END HIDDEN TESTS
```

<!-- grade_id: cell-4df55c5fb16520a4, points: 2, grade: true, solution: false -->
```python
assert type(hypothenuses) == dict
for side_length in sides:
    assert side_length in hypothenuses.keys()
    assert type(hypothenuses[side_length]) == dict
### BEGIN HIDDEN TESTS
assert len(hypothenuses.keys()) == 4

for side_length in sides:
    for other_side_length in sides:
        assert other_side_length in hypothenuses[side_length].keys() or side_length in hypothenuses[other_side_length].keys()

number_of_entries = 0
for key, item in hypothenuses.items():
    number_of_entries += len(item)
assert number_of_entries >= 10

assert abs(hypothenuses[1][1] - 1.4142135623730951) < 1e-6
assert abs(hypothenuses[2][2] - 2.8284271247461903) < 1e-6
assert abs(hypothenuses[3][3] - 4.242640687119285) < 1e-6
assert abs(hypothenuses[4][4] - 5.656854249492381) < 1e-6
### END HIDDEN TESTS
```