# Revision and Refreshers

### Exercise 1: If-Statements

**a)**

Rewrite the nested if conditions so that only a single if condition is used.

```python
temperature = 23.99999
day = "Mon"
working_days = ["Mon", "Tue", "Wed", "Thu", "Fr"]

if temperature >= 24:
    if day in working_days:
        print("Heat Free Day!")
else:
    print("Work hard!")
```

<!-- grade_id: cell-291b2405ee25ba0e, points: 0, grade: false, solution: true -->
```python
temperature = 23.99999
day = "Mon"
working_days = ["Mon", "Tue", "Wed", "Thu", "Fr"]

### BEGIN SOLUTION
if temperature >= 24 and day in working_days:
    print("Heat Free Day!")
else:
    print("Work hard!")
### END SOLUTION
```

**b)**

Rewrite the following if condition so that no more `and` or `or` operators are used, but the code remains functionally identical. 
Instead, several if conditions should be used.


```python
working_days = ["Mon", "Tue", "Wed", "Thu", "Fr"]

is_public_holiday = False
on_vacation = False
day = "Wed"

if day in working_days and not (is_public_holiday or on_vacation):
    print("Work hard!")
else:
    print("Relax and enjoy life!")
```

<!-- grade_id: cell-09cfba54e54f5fd3, points: 0, grade: false, solution: true -->
```python
working_days = ["Mon", "Tue", "Wed", "Thu", "Fr"]
is_public_holiday = False
on_vacation = False
day = "Wed"

### BEGIN SOLUTION
if day in working_days:
    if not is_public_holiday:
        if not on_vacation:
            print("Work hard!")
        else:
            print("Relax and enjoy life!")
    else:
        print("Relax and enjoy life!")
else:
    print("Relax and enjoy life!")
### END SOLUTION
```

### Exercise 2: Type conversion

A robot is equipped with a distance sensor that supplies measured values as a string in the format `XX.XX m` (e.g. `12.34 m`).
The following Python code is intended to convert these measured values into floating point numbers (float) so that they can be used for calculations.
However, a crucial line is missing in the code. Add the line to the code block.

*Note:* The place where the code is missing is underlined as follows ________. You may use one or two lines.

```python
1 def convert_distance_to_float(distance_str):
2   _________________________
3   _________________________
4   return distance_in_meters
5
6 distance_str = "25.87 m"
7 distance_float = convert_distance_to_float(distance_str)
8 print(f"Measured distance: {distance_float} Meter")
```

<!-- grade_id: cell-b2a34fbf6e4f0c82, points: 0, grade: false, solution: true -->
```python
def convert_distance_to_float(distance_str):
    ### BEGIN SOLUTION
    distance_in_meters = float(distance_str.split(" ")[0])
    ### END SOLUTION
    return distance_in_meters

distance_str = "25.87 m"
distance_float = convert_distance_to_float(distance_str)
print(f"Measured distance: {distance_float} Meter")
```

<!-- grade_id: cell-d5cae997a444d3fb, points: 0, grade: true, solution: false -->
```python
assert distance_float == 25.87
assert convert_distance_to_float("100.00 m") == 100
assert convert_distance_to_float("8.12 m") == 8.12
```

### Exercise 3: Objects and inheritance

Given are three classes that inherit from each other and the start of an associated main program (see code cell below).
The following output is already generated by this program:

```python
Creating Jin's motorized vehicle.
Creating car.
Creating Aya's motorized vehicle.
Creating motorcycle.
Creating Tom's motorized vehicle.
Creating motorcycle.
```

Your task is to implement additional lines that generate exactly the additional output below. 
No new variables may be declared and the existing list vehicles should contain all created instances of a MotorizedVehicle after completion of the program.


```python
Jane drives 10km in Jin's car.
Kilometer counter increased to 10km.
Tom sells the vehicle to Raj.
Raj rides 1km on Raj's motorbike.
Kilometer counter increased to 3913km.
Kilometer counter increased to 3914km.
Creating Tom's motorized vehicle.
Creating car.
Tom drives 1000km in Tom's car.
Kilometer counter increased to 1000km.
```

<!-- grade_id: cell-ec1ce6e641563768, points: 0, grade: false, solution: true -->
```python
class MotorizedVehicle:
    def __init__(self, seats, kW, kilometers, owner):
        self.seats = seats
        self.kW = kW
        self.kilometers = kilometers
        self.owner = owner
        print(f"Creating {self.owner}'s motorized vehicle.")

    def sell(self, new_owner):
        print(f"{self.owner} sells the vehicle to {new_owner}.")
        self.owner = new_owner

    def is_driven(self, distance):
        self.kilometers += distance
        print(f"Kilometer counter increased to {self.kilometers}km.")


class Car(MotorizedVehicle):
    def __init__(self, kW, kilometers, owner):
        super().__init__(5, kW, kilometers, owner)
        print(f"Creating car.")

    def drive(self, distance, driver):
        print(f"{driver} drives {distance}km in {self.owner}'s car.")
        self.is_driven(distance)


class Motorbike(MotorizedVehicle):
    def __init__(self, kW, kilometers, owner):
        super().__init__(2, kW, kilometers, owner)
        print(f"Creating motorbike.")

    def ride(self, distance, rider):
        print(f"{rider} rides {distance}km on {self.owner}'s motorbike.")
        self.is_driven(distance)

vehicles = [Car(100, 0, "Jin"),
            Motorbike(40, 2743, "Aya"),
            Motorbike(50, 3912, "Tom")]

### BEGIN SOLUTION
vehicles[0].drive(10, "Jane")
vehicles[2].sell("Raj")
vehicles[2].ride(1, "Raj")
vehicles[2].is_driven(1)
vehicles.append(Car(150, 0, "Tom"))
vehicles[3].drive(1000, "Tom")
### END SOLUTION
```

<!-- grade_id: cell-b496652f0d6eccf6, points: 0, grade: true, solution: false -->
```python
assert len(vehicles) == 4
assert vehicles[0].owner == "Jin"
assert vehicles[1].owner == "Aya"
assert vehicles[2].owner == "Raj"
assert vehicles[3].owner == "Tom"
```