## Aufgabe 03: Objektorientiertes Programmieren II


Allgemeiner Hinweis: Verwenden Sie die `print()`-Funktion, um Variablen anzuzeigen und Ihren Code zu debuggen.

### Aufgabe 1: Vererbung

#### 1 A)

Im Folgenden betrachten wir die Klasse `Smart_Device`, die ein intelligentes Smart Home Gerät abstrahiert.

Die Klasse hat ein boolesches Attribut „on“, das angibt, ob das Gerät gerade ein- oder ausgeschaltet ist. 
Außerdem gibt es ein Attribut `max_energy_consumption_kwh`, das den maximalen Stromverbrauch des Geräts in Kilowatt angibt.

Was Sie tun müssen:
- Implementieren Sie die Methode `turn_on_off` in der Klasse, die den Wert des booleschen Attributs `on` von `False` auf `True` umkehrt und umgekehrt.
- Erstelle eine Instanz von `Smart_Device`, deren Name `my_device` ist, und `max_energy_consumption_kW` ist 30.
- Benutze dann die Methode `turn_on_off`, um es einzuschalten. Schließlich wird der Wert des Attributs `on` ausgedruckt.

<!-- grade_id: cell-fdaca7585491de91, points: 0, grade: false, solution: true -->
```python
class Smart_Device:
    
    def __init__(self, max_energy_consumption_kW):
        self.on = False
        self.max_energy_consumption_kW = max_energy_consumption_kW
        
    ### BEGIN SOLUTION
    def turn_on_off(self):
        self.on = not self.on
    ### END SOLUTION

my_device = None

### BEGIN SOLUTION
my_device = Smart_Device(30)
my_device.turn_on_off()
print(my_device.on)
### END SOLUTION
```

<!-- grade_id: cell-98338f82d9fe88d4, points: 1, grade: true, solution: false -->
```python
assert my_device
assert hasattr(Smart_Device, "turn_on_off")
assert callable(Smart_Device.turn_on_off)
assert my_device.on
print("TEST PASSED!")
```

#### 1 B)

In diesem Teil werden wir eine Klasse namens `Smart_Light` implementieren, die von `Smart_Device` erbt.

Was Sie tun müssen:
- Implementieren Sie die Klasse `Smart_Light`, die von der Klasse `Smart_Device` erbt
- Implementiere einen Konstruktor, der die gleichen Argumente wie die Superklasse hat, zusätzlich zu zwei neuen Attributen, nämlich `brightness` und `color`.
- Setzen Sie die Standard-Helligkeit auf 70 und die Standard-Farbe auf weiß.

Tipp: Verwenden Sie die Funktion `super()`, um die Variablen der Superklasse zu initialisieren.

<!-- grade_id: cell-9871031a0fc87b74, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Smart_Light(Smart_Device):
    def __init__(self, max_energy_consumption_kW, brightness=70, color="white"):
        super().__init__(max_energy_consumption_kW)
        self.brightness = brightness
        self.color = color
### END SOLUTION
```

<!-- grade_id: cell-56cd708de73fb5bf, points: 1, grade: true, solution: false -->
```python
import inspect 

assert Smart_Light
assert issubclass(Smart_Light, Smart_Device)

signature = inspect.signature(Smart_Light.__dict__["__init__"])
params = signature.parameters.values()

param_names = [param.name for param in params]
assert "self" in param_names
assert "max_energy_consumption_kW" in param_names
assert "brightness" in param_names
assert "color" in param_names

assert signature.parameters["brightness"].default == 70
assert signature.parameters["color"].default == "white"

print("TEST PASSED!")
```

#### 1 C)

Nun geht es weiter mit der Implementierung der Klasse `Smart_Light` aus Teil B. Kopiere deine Lösung in die nächste Zelle. 
Wir wollen Instanzmethoden für die Einstellung der Helligkeit und der Farbe implementieren.

Was du tun musst:
- Kopieren Sie Ihren Code aus Teil B in die Zelle darunter.
- Füge die folgenden neuen Methoden zur gleichen Klasse `Smart_Light` hinzu:
  - Methode `set_color`: Diese Methode setzt die Lichtfarbe, indem sie das Attribut `color` neu zuweist.
  - Methode `set_brightness`: Mit dieser Methode wird die Helligkeit durch Neuzuweisung des Attributs `brightness` eingestellt. Stellen Sie sicher, dass die Helligkeit zwischen 0 und 100 liegt.
- Erstelle schließlich eine Instanz meiner `Smart_Light` mit dem Namen `my_light`, wobei `max_energy_consumption_kW` auf 0.03, `brightness` auf 75 und die Farbe auf „yellow“ gesetzt wird

<!-- grade_id: cell-03ed1f55f8ad5db7, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Smart_Light(Smart_Device):
    def __init__(self, max_energy_consumption_kW, brightness=70, color="white"):
        super().__init__(max_energy_consumption_kW)
        self.brightness = brightness
        self.color = color

    def set_color(self, color):
        self.color = color

    def set_brightness(self, brightness):
        if brightness >= 0 and brightness <= 100:
            self.brightness = brightness

my_light = Smart_Light(0.03, 75, "yellow")
### END SOLUTION
```

<!-- grade_id: cell-78d15dd824a63cdd, points: 0, grade: true, solution: false -->
```python
assert my_light
assert isinstance(my_light, Smart_Device)

assert hasattr(my_light, "max_energy_consumption_kW")
assert hasattr(my_light, "brightness")
assert hasattr(my_light, "color")

assert hasattr(my_light, "set_color")
assert hasattr(my_light, "set_brightness")

assert my_light.max_energy_consumption_kW == 0.03
assert my_light.color == "yellow"
assert my_light.brightness == 75

import inspect 

signature = inspect.signature(Smart_Light.__dict__["set_color"])
params = signature.parameters.values()
param_names = [param.name for param in params]
assert "self" in param_names
assert "color" in param_names

signature = inspect.signature(Smart_Light.__dict__["set_brightness"])
params = signature.parameters.values()
param_names = [param.name for param in params]
assert "self" in param_names
assert "brightness" in param_names

print("TEST PASSED!")
```

### Aufgabe 2: Magische Methods

#### 2 A)
Gegeben ist die Klasse `Robot`, die einen Roboter abstrahiert. Die Klasse hat die Attribute `name`, `battery_capacity` und `task`.

Implementieren Sie die Methoden `__str__` und `__repr__` wie folgt:
- Die magische Methode ``__str__`` soll einen String zurückgeben als `my name is <name>, my task is <task> and i can work for <battery_capacity> hours`
- Die magische Methode ``__repr__`` sollte eine String zurückgeben, die aus `name`, `task` und `battery_capacity` besteht. (z.B., `"spot,cleaning,50"`)

Erstelle nun eine Instanz mit dem Namen `my_robot`, deren Name `Jack` ist und deren Aufgabe `teaching` ist, und die eine Batteriekapazität von `12` Stunden hat. Schließlich wandeln Sie diese Klasse in eine String um und speichern sie in einer neuen Variablen namens `str_jack`. 

<!-- grade_id: cell-299917d4b5e5f00c, points: 0, grade: false, solution: true -->
```python
class Robot:
    
    def __init__(self, name, battery_capacity=50, task="cleaning"):
        self.name = name
        self.battery_capacity = battery_capacity
        self.task = task
        
    def __str__(self):
        ### BEGIN SOLUTION
        return f"my name is {self.name}, my task is {self.task} and I can do work for {self.battery_capacity} hours"
        ### END SOLUTION

    def __repr__(self):
        ### BEGIN SOLUTION
        return ",".join([self.name, self.task, str(self.battery_capacity)])
        ### END SOLUTION

my_robot = None
str_jack = None
### BEGIN SOLUTION
my_robot = Robot("Jack", 12, "teaching")
str_jack = str(my_robot)
print(str_jack)
my_robot.__repr__()
### END SOLUTION
```

<!-- grade_id: cell-80c70901041cfed5, points: 1, grade: true, solution: false -->
```python
assert my_robot
assert str_jack

assert isinstance(my_robot, Robot)
assert str_jack == "my name is Jack, my task is teaching and I can do work for 12 hours"

assert (my_robot.__repr__() == "Jack,teaching,12") or (my_robot.__repr__() == "Jack, teaching, 12")
print("TEST PASSED!")
```

#### 2 B)
Gegeben ist die Klasse `Grade`, die die Note eines Schülers abstrahiert. Die möglichen Noten sind `A+`, `B+`, `C+`, `D`, wobei `A+` die beste Note und `D` die schlechteste Note ist. 

Wir wollen die Methoden `__lt__`, `__gt__` und `__eq__` implementieren, um die Noten miteinander zu vergleichen. Zum Beispiel, wenn wir `A+` und `B+` vergleichen:
- `A+` > `B+` muss `True` zurückgeben
- `A+` < `B+` muss `False` zurückgeben
- `A+` == `B+` muss `False` zurückgeben

Hinweise:
- Sie können im Konstruktor zusätzliche Attribute initialisieren. Die Argumente müssen jedoch dieselben bleiben.
- Vielleicht können Ihnen Wörterbücher helfen?

<!-- grade_id: cell-7159c552b34bd068, points: 0, grade: false, solution: true -->
```python
class Grade():
    possible_grades = ["A+", "B+", "C+", "D"]
    def __init__(self, grade):
        if grade in self.possible_grades:
            self.grade = grade
        else:
            raise ValueError(f"grade must be on in {self.possible_grades}")
        ### BEGIN SOLUTION
        self.numeric_grade_dict = {
            "A+": 4,
            "B+": 3,
            "C+": 2,
            "D": 1
        }
        ### END SOLUTION

    ### BEGIN SOLUTION
    def __eq__(self, other):
        return self.numeric_grade_dict[self.grade] == self.numeric_grade_dict[other.grade]

    def __lt__(self, other):
        return self.numeric_grade_dict[self.grade] < self.numeric_grade_dict[other.grade]

    def __gt__(self, other):
        return self.numeric_grade_dict[self.grade] > self.numeric_grade_dict[other.grade]
    ### END SOLUTION
```

<!-- grade_id: cell-7a61c7e5047552d3, points: 0, grade: true, solution: false -->
```python
assert Grade

assert Grade("A+") == Grade("A+")
assert Grade("B+") == Grade("B+")
assert Grade("C+") == Grade("C+")
assert Grade("D") == Grade("D")

assert Grade("A+") != Grade("B+")
assert Grade("A+") != Grade("C+")
assert Grade("A+") != Grade("D")
assert Grade("B+") != Grade("C+")
assert Grade("B+") != Grade("D")
assert Grade("C+") != Grade("D")

assert Grade("A+") > Grade("B+")
assert Grade("A+") > Grade("C+")
assert Grade("A+") > Grade("D")
assert Grade("B+") > Grade("C+")
assert Grade("B+") > Grade("D")
assert Grade("C+") > Grade("D")

assert not Grade("A+") < Grade("B+") 
assert not Grade("A+") < Grade("C+")
assert not Grade("A+") < Grade("D")
assert not Grade("B+") < Grade("C+")
assert not Grade("B+") < Grade("D")
assert not Grade("C+") < Grade("D")

print("TEST PASSED!")
```